% SZS output start CNFRefutation for abs-problem-17-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf2, axiom,
    (X < 0 | ~ lgen(R, Y, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12)) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf3, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf3, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120), Y) | lgen(R, exp(X), Y))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(abs_problem_17, conjecture,
    (! [X] :
       ((1 < X | X < 0) =>
        abs(exp(X) * (12 - 6 * X + X ^ 2) - (12 + 6 * X + X ^ 2)) <=
        abs(X) ^ 5 * exp(abs(X)) / 60))).

fof(subgoal_0, plain,
    (! [X] :
       ((1 < X | X < 0) =>
        abs(exp(X) * (12 - 6 * X + X ^ 2) - (12 + 6 * X + X ^ 2)) <=
        abs(X) ^ 5 * exp(abs(X)) / 60)),
    inference(strip, [], [abs_problem_17])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((1 < X | X < 0) =>
          abs(exp(X) * (12 - 6 * X + X ^ 2) - (12 + 6 * X + X ^ 2)) <=
          abs(X) ^ 5 * exp(abs(X)) / 60)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (abs(X) ^ 5 * exp(abs(X)) / 60 <
        abs(exp(X) * (12 - 6 * X + X ^ 2) - (12 + 6 * X + X ^ 2)) &
        (1 < X | X < 0))), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (abs(skoXC1) ^ 5 * exp(abs(skoXC1)) / 60 <
     abs(exp(skoXC1) * (12 - 6 * skoXC1 + skoXC1 ^ 2) -
         (12 + 6 * skoXC1 + skoXC1 ^ 2)) & (skoXC1 < 0 | 1 < skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (abs(skoXC1) ^ 5 * exp(abs(skoXC1)) / 60 <
     abs(exp(skoXC1) * (12 - 6 * skoXC1 + skoXC1 ^ 2) -
         (12 + 6 * skoXC1 + skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 < 0 | 1 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -1440 + skoXC1 * (skoXC1 * (96 + skoXC1 * (skoXC1 * (-6 + skoXC1)))) |
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     0 |
     (-1440 +
      skoXC1 * (skoXC1 * (96 + skoXC1 * (skoXC1 * (-6 + skoXC1))))) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <=
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, exp(X_000098), X_000099) | exp(X_000098) <= X_000099),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (4 < X_000098 | X_000098 < 0 |
     ~ lgen(0,
            -(X_000098 ^ 3 + 12 * X_000098 ^ 2 + 60 * X_000098 + 120) /
            (X_000098 ^ 3 - 12 * X_000098 ^ 2 + 60 * X_000098 - 120),
            X_000099) | lgen(0, exp(X_000098), X_000099)),
    inference(subst, [], [exp_upper_bound_cf3])).

cnf(refute_0_4, plain,
    (4 < X_000098 | X_000098 < 0 |
     ~ lgen(0,
            -(X_000098 ^ 3 + 12 * X_000098 ^ 2 + 60 * X_000098 + 120) /
            (X_000098 ^ 3 - 12 * X_000098 ^ 2 + 60 * X_000098 - 120),
            X_000099) | exp(X_000098) <= X_000099),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (4 < X_000098 | X_000098 < 0 |
     X_000099 <
     (-120 + X_000098 * (-60 + X_000098 * (-12 + X_000098 * -1))) /
     (-120 + X_000098 * (60 + X_000098 * (-12 + X_000098))) |
     exp(X_000098) <= X_000099), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1 |
     exp(skoXC1) <=
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (abs(skoXC1) ^ 5 * exp(abs(skoXC1)) / 60 <
     abs(exp(skoXC1) * (12 - 6 * skoXC1 + skoXC1 ^ 2) -
         (12 + 6 * skoXC1 + skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(abs(skoXC1)) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60)))))),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_10, plain,
    (abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(abs(skoXC1)) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1),
    introduced(tautology, [equality])).

cnf(refute_0_11, plain,
    (skoXC1 < 0 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) |
     exp(abs(skoXC1)) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1),
    inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (skoXC1 < 0 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60)))))),
    inference(resolve, [], [refute_0_11, refute_0_8])).

cnf(refute_0_13, plain,
    (skoXC1 < 0 |
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)))),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) < 0 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) =
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_15, plain,
    (exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) !=
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_16, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) < 0 |
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60)))))),
    inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) < 0 | skoXC1 < 0 |
     exp(skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))),
    inference(resolve, [], [refute_0_16, refute_0_13])).

cnf(refute_0_18, plain,
    (skoXC1 < 0 |
     -12 + skoXC1 * (-6 + skoXC1 * -1) <
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60)))))),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (1440 + skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <=
     -12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_20, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) *
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_21, plain,
    (~ lgen(0, X_000085, exp(X_000084)) | X_000085 <= exp(X_000084)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_22, plain,
    (~ lgen(0, X_000085,
            -(X_000084 ^ 3 + 12 * X_000084 ^ 2 + 60 * X_000084 + 120) /
            (X_000084 ^ 3 - 12 * X_000084 ^ 2 + 60 * X_000084 - 120)) |
     0 <= X_000084 | lgen(0, X_000085, exp(X_000084))),
    inference(subst, [], [exp_lower_bound_cf3])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000085,
            -(X_000084 ^ 3 + 12 * X_000084 ^ 2 + 60 * X_000084 + 120) /
            (X_000084 ^ 3 - 12 * X_000084 ^ 2 + 60 * X_000084 - 120)) |
     X_000085 <= exp(X_000084) | 0 <= X_000084),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    ((-120 + X_000084 * (-60 + X_000084 * (-12 + X_000084 * -1))) /
     (-120 + X_000084 * (60 + X_000084 * (-12 + X_000084))) < X_000085 |
     X_000085 <= exp(X_000084) | 0 <= X_000084),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) | 0 <= skoXC1 |
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) <
     (1 + skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <=
     1 + skoXC1 * -1), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000050, exp(X_000049)) | X_000050 <= exp(X_000049)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (X_000049 < -1 | ~ lgen(0, X_000050, 1 + X_000049) |
     lgen(0, X_000050, exp(X_000049))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_29, plain,
    (X_000049 < -1 | ~ lgen(0, X_000050, 1 + X_000049) |
     X_000050 <= exp(X_000049)),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (1 + X_000049 < X_000050 | X_000049 < -1 | X_000050 <= exp(X_000049)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * -1 < -1 |
     1 + skoXC1 * -1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_33, plain,
    (abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(-skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) |
     abs(skoXC1) != -skoXC1 |
     exp(abs(skoXC1)) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1),
    introduced(tautology, [equality])).

cnf(refute_0_34, plain,
    (abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(-skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) |
     0 <= skoXC1 |
     exp(abs(skoXC1)) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain,
    (abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) * -1 <
     exp(-skoXC1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_34, refute_0_8])).

cnf(refute_0_36, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(skoXC1) != -skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_38, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * 1/60)))))),
    inference(resolve, [], [refute_0_32, refute_0_37])).

cnf(refute_0_39, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_38, refute_0_36])).

cnf(refute_0_40, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(skoXC1) != -skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_42, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_32, refute_0_41])).

cnf(refute_0_43, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (-skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_42, refute_0_40])).

cnf(refute_0_44, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (-144 + skoXC1 * (skoXC1 * (12 + skoXC1 * (skoXC1 * -1))) <
     (12 + skoXC1 * (6 + skoXC1)) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     (-144 + skoXC1 * (skoXC1 * (12 + skoXC1 * (skoXC1 * -1)))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= 12 + skoXC1 * (6 + skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_46, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) * (12 + skoXC1 * (-6 + skoXC1)) |
     12 + skoXC1 * (-6 + skoXC1) <= 0 |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_47, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) <
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)), introduced(tautology, [assume])).

cnf(refute_0_48, plain,
    (exp(skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(split, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (exp(skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain, (skoXC1 < 0 | 1 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_51, plain, (-12 < skoXC1 * (-6 + skoXC1)),
    inference(decision, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (exp(skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_49, refute_0_51])).

cnf(refute_0_53, plain,
    (~ lgen(0, X_000067, exp(X_000066)) | X_000067 <= exp(X_000066)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_54, plain,
    (X_000066 < 0 |
     ~ lgen(0, X_000067,
            (X_000066 ^ 2 + 6 * X_000066 + 12) /
            (X_000066 ^ 2 - 6 * X_000066 + 12)) |
     lgen(0, X_000067, exp(X_000066))),
    inference(subst, [], [exp_lower_bound_cf2])).

cnf(refute_0_55, plain,
    (X_000066 < 0 |
     ~ lgen(0, X_000067,
            (X_000066 ^ 2 + 6 * X_000066 + 12) /
            (X_000066 ^ 2 - 6 * X_000066 + 12)) |
     X_000067 <= exp(X_000066)),
    inference(resolve, [], [refute_0_54, refute_0_53])).

cnf(refute_0_56, plain,
    ((12 + X_000066 * (6 + X_000066)) / (12 + X_000066 * (-6 + X_000066)) <
     X_000067 | X_000066 < 0 | X_000067 <= exp(X_000066)),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_57, refute_0_52])).

cnf(refute_0_59, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) * (12 + skoXC1 * (-6 + skoXC1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | 12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_46, refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     (-144 + skoXC1 * (skoXC1 * (12 + skoXC1 * (skoXC1 * -1)))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) <
     (1 + skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= 1 + skoXC1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_62, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_63, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_62, refute_0_52])).

cnf(refute_0_64, plain,
    (skoXC1 < -1 |
     -12 + skoXC1 * (-6 + skoXC1 * -1) <
     (1 + skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_61, refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 < -1 | 0 < skoXC1 * (skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 < -1 | 0 < skoXC1 * (skoXC1 * (6 + skoXC1 * -1)) |
     -12 < skoXC1 * (-6 + skoXC1)),
    inference(decision, [], [refute_0_50])).

cnf(refute_0_67, plain,
    (skoXC1 < -1 | 0 < skoXC1 * (skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (0 < skoXC1 * (skoXC1 * (6 + skoXC1 * -1)) | -1 <= skoXC1),
    inference(decision, [], [refute_0_50])).

cnf(refute_0_69, plain,
    (0 < skoXC1 * (skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_68, refute_0_67])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 | -12 < skoXC1 * (-6 + skoXC1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(decision, [], [refute_0_69, refute_0_50])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     (-144 + skoXC1 * (skoXC1 * (12 + skoXC1 * (skoXC1 * -1)))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_60, refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 < 0 |
     -144 + skoXC1 * (skoXC1 * (12 + skoXC1 * (skoXC1 * -1))) <
     (12 + skoXC1 * (6 + skoXC1)) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_45, refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 < 0 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 < 0 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_73, refute_0_70])).

cnf(refute_0_75, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_44, refute_0_74])).

cnf(refute_0_76, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(skoXC1) != -skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * 1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_77, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * 1/60)))))),
    inference(resolve, [], [refute_0_32, refute_0_76])).

cnf(refute_0_78, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_77, refute_0_75])).

cnf(refute_0_79, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_79, refute_0_74])).

cnf(refute_0_81, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(skoXC1) != -skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_82, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_32, refute_0_81])).

cnf(refute_0_83, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_82, refute_0_80])).

cnf(refute_0_84, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_84, refute_0_74])).

cnf(refute_0_86, plain,
    (exp(skoXC1 * -1) *
     (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(skoXC1) != -skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (exp(skoXC1 * -1) *
     (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))))),
    inference(resolve, [], [refute_0_32, refute_0_86])).

cnf(refute_0_88, plain,
    (exp(skoXC1 * -1) *
     (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_87, refute_0_85])).

cnf(refute_0_89, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_89, refute_0_74])).

cnf(refute_0_91, plain,
    (0 <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) =
     -(-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_92, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -(-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) !=
     -(-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -(-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     0 <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -(-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) |
     0 <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_93, refute_0_90])).

cnf(refute_0_95, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) <
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (exp(skoXC1 * -1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(split, [], [refute_0_95])).

cnf(refute_0_97, plain,
    (exp(skoXC1 * -1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <= 0),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(decision, [], [refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_99, plain,
    (exp(skoXC1 * -1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 * -1 < -1 |
     1 + skoXC1 * -1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_31, refute_0_99])).

cnf(refute_0_101, plain,
    (1 + skoXC1 * -1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     1 < skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 <= 1),
    inference(decision, [], [refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_103, plain,
    (1 + skoXC1 * -1 <
     (-12 + skoXC1 * (-6 + skoXC1 * -1) +
      exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_102, refute_0_101])).

cnf(refute_0_104, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) <
     (1 + skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_26, refute_0_103])).

cnf(refute_0_105, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <= 0),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_105, refute_0_98])).

cnf(refute_0_107, plain,
    (exp(skoXC1) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(split, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (exp(skoXC1) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (-12 < skoXC1 * (-6 + skoXC1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(decision, [], [refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_110, plain,
    (exp(skoXC1) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) | 0 <= skoXC1 |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_25, refute_0_110])).

cnf(refute_0_112, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_111, refute_0_74])).

cnf(refute_0_113, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) *
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_20, refute_0_112])).

cnf(refute_0_114, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_114, refute_0_109])).

cnf(refute_0_116, plain,
    (1440 + skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_19, refute_0_115])).

cnf(refute_0_117, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-3 + skoXC1 * (6/5 + skoXC1 * (-13/60 + skoXC1 * 1/60)))))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     0 * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_119, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_120, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_121, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_120, refute_0_119])).

cnf(refute_0_122, plain,
    (0 <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (0 <
     (-12 +
      skoXC1 *
      (-6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60)))))) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_122, refute_0_110])).

cnf(refute_0_124, plain,
    (-12 +
     skoXC1 *
     (-6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (skoXC1 * (-1/60 + skoXC1 * 1/60))))) <
     0 * (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_118, refute_0_123])).

cnf(refute_0_125, plain,
    (-12 <
     skoXC1 *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (skoXC1 * (1/60 + skoXC1 * -1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (-12 < skoXC1 * (-6 + skoXC1) |
     -12 <
     skoXC1 *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (skoXC1 * (1/60 + skoXC1 * -1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(decision, [], [refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_127, plain,
    (-12 <
     skoXC1 *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (skoXC1 * (1/60 + skoXC1 * -1/60))))) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-3 + skoXC1 * (6/5 + skoXC1 * (-13/60 + skoXC1 * 1/60)))))))) <=
     0 | skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(decision, [],
              [refute_0_127, refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_129, plain,
    (exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(resolve, [], [refute_0_128, refute_0_117])).

cnf(refute_0_130, plain,
    (-120 < skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) |
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(decision, [],
              [refute_0_127, refute_0_50, refute_0_69, refute_0_74])).

cnf(refute_0_131, plain,
    (exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60)))))),
    inference(resolve, [], [refute_0_131, refute_0_18])).

cnf(refute_0_133, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_134, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60))))) |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) *
       (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_133])).

cnf(refute_0_135, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_134, refute_0_132])).

cnf(refute_0_136, plain,
    (~ lgen(0, exp(X_000073), X_000074) | exp(X_000073) <= X_000074),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_137, plain,
    (~ lgen(0,
            (X_000073 ^ 2 + 6 * X_000073 + 12) /
            (X_000073 ^ 2 - 6 * X_000073 + 12), X_000074) | 0 <= X_000073 |
     lgen(0, exp(X_000073), X_000074)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_138, plain,
    (~ lgen(0,
            (X_000073 ^ 2 + 6 * X_000073 + 12) /
            (X_000073 ^ 2 - 6 * X_000073 + 12), X_000074) | 0 <= X_000073 |
     exp(X_000073) <= X_000074),
    inference(resolve, [], [refute_0_137, refute_0_136])).

cnf(refute_0_139, plain,
    (X_000074 <
     (12 + X_000073 * (6 + X_000073)) / (12 + X_000073 * (-6 + X_000073)) |
     0 <= X_000073 | exp(X_000073) <= X_000074),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) |
     0 <= skoXC1 |
     exp(skoXC1) <=
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1))),
    inference(subst, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (12 + skoXC1 * (6 + skoXC1) +
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <
     0 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_142, plain,
    (0 <
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_121])).

cnf(refute_0_143, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_77, refute_0_44])).

cnf(refute_0_144, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_82, refute_0_144])).

cnf(refute_0_146, plain,
    (exp(skoXC1 * -1) *
     (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (exp(skoXC1 * -1) *
     (-skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_87, refute_0_146])).

cnf(refute_0_148, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain,
    (exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) !=
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    introduced(tautology, [equality])).

cnf(refute_0_150, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) < 0 |
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) |
     abs(-12 + skoXC1 * (-6 + skoXC1 * -1) +
         exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1))) <=
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_14, refute_0_149])).

cnf(refute_0_151, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) < 0 |
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -12 + skoXC1 * (-6 + skoXC1 * -1) +
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_150, refute_0_148])).

cnf(refute_0_152, plain,
    (12 + skoXC1 * (6 + skoXC1) +
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     -12 + skoXC1 * (-6 + skoXC1 * -1) <
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (12 + skoXC1 * (6 + skoXC1) +
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <
     exp(skoXC1 * -1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_131, refute_0_152])).

cnf(refute_0_154, plain,
    (exp(skoXC1 * -1) <
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     0 <= skoXC1), inference(split, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (exp(skoXC1 * -1) <
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <= 0),
    inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (-12 + skoXC1 * (-6 + skoXC1 * -1) <
     (-27648 + skoXC1 * (13824 + skoXC1 * -2304)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     -27648 + skoXC1 * (13824 + skoXC1 * -2304) <=
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_157, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) |
     0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_158, plain,
    (exp(X_000056) < X_000057 | ~ lgen(1, exp(X_000056), X_000057)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_159, plain,
    (0 < X_000056 |
     ~ lgen(1,
            2304 /
            (-(X_000056 ^ 3) + 6 * X_000056 ^ 2 - 24 * X_000056 + 48) ^ 2,
            X_000057) | lgen(1, exp(X_000056), X_000057)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_160, plain,
    (0 < X_000056 | exp(X_000056) < X_000057 |
     ~ lgen(1,
            2304 /
            (-(X_000056 ^ 3) + 6 * X_000056 ^ 2 - 24 * X_000056 + 48) ^ 2,
            X_000057)),
    inference(resolve, [], [refute_0_159, refute_0_158])).

cnf(refute_0_161, plain,
    (0 < X_000056 | exp(X_000056) < X_000057 |
     X_000057 <=
     2304 /
     (2304 +
      X_000056 *
      (-2304 +
       X_000056 *
       (1152 +
        X_000056 *
        (-384 + X_000056 * (84 + X_000056 * (-12 + X_000056))))))),
    inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (exp(skoXC1) <
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) | 0 < skoXC1 |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))))),
    inference(subst, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(split, [], [refute_0_131])).

cnf(refute_0_164, plain,
    (skoXC1 * (-6 + skoXC1) <= -12 |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    ((-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <= exp(skoXC1)),
    inference(resolve, [], [refute_0_164, refute_0_51])).

cnf(refute_0_166, plain,
    (0 < skoXC1 |
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) /
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))))),
    inference(resolve, [], [refute_0_165, refute_0_162])).

cnf(refute_0_167, plain,
    (0 < skoXC1 | 0 <= -12 + skoXC1 * (6 + skoXC1 * -1) |
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (-12 + skoXC1 * (6 + skoXC1 * -1)) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_166, refute_0_157])).

cnf(refute_0_168, plain,
    (0 < skoXC1 | skoXC1 * (-6 + skoXC1) <= -12 |
     (-27648 + skoXC1 * (13824 + skoXC1 * -2304)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_167])).

cnf(refute_0_169, plain, (-12 < skoXC1 * (-6 + skoXC1) | 0 < skoXC1),
    inference(decision, [], [refute_0_50])).

cnf(refute_0_170, plain,
    (0 < skoXC1 |
     (-27648 + skoXC1 * (13824 + skoXC1 * -2304)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <=
     -12 + skoXC1 * (-6 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (0 < skoXC1 |
     -27648 + skoXC1 * (13824 + skoXC1 * -2304) <=
     (-12 + skoXC1 * (-6 + skoXC1 * -1)) *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_170, refute_0_156])).

cnf(refute_0_172, plain,
    (0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-144 +
         skoXC1 * (-24 + skoXC1 * (24 + skoXC1 * (-6 + skoXC1))))))) <= 0 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_171])).

cnf(refute_0_173, plain,
    (-2304 <
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 + skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) |
     0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-144 +
         skoXC1 * (-24 + skoXC1 * (24 + skoXC1 * (-6 + skoXC1))))))) <= 0),
    inference(decision, [], [refute_0_50])).

cnf(refute_0_174, plain,
    (0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-144 +
         skoXC1 * (-24 + skoXC1 * (24 + skoXC1 * (-6 + skoXC1))))))) <= 0),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     0 <= skoXC1), inference(decision, [], [refute_0_50, refute_0_174])).

cnf(refute_0_176, plain,
    (exp(skoXC1 * -1) <
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_155, refute_0_175])).

cnf(refute_0_177, plain,
    (0 <
     (12 + skoXC1 * (6 + skoXC1) +
      exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1))) /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_142, refute_0_176])).

cnf(refute_0_178, plain,
    (12 + skoXC1 * (6 + skoXC1) +
     exp(skoXC1) * (-12 + skoXC1 * (6 + skoXC1 * -1)) <
     0 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60))))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_141, refute_0_177])).

cnf(refute_0_179, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <= 0),
    inference(arithmetic, [], [refute_0_178])).

cnf(refute_0_180, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) * (12 + skoXC1 * (-6 + skoXC1)) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_179, refute_0_175])).

cnf(refute_0_181, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     exp(skoXC1) | 0 <= skoXC1 | 12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(split, [], [refute_0_180])).

cnf(refute_0_182, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     exp(skoXC1) | 0 <= skoXC1 | skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain, (-12 < skoXC1 * (-6 + skoXC1) | 0 <= skoXC1),
    inference(decision, [], [refute_0_50, refute_0_174])).

cnf(refute_0_184, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     exp(skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) |
     0 <= skoXC1), inference(resolve, [], [refute_0_140, refute_0_184])).

cnf(refute_0_186, plain, (0 <= skoXC1),
    inference(arithmetic, [], [refute_0_185])).

cnf(refute_0_187, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_186, refute_0_135])).

cnf(refute_0_188, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60))))) |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_188])).

cnf(refute_0_190, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_189, refute_0_187])).

cnf(refute_0_191, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_186, refute_0_190])).

cnf(refute_0_192, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (abs(skoXC1) * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_192])).

cnf(refute_0_194, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_193, refute_0_191])).

cnf(refute_0_195, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_186, refute_0_194])).

cnf(refute_0_196, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_197, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) *
      (abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_196])).

cnf(refute_0_198, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_197, refute_0_195])).

cnf(refute_0_199, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_186, refute_0_198])).

cnf(refute_0_200, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     abs(skoXC1) != skoXC1 |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_201, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      abs(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_200])).

cnf(refute_0_202, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 + skoXC1 * (-6 + skoXC1) +
      skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_201, refute_0_199])).

cnf(refute_0_203, plain,
    (skoXC1 < 0 |
     12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     exp(skoXC1) *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_186, refute_0_203])).

cnf(refute_0_205, plain,
    ((12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     exp(skoXC1) |
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     0), inference(split, [], [refute_0_204])).

cnf(refute_0_206, plain,
    ((12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     exp(skoXC1) |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1 |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(resolve, [], [refute_0_6, refute_0_206])).

cnf(refute_0_208, plain,
    (12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <
     0 | 12 + skoXC1 * (6 + skoXC1) < 0 |
     0 <
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))),
    inference(split, [], [refute_0_204])).

cnf(refute_0_209, plain,
    (12 < skoXC1 * (-6 + skoXC1 * -1) |
     -12 <
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     12 <
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60))))),
    inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (-12 <
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     12 <
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) |
     skoXC1 * (-6 + skoXC1 * -1) <= 12),
    inference(decision, [], [refute_0_174, refute_0_50, refute_0_186])).

cnf(refute_0_211, plain,
    (-12 <
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     12 <
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60))))),
    inference(resolve, [], [refute_0_210, refute_0_209])).

cnf(refute_0_212, plain,
    (4 < skoXC1 | 0 <= skoXC1 |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12),
    inference(decision, [],
              [refute_0_211, refute_0_174, refute_0_50, refute_0_186])).

cnf(refute_0_213, plain,
    ((12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1 |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(resolve, [], [refute_0_212, refute_0_207])).

cnf(refute_0_214, plain,
    (skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12 | skoXC1 <= 4),
    inference(decision, [],
              [refute_0_211, refute_0_174, refute_0_50, refute_0_186])).

cnf(refute_0_215, plain,
    ((12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(resolve, [], [refute_0_214, refute_0_213])).

cnf(refute_0_216, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     0 *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     0 <=
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_217, plain,
    (0 <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_121])).

cnf(refute_0_218, plain,
    (exp(skoXC1) <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     0 <=
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))),
    inference(split, [], [refute_0_204])).

cnf(refute_0_219, plain,
    (exp(skoXC1) <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12), inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (0 <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12), inference(resolve, [], [refute_0_217, refute_0_219])).

cnf(refute_0_221, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     0 *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     0 <=
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) |
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12), inference(resolve, [], [refute_0_216, refute_0_220])).

cnf(refute_0_222, plain,
    (12 < skoXC1 * (-6 + skoXC1 * -1) |
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12), inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (skoXC1 * (-6 + skoXC1 * -1) <= 12 |
     skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12),
    inference(decision, [],
              [refute_0_211, refute_0_174, refute_0_50, refute_0_186])).

cnf(refute_0_224, plain,
    (skoXC1 * (6 + skoXC1 * (-1 + skoXC1 * (skoXC1 * (skoXC1 * 1/60)))) <=
     12), inference(resolve, [], [refute_0_223, refute_0_222])).

cnf(refute_0_225, plain,
    (-12 <
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))),
    inference(resolve, [], [refute_0_224, refute_0_211])).

cnf(refute_0_226, plain,
    ((12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_215, refute_0_225])).

cnf(refute_0_227, plain,
    (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <
     (12 + skoXC1 * (6 + skoXC1)) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))),
    inference(resolve, [], [refute_0_1, refute_0_226])).

cnf(refute_0_228, plain,
    (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <
     (-1440 +
      skoXC1 * (skoXC1 * (96 + skoXC1 * (skoXC1 * (-6 + skoXC1))))) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) |
     skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (-120 < skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1))),
    inference(decision, [],
              [refute_0_224, refute_0_211, refute_0_174, refute_0_50,
               refute_0_186])).

cnf(refute_0_230, plain,
    (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <
     (-1440 +
      skoXC1 * (skoXC1 * (96 + skoXC1 * (skoXC1 * (-6 + skoXC1))))) /
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))))),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) *
     (12 +
      skoXC1 *
      (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60))))) <
     -1440 + skoXC1 * (skoXC1 * (96 + skoXC1 * (skoXC1 * (-6 + skoXC1)))) |
     12 +
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_230])).

cnf(refute_0_232, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (-1 + skoXC1 * (-1/5 + skoXC1 * -1/60))))))) |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (-1 + skoXC1 * (-1/5 + skoXC1 * -1/60))))))) <=
     0 |
     skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12),
    inference(decision, [],
              [refute_0_224, refute_0_211, refute_0_174, refute_0_50,
               refute_0_186])).

cnf(refute_0_234, plain,
    (skoXC1 * (-6 + skoXC1 * (1 + skoXC1 * (skoXC1 * (skoXC1 * -1/60)))) <=
     -12), inference(resolve, [], [refute_0_233, refute_0_232])).

cnf(refute_0_235, plain, ($false),
    inference(resolve, [], [refute_0_234, refute_0_225])).
% SZS output end CNFRefutation for abs-problem-17-weak.tptp
