% SZS output start CNFRefutation for abs-problem-12.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(abs_problem_12, conjecture,
    (! [X] : exp(abs(X)) - 1 <= abs(X) * exp(abs(X)))).

fof(subgoal_0, plain, (! [X] : exp(abs(X)) - 1 <= abs(X) * exp(abs(X))),
    inference(strip, [], [abs_problem_12])).

fof(negate_0_0, plain, (~ ! [X] : exp(abs(X)) - 1 <= abs(X) * exp(abs(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain, (? [X] : abs(X) * exp(abs(X)) < exp(abs(X)) - 1),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (abs(skoXC1) * exp(abs(skoXC1)) < exp(abs(skoXC1)) - 1),
    inference(skolemize, [], [normalize_0_0])).

cnf(refute_0_0, plain,
    (1 *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 + skoXC1 * -21743271936 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0 |
     (21743271936 + skoXC1 * -21743271936) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <= 1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (1 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     1 / (1 + skoXC1 * -1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (abs(skoXC1) * exp(abs(skoXC1)) < exp(abs(skoXC1)) - 1),
    inference(canonicalize, [], [normalize_0_1])).

cnf(refute_0_3, plain, (1 < exp(abs(skoXC1)) * (1 + abs(skoXC1) * -1)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_5, plain,
    (1 < exp(abs(skoXC1)) * (1 + skoXC1 * -1) | abs(skoXC1) != skoXC1 |
     exp(abs(skoXC1)) * (1 + abs(skoXC1) * -1) <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_6, plain,
    (skoXC1 < 0 | 1 < exp(abs(skoXC1)) * (1 + skoXC1 * -1) |
     exp(abs(skoXC1)) * (1 + abs(skoXC1) * -1) <= 1),
    inference(resolve, [], [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain,
    (skoXC1 < 0 | 1 < exp(abs(skoXC1)) * (1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_6, refute_0_3])).

cnf(refute_0_8, plain,
    (skoXC1 < 0 | 1 / (1 + skoXC1 * -1) < exp(abs(skoXC1)) |
     1 + skoXC1 * -1 <= 0), inference(split, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (skoXC1 < 0 | 1 / (1 + skoXC1 * -1) < exp(abs(skoXC1)) | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (1 / (1 + skoXC1 * -1) < exp(skoXC1) | abs(skoXC1) != skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + skoXC1 * -1)),
    introduced(tautology, [equality])).

cnf(refute_0_11, plain,
    (skoXC1 < 0 | 1 / (1 + skoXC1 * -1) < exp(skoXC1) |
     exp(abs(skoXC1)) <= 1 / (1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_4, refute_0_10])).

cnf(refute_0_12, plain,
    (skoXC1 < 0 | 1 / (1 + skoXC1 * -1) < exp(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_11, refute_0_9])).

cnf(refute_0_13, plain,
    (~ lgen(0, exp(X_000063), X_000064) | exp(X_000063) <= X_000064),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_14, plain,
    (3 < X_000063 |
     ~ lgen(0,
            21743271936 /
            (-(X_000063 ^ 3) + 12 * X_000063 ^ 2 - 96 * X_000063 + 384) ^
            4, X_000064) | lgen(0, exp(X_000063), X_000064)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_15, plain,
    (3 < X_000063 |
     ~ lgen(0,
            21743271936 /
            (-(X_000063 ^ 3) + 12 * X_000063 ^ 2 - 96 * X_000063 + 384) ^
            4, X_000064) | exp(X_000063) <= X_000064),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (3 < X_000063 |
     X_000064 <
     21743271936 /
     (21743271936 +
      X_000063 *
      (-21743271936 +
       X_000063 *
       (10871635968 +
        X_000063 *
        (-3623878656 +
         X_000063 *
         (891813888 +
          X_000063 *
          (-169869312 +
           X_000063 *
           (25657344 +
            X_000063 *
            (-3096576 +
             X_000063 *
             (297216 +
              X_000063 *
              (-22272 +
               X_000063 * (1248 + X_000063 * (-48 + X_000063)))))))))))) |
     exp(X_000063) <= X_000064), inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (1 / (1 + skoXC1 * -1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | exp(skoXC1) <= 1 / (1 + skoXC1 * -1)),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 < 0 |
     1 / (1 + skoXC1 * -1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | 1 <= skoXC1),
    inference(resolve, [], [refute_0_17, refute_0_12])).

cnf(refute_0_19, plain,
    (1 < 0 | 1 + abs(skoXC1) * -1 < 0 | 0 < 1 + abs(skoXC1) * -1),
    inference(split, [], [refute_0_3])).

cnf(refute_0_20, plain, (abs(skoXC1) < 1 | 1 < abs(skoXC1)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1 < skoXC1 | abs(skoXC1) != skoXC1 | abs(skoXC1) <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_22, plain, (skoXC1 < 0 | 1 < skoXC1 | abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_4, refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 < 0 | abs(skoXC1) < 1 | 1 < skoXC1),
    inference(resolve, [], [refute_0_22, refute_0_20])).

cnf(refute_0_24, plain,
    (skoXC1 < 1 | abs(skoXC1) != skoXC1 | 1 <= abs(skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain, (skoXC1 < 0 | skoXC1 < 1 | 1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_4, refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 < 0 | skoXC1 < 1 | 1 < skoXC1),
    inference(resolve, [], [refute_0_25, refute_0_23])).

cnf(refute_0_27, plain, (skoXC1 < 1 | 1 < skoXC1 | 0 <= skoXC1),
    introduced(tautology, [decision])).

cnf(refute_0_28, plain, (skoXC1 < 1 | 1 < skoXC1),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (exp(abs(skoXC1)) < 1 / (1 + abs(skoXC1) * -1) |
     0 <= 1 + abs(skoXC1) * -1), inference(split, [], [refute_0_3])).

cnf(refute_0_30, plain,
    (exp(abs(skoXC1)) < 1 / (1 + abs(skoXC1) * -1) | abs(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_32, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_33, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (0 < 1 / (1 + abs(skoXC1) * -1) |
     1 / (1 + abs(skoXC1) * -1) <= exp(abs(skoXC1))),
    inference(subst, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (0 < 1 / (1 + abs(skoXC1) * -1) | abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_34, refute_0_30])).

cnf(refute_0_36, plain,
    (0 < 1 / (1 + skoXC1 * -1) | abs(skoXC1) != skoXC1 |
     1 / (1 + abs(skoXC1) * -1) <= 0), introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    (skoXC1 < 0 | 0 < 1 / (1 + skoXC1 * -1) |
     1 / (1 + abs(skoXC1) * -1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 < 0 | 0 < 1 / (1 + skoXC1 * -1) | abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_37, refute_0_35])).

cnf(refute_0_39, plain,
    (skoXC1 < 0 | abs(skoXC1) < 1 | 0 < 1 / (1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_38, refute_0_20])).

cnf(refute_0_40, plain,
    (0 < 1 / (1 + skoXC1 * -1) | 1 / (1 + skoXC1 * -1) <= exp(skoXC1)),
    inference(subst, [], [refute_0_33])).

cnf(refute_0_41, plain,
    (1 / (1 + abs(skoXC1) * -1) < exp(abs(skoXC1)) |
     1 + abs(skoXC1) * -1 <= 0), inference(split, [], [refute_0_3])).

cnf(refute_0_42, plain,
    (1 / (1 + abs(skoXC1) * -1) < exp(abs(skoXC1)) | 1 <= abs(skoXC1)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (1 / (1 + abs(skoXC1) * -1) < exp(skoXC1) | abs(skoXC1) != skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + abs(skoXC1) * -1)),
    introduced(tautology, [equality])).

cnf(refute_0_44, plain,
    (skoXC1 < 0 | 1 / (1 + abs(skoXC1) * -1) < exp(skoXC1) |
     exp(abs(skoXC1)) <= 1 / (1 + abs(skoXC1) * -1)),
    inference(resolve, [], [refute_0_4, refute_0_43])).

cnf(refute_0_45, plain,
    (skoXC1 < 0 | 1 / (1 + abs(skoXC1) * -1) < exp(skoXC1) |
     1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_44, refute_0_42])).

cnf(refute_0_46, plain,
    (1 < exp(skoXC1) * (1 + abs(skoXC1) * -1) | 1 + abs(skoXC1) * -1 <= 0 |
     exp(skoXC1) <= 1 / (1 + abs(skoXC1) * -1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_47, plain,
    (skoXC1 < 0 | 1 < exp(skoXC1) * (1 + abs(skoXC1) * -1) |
     1 <= abs(skoXC1) | 1 + abs(skoXC1) * -1 <= 0),
    inference(resolve, [], [refute_0_46, refute_0_45])).

cnf(refute_0_48, plain,
    (skoXC1 < 0 | 1 < exp(skoXC1) * (1 + abs(skoXC1) * -1) |
     1 <= abs(skoXC1)), inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (1 < exp(skoXC1) * (1 + skoXC1 * -1) | abs(skoXC1) != skoXC1 |
     exp(skoXC1) * (1 + abs(skoXC1) * -1) <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_50, plain,
    (skoXC1 < 0 | 1 < exp(skoXC1) * (1 + skoXC1 * -1) |
     exp(skoXC1) * (1 + abs(skoXC1) * -1) <= 1),
    inference(resolve, [], [refute_0_4, refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 < 0 | 1 < exp(skoXC1) * (1 + skoXC1 * -1) | 1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_50, refute_0_48])).

cnf(refute_0_52, plain,
    (skoXC1 < 0 | exp(skoXC1) < 1 / (1 + skoXC1 * -1) |
     0 <= 1 + skoXC1 * -1 | 1 <= abs(skoXC1)),
    inference(split, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (skoXC1 < 0 | exp(skoXC1) < 1 / (1 + skoXC1 * -1) | 1 <= abs(skoXC1) |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_55, plain,
    (1 < -skoXC1 | abs(skoXC1) != -skoXC1 | abs(skoXC1) <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_56, plain, (1 < -skoXC1 | 0 <= skoXC1 | abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain, (abs(skoXC1) < 1 | 1 < -skoXC1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_56, refute_0_20])).

cnf(refute_0_58, plain, (skoXC1 < -1 | abs(skoXC1) < 1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (-skoXC1 < 1 | abs(skoXC1) != -skoXC1 | 1 <= abs(skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_60, plain, (-skoXC1 < 1 | 0 <= skoXC1 | 1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_54, refute_0_59])).

cnf(refute_0_61, plain, (skoXC1 < -1 | -skoXC1 < 1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_60, refute_0_58])).

cnf(refute_0_62, plain, (skoXC1 < -1 | -1 < skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain, (skoXC1 < -1 | skoXC1 < 0 | -1 < skoXC1),
    inference(decision, [], [refute_0_28])).

cnf(refute_0_64, plain, (skoXC1 < -1 | -1 < skoXC1),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain, (0 <= skoXC1 | skoXC1 <= 1),
    inference(decision, [], [refute_0_64, refute_0_28])).

cnf(refute_0_66, plain,
    (exp(skoXC1) < 1 / (1 + skoXC1 * -1) | 1 <= abs(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_65, refute_0_53])).

cnf(refute_0_67, plain,
    (0 < 1 / (1 + skoXC1 * -1) | 1 <= abs(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_40, refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 < 0 | 0 < 1 / (1 + skoXC1 * -1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_67, refute_0_39])).

cnf(refute_0_69, plain, (0 < 1 / (1 + skoXC1 * -1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_65, refute_0_68])).

cnf(refute_0_70, plain,
    (1 < 0 * (1 + skoXC1 * -1) | 0 <= 1 + skoXC1 * -1 |
     1 / (1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_71, plain,
    (1 < 0 * (1 + skoXC1 * -1) | 0 <= 1 + skoXC1 * -1 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_70, refute_0_69])).

cnf(refute_0_72, plain, (skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain, (skoXC1 < 1),
    inference(resolve, [], [refute_0_72, refute_0_28])).

cnf(refute_0_74, plain,
    (skoXC1 < 0 |
     1 / (1 + skoXC1 * -1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1), inference(resolve, [], [refute_0_18, refute_0_73])).

cnf(refute_0_75, plain,
    (1 *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <
     2304 + skoXC1 * -2304 |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0 |
     (2304 + skoXC1 * -2304) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <= 1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_76, plain,
    (1 <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <=
     1 / (1 + skoXC1 * -1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_77, plain,
    (~ lgen(0, exp(X_000056), X_000057) | exp(X_000056) <= X_000057),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_78, plain,
    (0 < X_000056 |
     ~ lgen(0,
            2304 /
            (-(X_000056 ^ 3) + 6 * X_000056 ^ 2 - 24 * X_000056 + 48) ^ 2,
            X_000057) | lgen(0, exp(X_000056), X_000057)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_79, plain,
    (0 < X_000056 |
     ~ lgen(0,
            2304 /
            (-(X_000056 ^ 3) + 6 * X_000056 ^ 2 - 24 * X_000056 + 48) ^ 2,
            X_000057) | exp(X_000056) <= X_000057),
    inference(resolve, [], [refute_0_78, refute_0_77])).

cnf(refute_0_80, plain,
    (0 < X_000056 |
     X_000057 <
     2304 /
     (2304 +
      X_000056 *
      (-2304 +
       X_000056 *
       (1152 +
        X_000056 *
        (-384 + X_000056 * (84 + X_000056 * (-12 + X_000056)))))) |
     exp(X_000056) <= X_000057), inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (1 / (1 + skoXC1 * -1) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1 | exp(skoXC1) <= 1 / (1 + skoXC1 * -1)),
    inference(subst, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (skoXC1 < 0 |
     1 / (1 + skoXC1 * -1) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1 | 1 <= skoXC1),
    inference(resolve, [], [refute_0_81, refute_0_12])).

cnf(refute_0_83, plain,
    (skoXC1 < 0 |
     1 / (1 + skoXC1 * -1) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_82, refute_0_73])).

cnf(refute_0_84, plain,
    (skoXC1 < 0 |
     1 <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (1 + skoXC1 * -1) | 0 < skoXC1 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_76, refute_0_83])).

cnf(refute_0_85, plain,
    (skoXC1 < 0 |
     1 <
     (2304 + skoXC1 * -2304) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1 | 1 <= skoXC1), inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (0 < 1 / (1 + -skoXC1 * -1) | abs(skoXC1) != -skoXC1 |
     1 / (1 + abs(skoXC1) * -1) <= 0), introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (0 < 1 / (1 + -skoXC1 * -1) | 0 <= skoXC1 |
     1 / (1 + abs(skoXC1) * -1) <= 0),
    inference(resolve, [], [refute_0_54, refute_0_86])).

cnf(refute_0_88, plain,
    (0 < 1 / (1 + -skoXC1 * -1) | 0 <= skoXC1 | abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_87, refute_0_35])).

cnf(refute_0_89, plain,
    (0 < 1 / (1 + skoXC1) | 0 <= skoXC1 | abs(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (1 < abs(skoXC1) | abs(skoXC1) != -skoXC1 | -skoXC1 <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain, (1 < abs(skoXC1) | 0 <= skoXC1 | -skoXC1 <= 1),
    inference(resolve, [], [refute_0_54, refute_0_90])).

cnf(refute_0_92, plain,
    (0 < 1 / (1 + skoXC1) | 0 <= skoXC1 | -skoXC1 <= 1),
    inference(resolve, [], [refute_0_89, refute_0_91])).

cnf(refute_0_93, plain,
    (0 < 1 / (1 + skoXC1) | -1 <= skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain, (skoXC1 < 0 | -1 <= skoXC1),
    inference(decision, [], [refute_0_64, refute_0_28])).

cnf(refute_0_95, plain, (0 < 1 / (1 + skoXC1) | -1 <= skoXC1),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (1 < 0 * (1 + skoXC1) | 0 <= 1 + skoXC1 | 1 / (1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_97, plain,
    (1 < 0 * (1 + skoXC1) | -1 <= skoXC1 | 0 <= 1 + skoXC1),
    inference(resolve, [], [refute_0_96, refute_0_95])).

cnf(refute_0_98, plain, (-1 <= skoXC1),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain, (skoXC1 < 0 | skoXC1 < 1 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_98, refute_0_72, refute_0_64, refute_0_28])).

cnf(refute_0_100, plain,
    (skoXC1 < 0 |
     1 <
     (2304 + skoXC1 * -2304) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_85, refute_0_99])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 |
     1 *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <
     2304 + skoXC1 * -2304 | 0 < skoXC1 |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_75, refute_0_100])).

cnf(refute_0_102, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (-1152 +
       skoXC1 * (384 + skoXC1 * (-84 + skoXC1 * (12 + skoXC1 * -1))))) |
     0 < skoXC1 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-1152 +
       skoXC1 * (384 + skoXC1 * (-84 + skoXC1 * (12 + skoXC1 * -1))))) <=
     0 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304),
    inference(decision, [],
              [refute_0_98, refute_0_72, refute_0_64, refute_0_28])).

cnf(refute_0_104, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(resolve, [], [refute_0_103, refute_0_102])).

cnf(refute_0_105, plain,
    (skoXC1 < 0 |
     -2304 <
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 + skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_98, refute_0_72, refute_0_64, refute_0_28])).

cnf(refute_0_106, plain, (skoXC1 < 0 | 0 < skoXC1),
    inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain, (skoXC1 < 0 | skoXC1 <= 3),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_108, plain,
    (skoXC1 < 0 |
     1 / (1 + skoXC1 * -1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))))),
    inference(resolve, [], [refute_0_107, refute_0_74])).

cnf(refute_0_109, plain,
    (skoXC1 < 0 |
     1 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_1, refute_0_108])).

cnf(refute_0_110, plain,
    (skoXC1 < 0 |
     1 <
     (21743271936 + skoXC1 * -21743271936) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     1 <= skoXC1), inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain, (skoXC1 < 0 | skoXC1 < 1),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_112, plain,
    (skoXC1 < 0 |
     1 <
     (21743271936 + skoXC1 * -21743271936) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))))),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (skoXC1 < 0 |
     1 *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 + skoXC1 * -21743271936 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (-10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (-25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (-297216 +
             skoXC1 *
             (22272 +
              skoXC1 * (-1248 + skoXC1 * (48 + skoXC1 * -1))))))))))) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (1 *
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) <
     21743271936 + skoXC1 * 21743271936 |
     21743271936 +
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) <=
     0 |
     (21743271936 + skoXC1 * 21743271936) /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 +
               skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) <= 1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_116, plain,
    (1 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) *
     (1 + skoXC1) | 1 + skoXC1 <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 +
               skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) <=
     1 / (1 + skoXC1)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_117, plain,
    (1 / (1 + -skoXC1 * -1) < exp(abs(skoXC1)) | abs(skoXC1) != -skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + abs(skoXC1) * -1)),
    introduced(tautology, [equality])).

cnf(refute_0_118, plain,
    (1 / (1 + -skoXC1 * -1) < exp(abs(skoXC1)) | 0 <= skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + abs(skoXC1) * -1)),
    inference(resolve, [], [refute_0_54, refute_0_117])).

cnf(refute_0_119, plain,
    (1 / (1 + -skoXC1 * -1) < exp(abs(skoXC1)) | 0 <= skoXC1 |
     1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_118, refute_0_42])).

cnf(refute_0_120, plain,
    (1 / (1 + skoXC1) < exp(abs(skoXC1)) | 0 <= skoXC1 | 1 <= abs(skoXC1)),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (1 / (1 + skoXC1) < exp(-skoXC1) | abs(skoXC1) != -skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_122, plain,
    (1 / (1 + skoXC1) < exp(-skoXC1) | 0 <= skoXC1 |
     exp(abs(skoXC1)) <= 1 / (1 + skoXC1)),
    inference(resolve, [], [refute_0_54, refute_0_121])).

cnf(refute_0_123, plain,
    (1 / (1 + skoXC1) < exp(-skoXC1) | 0 <= skoXC1 | 1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_122, refute_0_120])).

cnf(refute_0_124, plain,
    (1 / (1 + skoXC1) < exp(skoXC1 * -1) | 0 <= skoXC1 | 1 <= abs(skoXC1)),
    inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (3 < skoXC1 * -1 |
     1 / (1 + skoXC1) <
     21743271936 /
     (21743271936 +
      skoXC1 * -1 *
      (-21743271936 +
       skoXC1 * -1 *
       (10871635968 +
        skoXC1 * -1 *
        (-3623878656 +
         skoXC1 * -1 *
         (891813888 +
          skoXC1 * -1 *
          (-169869312 +
           skoXC1 * -1 *
           (25657344 +
            skoXC1 * -1 *
            (-3096576 +
             skoXC1 * -1 *
             (297216 +
              skoXC1 * -1 *
              (-22272 +
               skoXC1 * -1 *
               (1248 + skoXC1 * -1 * (-48 + skoXC1 * -1)))))))))))) |
     exp(skoXC1 * -1) <= 1 / (1 + skoXC1)),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_126, plain,
    (3 < skoXC1 * -1 |
     1 / (1 + skoXC1) <
     21743271936 /
     (21743271936 +
      skoXC1 * -1 *
      (-21743271936 +
       skoXC1 * -1 *
       (10871635968 +
        skoXC1 * -1 *
        (-3623878656 +
         skoXC1 * -1 *
         (891813888 +
          skoXC1 * -1 *
          (-169869312 +
           skoXC1 * -1 *
           (25657344 +
            skoXC1 * -1 *
            (-3096576 +
             skoXC1 * -1 *
             (297216 +
              skoXC1 * -1 *
              (-22272 +
               skoXC1 * -1 *
               (1248 + skoXC1 * -1 * (-48 + skoXC1 * -1)))))))))))) |
     0 <= skoXC1 | 1 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_125, refute_0_124])).

cnf(refute_0_127, plain,
    (1 < 0 * (1 + abs(skoXC1) * -1) | 0 <= 1 + abs(skoXC1) * -1 |
     1 / (1 + abs(skoXC1) * -1) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_128, plain,
    (1 < 0 * (1 + abs(skoXC1) * -1) | 0 <= 1 + abs(skoXC1) * -1 |
     abs(skoXC1) <= 1),
    inference(resolve, [], [refute_0_127, refute_0_35])).

cnf(refute_0_129, plain, (abs(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain, (abs(skoXC1) < 1),
    inference(resolve, [], [refute_0_129, refute_0_20])).

cnf(refute_0_131, plain,
    (3 < skoXC1 * -1 |
     1 / (1 + skoXC1) <
     21743271936 /
     (21743271936 +
      skoXC1 * -1 *
      (-21743271936 +
       skoXC1 * -1 *
       (10871635968 +
        skoXC1 * -1 *
        (-3623878656 +
         skoXC1 * -1 *
         (891813888 +
          skoXC1 * -1 *
          (-169869312 +
           skoXC1 * -1 *
           (25657344 +
            skoXC1 * -1 *
            (-3096576 +
             skoXC1 * -1 *
             (297216 +
              skoXC1 * -1 *
              (-22272 +
               skoXC1 * -1 *
               (1248 + skoXC1 * -1 * (-48 + skoXC1 * -1)))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_126, refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 < -3 |
     1 / (1 + skoXC1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) |
     0 <= skoXC1), inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain, (-3 <= skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_134, plain,
    (1 / (1 + skoXC1) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_133, refute_0_132])).

cnf(refute_0_135, plain,
    (1 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) *
     (1 + skoXC1) | 0 <= skoXC1 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_116, refute_0_134])).

cnf(refute_0_136, plain,
    (1 <
     (21743271936 + skoXC1 * 21743271936) /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) |
     0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain, (-1 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_138, plain,
    (1 <
     (21743271936 + skoXC1 * 21743271936) /
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_136, refute_0_137])).

cnf(refute_0_139, plain,
    (1 *
     (21743271936 +
      skoXC1 *
      (21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1)))))))))))) <
     21743271936 + skoXC1 * 21743271936 | 0 <= skoXC1 |
     21743271936 +
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) <=
     0), inference(resolve, [], [refute_0_115, refute_0_138])).

cnf(refute_0_140, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (-891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (-25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (-297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (-1248 + skoXC1 * (-48 + skoXC1 * -1))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (-891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (-25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (-297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (-1248 + skoXC1 * (-48 + skoXC1 * -1))))))))))) <=
     0 |
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) <=
     -21743271936),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_142, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_141, refute_0_140])).

cnf(refute_0_143, plain,
    (-21743271936 <
     skoXC1 *
     (21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (22272 + skoXC1 * (1248 + skoXC1 * (48 + skoXC1))))))))))) |
     0 <= skoXC1),
    inference(decision, [],
              [refute_0_106, refute_0_98, refute_0_72, refute_0_64,
               refute_0_28])).

cnf(refute_0_144, plain, (0 <= skoXC1),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (-10871635968 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-891813888 +
         skoXC1 *
         (169869312 +
          skoXC1 *
          (-25657344 +
           skoXC1 *
           (3096576 +
            skoXC1 *
            (-297216 +
             skoXC1 *
             (22272 +
              skoXC1 * (-1248 + skoXC1 * (48 + skoXC1 * -1))))))))))) <=
     0 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936),
    inference(decision, [],
              [refute_0_144, refute_0_106, refute_0_98, refute_0_72,
               refute_0_64, refute_0_28])).

cnf(refute_0_146, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_145, refute_0_114])).

cnf(refute_0_147, plain,
    (skoXC1 < 0 |
     -21743271936 <
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 + skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))),
    inference(decision, [],
              [refute_0_144, refute_0_106, refute_0_98, refute_0_72,
               refute_0_64, refute_0_28])).

cnf(refute_0_148, plain, (skoXC1 < 0),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain, ($false),
    inference(resolve, [], [refute_0_144, refute_0_148])).
% SZS output end CNFRefutation for abs-problem-12.tptp
