% SZS output start CNFRefutation for abs-problem-5-sqrt-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(ln_upper_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (5 * X ^ 5 + 647 * X ^ 4 + 4397 * X ^ 3 + 6397 * X ^ 2 +
             2272 * X + 142) * (X - 1) /
            (6 * X ^ 5 + 75 * X ^ 4 + 200 * X ^ 3 + 150 * X ^ 2 + 30 * X +
             1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/70 *
            (353 + 1921 * X ^ 5 + 8389 * X + 38524 * X ^ 2 +
             20149 * X ^ 4 + 50774 * X ^ 3 + 10 * X ^ 6) * (X - 1) /
            (1 + 126 * X ^ 5 + 42 * X + 315 * X ^ 2 + 525 * X ^ 4 +
             700 * X ^ 3 + 7 * X ^ 6), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf15, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/280 *
            (1487 + 141123 * X ^ 5 + 49199 * X + 334575 * X ^ 2 +
             35 * X ^ 7 + 547235 * X ^ 4 + 718735 * X ^ 3 + 9411 * X ^ 6) *
            (X - 1) /
            (1 + 1176 * X ^ 5 + 56 * X + 588 * X ^ 2 + 8 * X ^ 7 +
             2450 * X ^ 4 + 1960 * X ^ 3 + 196 * X ^ 6), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_lower_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (142 * X ^ 5 + 2272 * X ^ 4 + 6397 * X ^ 3 + 4397 * X ^ 2 +
             647 * X + 5) * (X - 1) /
            (X *
             (X ^ 5 + 30 * X ^ 4 + 150 * X ^ 3 + 200 * X ^ 2 + 75 * X +
              6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/70 *
            (10 + 1921 * X + 20149 * X ^ 2 + 50774 * X ^ 3 +
             38524 * X ^ 4 + 8389 * X ^ 5 + 353 * X ^ 6) * (X - 1) /
            (X *
             (7 + 126 * X + 525 * X ^ 2 + 700 * X ^ 3 + 315 * X ^ 4 +
              42 * X ^ 5 + X ^ 6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf15, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/280 *
            (35 + 9411 * X + 141123 * X ^ 2 + 547235 * X ^ 3 +
             718735 * X ^ 4 + 334575 * X ^ 5 + 49199 * X ^ 6 +
             1487 * X ^ 7) * (X - 1) /
            (X *
             (8 + 196 * X + 1176 * X ^ 2 + 2450 * X ^ 3 + 1960 * X ^ 4 +
              588 * X ^ 5 + 56 * X ^ 6 + X ^ 7))) | lgen(R, Y, ln(X)))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_small, axiom,
    (X < 0 | 1 <= X |
     ~ lgen(R,
            1/32 *
            (65536 * X ^ 4 + 114688 * X ^ 3 + 17920 * X ^ 2 + 448 * X +
             1) / ((16 * X + 1) * (256 * X ^ 2 + 96 * X + 1)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_id, axiom,
    (X < 0 | 1 < X | ~ lgen(R, Y, X) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(abs_problem_5, conjecture,
    (! [X] :
       ((-999/1000 < X & X < 1200) =>
        abs(ln(1 + X)) <= 1/12 + abs(X) / sqrt(1 + X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((-999/1000 < X & X < 1200) =>
        abs(ln(1 + X)) <= 1/12 + abs(X) / sqrt(1 + X))),
    inference(strip, [], [abs_problem_5])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((-999/1000 < X & X < 1200) =>
          abs(ln(1 + X)) <= 1/12 + abs(X) / sqrt(1 + X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-999/1000 < X & X < 1200 &
        1/12 + abs(X) / sqrt(1 + X) < abs(ln(1 + X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 1200 &
     1/12 + abs(skoXC1) / sqrt(1 + skoXC1) < abs(ln(1 + skoXC1)) &
     -999/1000 < skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/12 + abs(skoXC1) / sqrt(1 + skoXC1) < abs(ln(1 + skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (-999/1000 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 1200),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((198593/384 +
      skoXC1 *
      (614065/96 +
       skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384)))) *
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <
     skoXC1 *
     (1277945955/32 +
      skoXC1 *
      (10218694629/64 +
       skoXC1 *
       (8480990661/32 +
        skoXC1 *
        (30303553841/128 +
         skoXC1 *
         (3948669967/32 +
          skoXC1 *
          (12195818119/320 +
           skoXC1 *
           (7621205769/1120 +
            skoXC1 *
            (5930325721/8960 +
             skoXC1 *
             (70606117/2240 +
              skoXC1 *
              (2618643/4480 +
               skoXC1 * (6369/2240 + skoXC1 * 1/256))))))))))) |
     6435 +
     skoXC1 *
     (24024 +
      skoXC1 *
      (36036 +
       skoXC1 *
       (27720 +
        skoXC1 *
        (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8)))))) <= 0 |
     skoXC1 *
     (1277945955/32 +
      skoXC1 *
      (10218694629/64 +
       skoXC1 *
       (8480990661/32 +
        skoXC1 *
        (30303553841/128 +
         skoXC1 *
         (3948669967/32 +
          skoXC1 *
          (12195818119/320 +
           skoXC1 *
           (7621205769/1120 +
            skoXC1 *
            (5930325721/8960 +
             skoXC1 *
             (70606117/2240 +
              skoXC1 *
              (2618643/4480 +
               skoXC1 * (6369/2240 + skoXC1 * 1/256))))))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <=
     198593/384 +
     skoXC1 *
     (614065/96 +
      skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (198593/384 +
     skoXC1 *
     (614065/96 +
      skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) *
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) |
     198593/32 +
     skoXC1 *
     (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32))) <=
     0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <=
     (198593/384 +
      skoXC1 *
      (614065/96 +
       skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384)))) /
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (~ lgen(0, ln(X_000113), X_000114) | ln(X_000113) <= X_000114),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (~ lgen(0,
            1/280 *
            (1487 + 141123 * X_000113 ^ 5 + 49199 * X_000113 +
             334575 * X_000113 ^ 2 + 35 * X_000113 ^ 7 +
             547235 * X_000113 ^ 4 + 718735 * X_000113 ^ 3 +
             9411 * X_000113 ^ 6) * (X_000113 - 1) /
            (1 + 1176 * X_000113 ^ 5 + 56 * X_000113 + 588 * X_000113 ^ 2 +
             8 * X_000113 ^ 7 + 2450 * X_000113 ^ 4 + 1960 * X_000113 ^ 3 +
             196 * X_000113 ^ 6), X_000114) | X_000113 <= 0 |
     lgen(0, ln(X_000113), X_000114)),
    inference(subst, [], [ln_upper_bound_cf15])).

cnf(refute_0_4, plain,
    (~ lgen(0,
            1/280 *
            (1487 + 141123 * X_000113 ^ 5 + 49199 * X_000113 +
             334575 * X_000113 ^ 2 + 35 * X_000113 ^ 7 +
             547235 * X_000113 ^ 4 + 718735 * X_000113 ^ 3 +
             9411 * X_000113 ^ 6) * (X_000113 - 1) /
            (1 + 1176 * X_000113 ^ 5 + 56 * X_000113 + 588 * X_000113 ^ 2 +
             8 * X_000113 ^ 7 + 2450 * X_000113 ^ 4 + 1960 * X_000113 ^ 3 +
             196 * X_000113 ^ 6), X_000114) | X_000113 <= 0 |
     ln(X_000113) <= X_000114),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000114 <
     (-1487/280 +
      X_000113 *
      (-852/5 +
       X_000113 *
       (-5096/5 +
        X_000113 *
        (-1372 +
         X_000113 *
         (1225/2 +
          X_000113 *
          (7252/5 +
           X_000113 *
           (2352/5 + X_000113 * (1172/35 + X_000113 * 1/8)))))))) /
     (1 +
      X_000113 *
      (56 +
       X_000113 *
       (588 +
        X_000113 *
        (1960 +
         X_000113 *
         (2450 + X_000113 * (1176 + X_000113 * (196 + X_000113 * 8))))))) |
     X_000113 <= 0 | ln(X_000113) <= X_000114),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((198593/384 +
      skoXC1 *
      (614065/96 +
       skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384)))) /
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) <
     (-1487/280 +
      (1 + skoXC1) *
      (-852/5 +
       (1 + skoXC1) *
       (-5096/5 +
        (1 + skoXC1) *
        (-1372 +
         (1 + skoXC1) *
         (1225/2 +
          (1 + skoXC1) *
          (7252/5 +
           (1 + skoXC1) *
           (2352/5 + (1 + skoXC1) * (1172/35 + (1 + skoXC1) * 1/8)))))))) /
     (1 +
      (1 + skoXC1) *
      (56 +
       (1 + skoXC1) *
       (588 +
        (1 + skoXC1) *
        (1960 +
         (1 + skoXC1) *
         (2450 +
          (1 + skoXC1) *
          (1176 + (1 + skoXC1) * (196 + (1 + skoXC1) * 8))))))) |
     1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <=
     (198593/384 +
      skoXC1 *
      (614065/96 +
       skoXC1 * (397283/192 + skoXC1 * (11153/96 + skoXC1 * 385/384)))) /
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    ((-198593/32 +
      skoXC1 *
      (-37969/8 +
       skoXC1 * (-9635/16 + skoXC1 * (-113/8 + skoXC1 * -1/32)))) *
     (-1/12 + ln(1 + skoXC1)) <
     skoXC1 * (-6001 + skoXC1 * (-2019 + skoXC1 * (-115 + skoXC1 * -1))) |
     -1/12 + ln(1 + skoXC1) <= 0 |
     skoXC1 * (-6001 + skoXC1 * (-2019 + skoXC1 * (-115 + skoXC1 * -1))) /
     (-1/12 + ln(1 + skoXC1)) <=
     -198593/32 +
     skoXC1 *
     (-37969/8 +
      skoXC1 * (-9635/16 + skoXC1 * (-113/8 + skoXC1 * -1/32)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_8, plain,
    (skoXC1 / (-1/12 + ln(1 + skoXC1)) *
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <
     198593/32 +
     skoXC1 *
     (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32))) |
     6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= 0 |
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <=
     skoXC1 / (-1/12 + ln(1 + skoXC1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_9, plain,
    (~ lgen(0, sqrt(X_000057), X_000058) | sqrt(X_000057) <= X_000058),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (X_000057 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000057 + 17920 * X_000057 ^ 2 +
             448 * X_000057 ^ 3 + X_000057 ^ 4) /
            ((X_000057 + 16) * (256 + 96 * X_000057 + X_000057 ^ 2)),
            X_000058) | lgen(0, sqrt(X_000057), X_000058)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_11, plain,
    (X_000057 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000057 + 17920 * X_000057 ^ 2 +
             448 * X_000057 ^ 3 + X_000057 ^ 4) /
            ((X_000057 + 16) * (256 + 96 * X_000057 + X_000057 ^ 2)),
            X_000058) | sqrt(X_000057) <= X_000058),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (X_000057 < 1 |
     X_000058 <
     (2048 +
      X_000057 *
      (3584 + X_000057 * (560 + X_000057 * (14 + X_000057 * 1/32)))) /
     (4096 + X_000057 * (1792 + X_000057 * (112 + X_000057))) |
     sqrt(X_000057) <= X_000058),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (1 + skoXC1 < 1 |
     skoXC1 / (-1/12 + ln(1 + skoXC1)) <
     (2048 +
      (1 + skoXC1) *
      (3584 +
       (1 + skoXC1) * (560 + (1 + skoXC1) * (14 + (1 + skoXC1) * 1/32)))) /
     (4096 + (1 + skoXC1) * (1792 + (1 + skoXC1) * (112 + (1 + skoXC1)))) |
     sqrt(1 + skoXC1) <= skoXC1 / (-1/12 + ln(1 + skoXC1))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (1/12 + abs(skoXC1) / sqrt(1 + skoXC1) < abs(ln(1 + skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_15, plain,
    (abs(skoXC1) / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1))),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_17, plain,
    (skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     abs(skoXC1) != skoXC1 |
     -1/12 + abs(ln(1 + skoXC1)) <= abs(skoXC1) / sqrt(1 + skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_18, plain,
    (skoXC1 < 0 | skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     -1/12 + abs(ln(1 + skoXC1)) <= abs(skoXC1) / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (skoXC1 < 0 | skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1))),
    inference(resolve, [], [refute_0_18, refute_0_15])).

cnf(refute_0_20, plain,
    (skoXC1 < (-1/12 + abs(ln(1 + skoXC1))) * sqrt(1 + skoXC1) |
     -1/12 + abs(ln(1 + skoXC1)) <= skoXC1 / sqrt(1 + skoXC1) |
     sqrt(1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_21, plain,
    (skoXC1 < 0 |
     skoXC1 < (-1/12 + abs(ln(1 + skoXC1))) * sqrt(1 + skoXC1) |
     sqrt(1 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (skoXC1 < 0 |
     skoXC1 + sqrt(1 + skoXC1) * 1/12 <
     abs(ln(1 + skoXC1)) * sqrt(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain, (-999/1000 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_24, plain, (skoXC1 < 1200),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_25, plain, (skoXC1 < 0 | -1 < skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 < 0 |
     skoXC1 + sqrt(1 + skoXC1) * 1/12 <
     abs(ln(1 + skoXC1)) * sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain,
    (ln(1 + skoXC1) < 0 | abs(ln(1 + skoXC1)) = ln(1 + skoXC1)),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_28, plain,
    (skoXC1 + sqrt(1 + skoXC1) * 1/12 < ln(1 + skoXC1) * sqrt(1 + skoXC1) |
     abs(ln(1 + skoXC1)) != ln(1 + skoXC1) |
     abs(ln(1 + skoXC1)) * sqrt(1 + skoXC1) <=
     skoXC1 + sqrt(1 + skoXC1) * 1/12), introduced(tautology, [equality])).

cnf(refute_0_29, plain,
    (ln(1 + skoXC1) < 0 |
     skoXC1 + sqrt(1 + skoXC1) * 1/12 < ln(1 + skoXC1) * sqrt(1 + skoXC1) |
     abs(ln(1 + skoXC1)) * sqrt(1 + skoXC1) <=
     skoXC1 + sqrt(1 + skoXC1) * 1/12),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (ln(1 + skoXC1) < 0 | skoXC1 < 0 |
     skoXC1 + sqrt(1 + skoXC1) * 1/12 < ln(1 + skoXC1) * sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_29, refute_0_26])).

cnf(refute_0_31, plain,
    (skoXC1 < 0 | skoXC1 < sqrt(1 + skoXC1) * (-1/12 + ln(1 + skoXC1)) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (skoXC1 < 0 | skoXC1 < sqrt(1 + skoXC1) * (-1/12 + ln(1 + skoXC1))),
    inference(resolve, [], [refute_0_31, refute_0_25])).

cnf(refute_0_33, plain,
    (skoXC1 < 0 | skoXC1 / (-1/12 + ln(1 + skoXC1)) < sqrt(1 + skoXC1) |
     -1/12 + ln(1 + skoXC1) <= 0), inference(split, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 < 0 | skoXC1 / (-1/12 + ln(1 + skoXC1)) < sqrt(1 + skoXC1) |
     ln(1 + skoXC1) <= 1/12), inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     1/12 |
     1/12 *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_36, plain,
    (ln(X_000078) < X_000079 | ~ lgen(1, ln(X_000078), X_000079)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_37, plain,
    (~ lgen(1,
            1/30 *
            (5 * X_000078 ^ 5 + 647 * X_000078 ^ 4 + 4397 * X_000078 ^ 3 +
             6397 * X_000078 ^ 2 + 2272 * X_000078 + 142) *
            (X_000078 - 1) /
            (6 * X_000078 ^ 5 + 75 * X_000078 ^ 4 + 200 * X_000078 ^ 3 +
             150 * X_000078 ^ 2 + 30 * X_000078 + 1), X_000079) |
     X_000078 <= 0 | lgen(1, ln(X_000078), X_000079)),
    inference(subst, [], [ln_upper_bound_cf11])).

cnf(refute_0_38, plain,
    (ln(X_000078) < X_000079 |
     ~ lgen(1,
            1/30 *
            (5 * X_000078 ^ 5 + 647 * X_000078 ^ 4 + 4397 * X_000078 ^ 3 +
             6397 * X_000078 ^ 2 + 2272 * X_000078 + 142) *
            (X_000078 - 1) /
            (6 * X_000078 ^ 5 + 75 * X_000078 ^ 4 + 200 * X_000078 ^ 3 +
             150 * X_000078 ^ 2 + 30 * X_000078 + 1), X_000079) |
     X_000078 <= 0), inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (ln(X_000078) < X_000079 | X_000078 <= 0 |
     X_000079 <=
     (-71/15 +
      X_000078 *
      (-71 +
       X_000078 *
       (-275/2 +
        X_000078 *
        (200/3 +
         X_000078 * (125 + X_000078 * (107/5 + X_000078 * 1/6)))))) /
     (1 +
      X_000078 *
      (30 +
       X_000078 *
       (150 + X_000078 * (200 + X_000078 * (75 + X_000078 * 6)))))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (ln(1 + skoXC1) < 1/12 |
     1/12 <=
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     1 + skoXC1 <= 0), inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    ((1/12 + skoXC1 * -11/12) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1743 +
       skoXC1 *
       (-2569 +
        skoXC1 *
        (-3703/2 +
         skoXC1 * (-6629/10 + skoXC1 * (-1562/15 + skoXC1 * -71/15)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0 |
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1743 +
       skoXC1 *
       (-2569 +
        skoXC1 *
        (-3703/2 +
         skoXC1 * (-6629/10 + skoXC1 * (-1562/15 + skoXC1 * -71/15)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     1/12 + skoXC1 * -11/12),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_42, plain,
    (1/12 + skoXC1 * -11/12 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-1 + skoXC1 * -1) | 0 <= -1 + skoXC1 * -1 |
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_43, plain,
    (~ lgen(0, X_000086, ln(X_000085)) | X_000086 <= ln(X_000085)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_44, plain,
    (~ lgen(0, X_000086,
            1/30 *
            (142 * X_000085 ^ 5 + 2272 * X_000085 ^ 4 +
             6397 * X_000085 ^ 3 + 4397 * X_000085 ^ 2 + 647 * X_000085 +
             5) * (X_000085 - 1) /
            (X_000085 *
             (X_000085 ^ 5 + 30 * X_000085 ^ 4 + 150 * X_000085 ^ 3 +
              200 * X_000085 ^ 2 + 75 * X_000085 + 6))) | X_000085 <= 0 |
     lgen(0, X_000086, ln(X_000085))),
    inference(subst, [], [ln_lower_bound_cf11])).

cnf(refute_0_45, plain,
    (~ lgen(0, X_000086,
            1/30 *
            (142 * X_000085 ^ 5 + 2272 * X_000085 ^ 4 +
             6397 * X_000085 ^ 3 + 4397 * X_000085 ^ 2 + 647 * X_000085 +
             5) * (X_000085 - 1) /
            (X_000085 *
             (X_000085 ^ 5 + 30 * X_000085 ^ 4 + 150 * X_000085 ^ 3 +
              200 * X_000085 ^ 2 + 75 * X_000085 + 6))) | X_000085 <= 0 |
     X_000086 <= ln(X_000085)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    ((-1/6 +
      X_000085 *
      (-107/5 +
       X_000085 *
       (-125 +
        X_000085 *
        (-200/3 +
         X_000085 * (275/2 + X_000085 * (71 + X_000085 * 71/15)))))) /
     (X_000085 *
      (6 +
       X_000085 *
       (75 +
        X_000085 *
        (200 + X_000085 * (150 + X_000085 * (30 + X_000085)))))) <
     X_000086 | X_000085 <= 0 | X_000086 <= ln(X_000085)),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) | 1 + skoXC1 <= 0 |
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) <= ln(1 + skoXC1)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (skoXC1 * -1 < (1 + skoXC1) * (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= -1/12 + ln(1 + skoXC1) * -1 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <= 1 + skoXC1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_49, plain,
    (~ lgen(0, X_000107, sqrt(X_000106)) | X_000107 <= sqrt(X_000106)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_50, plain,
    (1 < X_000106 | X_000106 < 0 | ~ lgen(0, X_000107, X_000106) |
     lgen(0, X_000107, sqrt(X_000106))),
    inference(subst, [], [sqrt_lower_bound_id])).

cnf(refute_0_51, plain,
    (1 < X_000106 | X_000106 < 0 | ~ lgen(0, X_000107, X_000106) |
     X_000107 <= sqrt(X_000106)),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (1 < X_000106 | X_000106 < X_000107 | X_000106 < 0 |
     X_000107 <= sqrt(X_000106)),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (1 + skoXC1 < 0 | 1 < 1 + skoXC1 |
     1 + skoXC1 < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <= sqrt(1 + skoXC1)),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_54, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_55, plain,
    (-skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     abs(skoXC1) != -skoXC1 |
     -1/12 + abs(ln(1 + skoXC1)) <= abs(skoXC1) / sqrt(1 + skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_56, plain,
    (-skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     0 <= skoXC1 |
     -1/12 + abs(ln(1 + skoXC1)) <= abs(skoXC1) / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (-skoXC1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     0 <= skoXC1), inference(resolve, [], [refute_0_56, refute_0_15])).

cnf(refute_0_58, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + abs(ln(1 + skoXC1)) |
     0 <= skoXC1), inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (0 <= ln(1 + skoXC1) | abs(ln(1 + skoXC1)) = -ln(1 + skoXC1)),
    inference(subst, [], [abs_negative])).

cnf(refute_0_60, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + -ln(1 + skoXC1) |
     abs(ln(1 + skoXC1)) != -ln(1 + skoXC1) |
     -1/12 + abs(ln(1 + skoXC1)) <= skoXC1 * -1 / sqrt(1 + skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + -ln(1 + skoXC1) |
     0 <= ln(1 + skoXC1) |
     -1/12 + abs(ln(1 + skoXC1)) <= skoXC1 * -1 / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + -ln(1 + skoXC1) |
     0 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_61, refute_0_58])).

cnf(refute_0_63, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + ln(1 + skoXC1) * -1 |
     0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain, (-1 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_65, plain,
    (skoXC1 * -1 / sqrt(1 + skoXC1) < -1/12 + ln(1 + skoXC1) * -1 |
     0 <= skoXC1), inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 * -1 < (-1/12 + ln(1 + skoXC1) * -1) * sqrt(1 + skoXC1) |
     -1/12 + ln(1 + skoXC1) * -1 <= skoXC1 * -1 / sqrt(1 + skoXC1) |
     sqrt(1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_67, plain,
    (skoXC1 * -1 < (-1/12 + ln(1 + skoXC1) * -1) * sqrt(1 + skoXC1) |
     0 <= skoXC1 | sqrt(1 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_66, refute_0_65])).

cnf(refute_0_68, plain,
    (skoXC1 * -1 < sqrt(1 + skoXC1) * (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (skoXC1 * -1 < sqrt(1 + skoXC1) * (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1), inference(resolve, [], [refute_0_68, refute_0_64])).

cnf(refute_0_70, plain,
    (sqrt(1 + skoXC1) < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= -1/12 + ln(1 + skoXC1) * -1 | 0 <= skoXC1),
    inference(split, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (sqrt(1 + skoXC1) < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1 | ln(1 + skoXC1) <= -1/12),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (-1/12 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <=
     -1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_73, plain,
    (X_000086 < ln(X_000085) | ~ lgen(1, X_000086, ln(X_000085))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_74, plain,
    (~ lgen(1, X_000086,
            1/30 *
            (142 * X_000085 ^ 5 + 2272 * X_000085 ^ 4 +
             6397 * X_000085 ^ 3 + 4397 * X_000085 ^ 2 + 647 * X_000085 +
             5) * (X_000085 - 1) /
            (X_000085 *
             (X_000085 ^ 5 + 30 * X_000085 ^ 4 + 150 * X_000085 ^ 3 +
              200 * X_000085 ^ 2 + 75 * X_000085 + 6))) | X_000085 <= 0 |
     lgen(1, X_000086, ln(X_000085))),
    inference(subst, [], [ln_lower_bound_cf11])).

cnf(refute_0_75, plain,
    (X_000086 < ln(X_000085) |
     ~ lgen(1, X_000086,
            1/30 *
            (142 * X_000085 ^ 5 + 2272 * X_000085 ^ 4 +
             6397 * X_000085 ^ 3 + 4397 * X_000085 ^ 2 + 647 * X_000085 +
             5) * (X_000085 - 1) /
            (X_000085 *
             (X_000085 ^ 5 + 30 * X_000085 ^ 4 + 150 * X_000085 ^ 3 +
              200 * X_000085 ^ 2 + 75 * X_000085 + 6))) | X_000085 <= 0),
    inference(resolve, [], [refute_0_74, refute_0_73])).

cnf(refute_0_76, plain,
    (X_000086 < ln(X_000085) | X_000085 <= 0 |
     (-1/6 +
      X_000085 *
      (-107/5 +
       X_000085 *
       (-125 +
        X_000085 *
        (-200/3 +
         X_000085 * (275/2 + X_000085 * (71 + X_000085 * 71/15)))))) /
     (X_000085 *
      (6 +
       X_000085 *
       (75 +
        X_000085 *
        (200 + X_000085 * (150 + X_000085 * (30 + X_000085)))))) <=
     X_000086), inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (-1/12 < ln(1 + skoXC1) | 1 + skoXC1 <= 0 |
     (-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <=
     -1/12), inference(subst, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (-1/12 < ln(1 + skoXC1) |
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     -1/12 | skoXC1 <= -1), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    ((1/6 + skoXC1 * (13/6 + skoXC1)) *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     skoXC1 *
     (924 +
      skoXC1 *
      (2982 +
       skoXC1 *
       (3626 +
        skoXC1 *
        (2037 + skoXC1 * (2569/5 + skoXC1 * (677/15 + skoXC1 * 1/3)))))) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0 |
     skoXC1 *
     (924 +
      skoXC1 *
      (2982 +
       skoXC1 *
       (3626 +
        skoXC1 *
        (2037 + skoXC1 * (2569/5 + skoXC1 * (677/15 + skoXC1 * 1/3)))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <=
     1/6 + skoXC1 * (13/6 + skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_80, plain,
    (1/6 + skoXC1 * (13/6 + skoXC1) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) *
     (2 + skoXC1 * 2) | 2 + skoXC1 * 2 <= 0 |
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <=
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_81, plain,
    (skoXC1 * (2 + skoXC1) <
     -1/6 + skoXC1 * -1/6 + ln(1 + skoXC1) * (2 + skoXC1 * 2) |
     2 + skoXC1 <= 0 |
     (-1/6 + skoXC1 * -1/6 + ln(1 + skoXC1) * (2 + skoXC1 * 2)) /
     (2 + skoXC1) <= skoXC1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_82, plain,
    (skoXC1 < (2 + skoXC1 * 2) / (2 + skoXC1) * (-1/12 + ln(1 + skoXC1)) |
     0 <= -1/12 + ln(1 + skoXC1) |
     skoXC1 / (-1/12 + ln(1 + skoXC1)) <= (2 + skoXC1 * 2) / (2 + skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_83, plain,
    (skoXC1 < 0 | sqrt(1 + skoXC1) < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     0 <= -1/12 + ln(1 + skoXC1)), inference(split, [], [refute_0_32])).

cnf(refute_0_84, plain,
    (skoXC1 < 0 | sqrt(1 + skoXC1) < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     1/12 <= ln(1 + skoXC1)), inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (~ lgen(0, X_000044, sqrt(X_000043)) | X_000044 <= sqrt(X_000043)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_86, plain,
    (X_000043 < 0 | ~ lgen(0, X_000044, 2 * X_000043 / (X_000043 + 1)) |
     lgen(0, X_000044, sqrt(X_000043))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_87, plain,
    (X_000043 < 0 | ~ lgen(0, X_000044, 2 * X_000043 / (X_000043 + 1)) |
     X_000044 <= sqrt(X_000043)),
    inference(resolve, [], [refute_0_86, refute_0_85])).

cnf(refute_0_88, plain,
    (X_000043 * 2 / (1 + X_000043) < X_000044 | X_000043 < 0 |
     X_000044 <= sqrt(X_000043)),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (1 + skoXC1 < 0 |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <
     skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     skoXC1 / (-1/12 + ln(1 + skoXC1)) <= sqrt(1 + skoXC1)),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (1 + skoXC1 < 0 | skoXC1 < 0 |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <
     skoXC1 / (-1/12 + ln(1 + skoXC1)) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_89, refute_0_84])).

cnf(refute_0_91, plain,
    (skoXC1 < -1 | skoXC1 < 0 |
     (2 + skoXC1 * 2) / (2 + skoXC1) < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     1/12 <= ln(1 + skoXC1)), inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain, (skoXC1 < 0 | -1 <= skoXC1),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_93, plain,
    (skoXC1 < 0 |
     (2 + skoXC1 * 2) / (2 + skoXC1) < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (skoXC1 < 0 |
     skoXC1 < (2 + skoXC1 * 2) / (2 + skoXC1) * (-1/12 + ln(1 + skoXC1)) |
     0 <= -1/12 + ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_82, refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 < 0 |
     skoXC1 <
     (-1/6 + skoXC1 * -1/6 + ln(1 + skoXC1) * (2 + skoXC1 * 2)) /
     (2 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1) <
     -1/6 + skoXC1 * -1/6 + ln(1 + skoXC1) * (2 + skoXC1 * 2) |
     1/12 <= ln(1 + skoXC1) | 2 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_81, refute_0_95])).

cnf(refute_0_97, plain,
    (skoXC1 < 0 |
     1/6 + skoXC1 * (13/6 + skoXC1) < ln(1 + skoXC1) * (2 + skoXC1 * 2) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain, (skoXC1 < 0 | -2 < skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_99, plain,
    (skoXC1 < 0 |
     1/6 + skoXC1 * (13/6 + skoXC1) < ln(1 + skoXC1) * (2 + skoXC1 * 2) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) < ln(1 + skoXC1) |
     1/12 <= ln(1 + skoXC1) | 2 + skoXC1 * 2 <= 0),
    inference(split, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) < ln(1 + skoXC1) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) < ln(1 + skoXC1) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_101, refute_0_25])).

cnf(refute_0_103, plain,
    (~ lgen(0, ln(X_000078), X_000079) | ln(X_000078) <= X_000079),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_104, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000078 ^ 5 + 647 * X_000078 ^ 4 + 4397 * X_000078 ^ 3 +
             6397 * X_000078 ^ 2 + 2272 * X_000078 + 142) *
            (X_000078 - 1) /
            (6 * X_000078 ^ 5 + 75 * X_000078 ^ 4 + 200 * X_000078 ^ 3 +
             150 * X_000078 ^ 2 + 30 * X_000078 + 1), X_000079) |
     X_000078 <= 0 | lgen(0, ln(X_000078), X_000079)),
    inference(subst, [], [ln_upper_bound_cf11])).

cnf(refute_0_105, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000078 ^ 5 + 647 * X_000078 ^ 4 + 4397 * X_000078 ^ 3 +
             6397 * X_000078 ^ 2 + 2272 * X_000078 + 142) *
            (X_000078 - 1) /
            (6 * X_000078 ^ 5 + 75 * X_000078 ^ 4 + 200 * X_000078 ^ 3 +
             150 * X_000078 ^ 2 + 30 * X_000078 + 1), X_000079) |
     X_000078 <= 0 | ln(X_000078) <= X_000079),
    inference(resolve, [], [refute_0_104, refute_0_103])).

cnf(refute_0_106, plain,
    (X_000079 <
     (-71/15 +
      X_000078 *
      (-71 +
       X_000078 *
       (-275/2 +
        X_000078 *
        (200/3 +
         X_000078 * (125 + X_000078 * (107/5 + X_000078 * 1/6)))))) /
     (1 +
      X_000078 *
      (30 +
       X_000078 *
       (150 + X_000078 * (200 + X_000078 * (75 + X_000078 * 6))))) |
     X_000078 <= 0 | ln(X_000078) <= X_000079),
    inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain,
    ((1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) <
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <=
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2)),
    inference(subst, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) <
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_107, refute_0_102])).

cnf(refute_0_109, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) / (2 + skoXC1 * 2) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_109, refute_0_25])).

cnf(refute_0_111, plain,
    (skoXC1 < 0 |
     1/6 + skoXC1 * (13/6 + skoXC1) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) *
     (2 + skoXC1 * 2) | 1/12 <= ln(1 + skoXC1) | 2 + skoXC1 * 2 <= 0),
    inference(resolve, [], [refute_0_80, refute_0_110])).

cnf(refute_0_112, plain,
    (skoXC1 < 0 |
     1/6 + skoXC1 * (13/6 + skoXC1) <
     skoXC1 *
     (924 +
      skoXC1 *
      (2982 +
       skoXC1 *
       (3626 +
        skoXC1 *
        (2037 + skoXC1 * (2569/5 + skoXC1 * (677/15 + skoXC1 * 1/3)))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (skoXC1 < 0 |
     1/6 + skoXC1 * (13/6 + skoXC1) <
     skoXC1 *
     (924 +
      skoXC1 *
      (2982 +
       skoXC1 *
       (3626 +
        skoXC1 *
        (2037 + skoXC1 * (2569/5 + skoXC1 * (677/15 + skoXC1 * 1/3)))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_112, refute_0_25])).

cnf(refute_0_114, plain,
    (skoXC1 < 0 |
     (1/6 + skoXC1 * (13/6 + skoXC1)) *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     skoXC1 *
     (924 +
      skoXC1 *
      (2982 +
       skoXC1 *
       (3626 +
        skoXC1 *
        (2037 + skoXC1 * (2569/5 + skoXC1 * (677/15 + skoXC1 * 1/3)))))) |
     1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0), inference(resolve, [], [refute_0_79, refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 < 0 |
     77 <
     skoXC1 *
     (-287 +
      skoXC1 *
      (-420 +
       skoXC1 *
       (-1372/3 +
        skoXC1 *
        (-2723/6 +
         skoXC1 * (-2747/10 + skoXC1 * (-1093/15 + skoXC1 * -17/3)))))) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     -462), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (-287 +
      skoXC1 *
      (-420 +
       skoXC1 *
       (-1372/3 +
        skoXC1 *
        (-2723/6 +
         skoXC1 * (-2747/10 + skoXC1 * (-1093/15 + skoXC1 * -17/3)))))) <=
     77 |
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     -462), inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_117, plain,
    (skoXC1 < 0 | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     -462), inference(resolve, [], [refute_0_116, refute_0_115])).

cnf(refute_0_118, plain,
    (skoXC1 < 0 |
     -462 <
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_119, plain, (skoXC1 < 0 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_117, refute_0_118])).

cnf(refute_0_120, plain, (-1 < skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [], [refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_121, plain,
    (-1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     -1/12), inference(resolve, [], [refute_0_78, refute_0_120])).

cnf(refute_0_122, plain,
    (-1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <=
     -1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_121, refute_0_72])).

cnf(refute_0_123, plain,
    (-1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (588 +
      skoXC1 *
      (2877/2 +
       skoXC1 *
       (4144/3 +
        skoXC1 * (2359/4 + skoXC1 * (512/5 + skoXC1 * 289/60))))) <=
     -77/2 |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (skoXC1 * -1 < 0 | -1/12 + ln(1 + skoXC1) * -1 < 0 |
     0 < -1/12 + ln(1 + skoXC1) * -1 | 0 <= skoXC1),
    inference(split, [], [refute_0_69])).

cnf(refute_0_125, plain,
    (ln(1 + skoXC1) < -1/12 | -1/12 < ln(1 + skoXC1) | 0 < skoXC1 |
     0 <= skoXC1), inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain, (0 <= skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_127, plain,
    (ln(1 + skoXC1) < -1/12 | -1/12 < ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_126, refute_0_125])).

cnf(refute_0_128, plain,
    ((198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     skoXC1 *
     (-85196397/8 +
      skoXC1 *
      (-1108314933/16 +
       skoXC1 *
       (-196825640 +
        skoXC1 *
        (-5122870785/16 +
         skoXC1 *
         (-26253998839/80 +
          skoXC1 *
          (-14058048697/64 +
           skoXC1 *
           (-215360364641/2240 +
            skoXC1 *
            (-935424479/35 +
             skoXC1 *
             (-153852952/35 +
              skoXC1 * (-12837632/35 + skoXC1 * -361472/35)))))))))) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0 |
     skoXC1 *
     (-85196397/8 +
      skoXC1 *
      (-1108314933/16 +
       skoXC1 *
       (-196825640 +
        skoXC1 *
        (-5122870785/16 +
         skoXC1 *
         (-26253998839/80 +
          skoXC1 *
          (-14058048697/64 +
           skoXC1 *
           (-215360364641/2240 +
            skoXC1 *
            (-935424479/35 +
             skoXC1 *
             (-153852952/35 +
              skoXC1 * (-12837632/35 + skoXC1 * -361472/35)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <=
     198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_129, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) *
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     0 <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) |
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) <=
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_130, plain,
    (~ lgen(0, X_000100, ln(X_000099)) | X_000100 <= ln(X_000099)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_131, plain,
    (~ lgen(0, X_000100,
            1/70 *
            (10 + 1921 * X_000099 + 20149 * X_000099 ^ 2 +
             50774 * X_000099 ^ 3 + 38524 * X_000099 ^ 4 +
             8389 * X_000099 ^ 5 + 353 * X_000099 ^ 6) * (X_000099 - 1) /
            (X_000099 *
             (7 + 126 * X_000099 + 525 * X_000099 ^ 2 +
              700 * X_000099 ^ 3 + 315 * X_000099 ^ 4 + 42 * X_000099 ^ 5 +
              X_000099 ^ 6))) | X_000099 <= 0 |
     lgen(0, X_000100, ln(X_000099))),
    inference(subst, [], [ln_lower_bound_cf13])).

cnf(refute_0_132, plain,
    (~ lgen(0, X_000100,
            1/70 *
            (10 + 1921 * X_000099 + 20149 * X_000099 ^ 2 +
             50774 * X_000099 ^ 3 + 38524 * X_000099 ^ 4 +
             8389 * X_000099 ^ 5 + 353 * X_000099 ^ 6) * (X_000099 - 1) /
            (X_000099 *
             (7 + 126 * X_000099 + 525 * X_000099 ^ 2 +
              700 * X_000099 ^ 3 + 315 * X_000099 ^ 4 + 42 * X_000099 ^ 5 +
              X_000099 ^ 6))) | X_000099 <= 0 | X_000100 <= ln(X_000099)),
    inference(resolve, [], [refute_0_131, refute_0_130])).

cnf(refute_0_133, plain,
    ((-1/7 +
      X_000099 *
      (-273/10 +
       X_000099 *
       (-1302/5 +
        X_000099 *
        (-875/2 +
         X_000099 *
         (175 +
          X_000099 *
          (861/2 + X_000099 * (574/5 + X_000099 * 353/70))))))) /
     (X_000099 *
      (7 +
       X_000099 *
       (126 +
        X_000099 *
        (525 +
         X_000099 *
         (700 + X_000099 * (315 + X_000099 * (42 + X_000099))))))) <
     X_000100 | X_000099 <= 0 | X_000100 <= ln(X_000099)),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    ((-1/7 +
      (1 + skoXC1) *
      (-273/10 +
       (1 + skoXC1) *
       (-1302/5 +
        (1 + skoXC1) *
        (-875/2 +
         (1 + skoXC1) *
         (175 +
          (1 + skoXC1) *
          (861/2 + (1 + skoXC1) * (574/5 + (1 + skoXC1) * 353/70))))))) /
     ((1 + skoXC1) *
      (7 +
       (1 + skoXC1) *
       (126 +
        (1 + skoXC1) *
        (525 +
         (1 + skoXC1) *
         (700 +
          (1 + skoXC1) * (315 + (1 + skoXC1) * (42 + (1 + skoXC1)))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     1 + skoXC1 <= 0 |
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) <=
     ln(1 + skoXC1)), inference(subst, [], [refute_0_133])).

cnf(refute_0_135, plain,
    ((-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) *
     (-1/12 + ln(1 + skoXC1) * -1) <
     skoXC1 *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) |
     -1/12 + ln(1 + skoXC1) * -1 <= 0 |
     skoXC1 *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) /
     (-1/12 + ln(1 + skoXC1) * -1) <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_136, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) <
     198593/32 +
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) |
     6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096)) <= 0 |
     (198593/32 +
      skoXC1 *
      (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048)))) /
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) <=
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_137, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) < sqrt(1 + skoXC1) |
     0 <= skoXC1 | -1/12 + ln(1 + skoXC1) * -1 <= 0),
    inference(split, [], [refute_0_69])).

cnf(refute_0_138, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) < sqrt(1 + skoXC1) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_137])).

cnf(refute_0_139, plain,
    (~ lgen(0, sqrt(X_000137), X_000138) | sqrt(X_000137) <= X_000138),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_140, plain,
    (X_000137 < 0 |
     ~ lgen(0,
            1/32 *
            (65536 * X_000137 ^ 4 + 114688 * X_000137 ^ 3 +
             17920 * X_000137 ^ 2 + 448 * X_000137 + 1) /
            ((16 * X_000137 + 1) *
             (256 * X_000137 ^ 2 + 96 * X_000137 + 1)), X_000138) |
     1 <= X_000137 | lgen(0, sqrt(X_000137), X_000138)),
    inference(subst, [], [sqrt_upper_bound_2_small])).

cnf(refute_0_141, plain,
    (X_000137 < 0 |
     ~ lgen(0,
            1/32 *
            (65536 * X_000137 ^ 4 + 114688 * X_000137 ^ 3 +
             17920 * X_000137 ^ 2 + 448 * X_000137 + 1) /
            ((16 * X_000137 + 1) *
             (256 * X_000137 ^ 2 + 96 * X_000137 + 1)), X_000138) |
     1 <= X_000137 | sqrt(X_000137) <= X_000138),
    inference(resolve, [], [refute_0_140, refute_0_139])).

cnf(refute_0_142, plain,
    (X_000137 < 0 |
     X_000138 <
     (1/32 +
      X_000137 *
      (14 + X_000137 * (560 + X_000137 * (3584 + X_000137 * 2048)))) /
     (1 + X_000137 * (112 + X_000137 * (1792 + X_000137 * 4096))) |
     1 <= X_000137 | sqrt(X_000137) <= X_000138),
    inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (1/32 +
      (1 + skoXC1) *
      (14 +
       (1 + skoXC1) *
       (560 + (1 + skoXC1) * (3584 + (1 + skoXC1) * 2048)))) /
     (1 +
      (1 + skoXC1) * (112 + (1 + skoXC1) * (1792 + (1 + skoXC1) * 4096))) |
     1 <= 1 + skoXC1 |
     sqrt(1 + skoXC1) <= skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (1/32 +
      (1 + skoXC1) *
      (14 +
       (1 + skoXC1) *
       (560 + (1 + skoXC1) * (3584 + (1 + skoXC1) * 2048)))) /
     (1 +
      (1 + skoXC1) * (112 + (1 + skoXC1) * (1792 + (1 + skoXC1) * 4096))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | 1 <= 1 + skoXC1),
    inference(resolve, [], [refute_0_143, refute_0_138])).

cnf(refute_0_145, plain,
    (skoXC1 < -1 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (198593/32 +
      skoXC1 *
      (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048)))) /
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain,
    ((4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     skoXC1 *
     (-27456 +
      skoXC1 *
      (-144672 +
       skoXC1 *
       (-326480 +
        skoXC1 *
        (-411960 +
         skoXC1 *
         (-3182897/10 +
          skoXC1 *
          (-619957/4 +
           skoXC1 *
           (-1651016/35 +
            skoXC1 *
            (-1202939/140 +
             skoXC1 *
             (-119617/140 +
              skoXC1 * (-21803/560 + skoXC1 * -353/560)))))))))) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0 |
     skoXC1 *
     (-27456 +
      skoXC1 *
      (-144672 +
       skoXC1 *
       (-326480 +
        skoXC1 *
        (-411960 +
         skoXC1 *
         (-3182897/10 +
          skoXC1 *
          (-619957/4 +
           skoXC1 *
           (-1651016/35 +
            skoXC1 *
            (-1202939/140 +
             skoXC1 *
             (-119617/140 +
              skoXC1 * (-21803/560 + skoXC1 * -353/560)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <=
     4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_147, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) *
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     0 <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) |
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) <=
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_148, plain,
    ((-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) *
     (-1/12 + ln(1 + skoXC1) * -1) <
     skoXC1 * (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -1/12 + ln(1 + skoXC1) * -1 <= 0 |
     skoXC1 * (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) /
     (-1/12 + ln(1 + skoXC1) * -1) <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_149, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0 |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_150, plain,
    (~ lgen(0, sqrt(X_000050), X_000051) | sqrt(X_000050) <= X_000051),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_151, plain,
    (X_000050 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000050 ^ 4 + 28 * X_000050 ^ 3 + 70 * X_000050 ^ 2 +
             28 * X_000050 + 1) /
            ((X_000050 + 1) * (X_000050 ^ 2 + 6 * X_000050 + 1)),
            X_000051) | lgen(0, sqrt(X_000050), X_000051)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_152, plain,
    (X_000050 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000050 ^ 4 + 28 * X_000050 ^ 3 + 70 * X_000050 ^ 2 +
             28 * X_000050 + 1) /
            ((X_000050 + 1) * (X_000050 ^ 2 + 6 * X_000050 + 1)),
            X_000051) | sqrt(X_000050) <= X_000051),
    inference(resolve, [], [refute_0_151, refute_0_150])).

cnf(refute_0_153, plain,
    (X_000050 < 0 |
     X_000051 <
     (1/8 +
      X_000050 *
      (7/2 + X_000050 * (35/4 + X_000050 * (7/2 + X_000050 * 1/8)))) /
     (1 + X_000050 * (7 + X_000050 * (7 + X_000050))) |
     sqrt(X_000050) <= X_000051),
    inference(arithmetic, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1)))) |
     sqrt(1 + skoXC1) <= skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_154, refute_0_138])).

cnf(refute_0_156, plain,
    (skoXC1 < -1 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain, (-1 <= skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_158, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_157, refute_0_156])).

cnf(refute_0_159, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_149, refute_0_158])).

cnf(refute_0_160, plain,
    (-16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) <
     skoXC1 * (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) /
     (-1/12 + ln(1 + skoXC1) * -1) | -1/12 <= ln(1 + skoXC1) |
     0 <= skoXC1 | skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1)) | 0 <= skoXC1),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_162, plain,
    (-16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) <
     skoXC1 * (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) /
     (-1/12 + ln(1 + skoXC1) * -1) | -1/12 <= ln(1 + skoXC1) |
     0 <= skoXC1), inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    ((-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) *
     (-1/12 + ln(1 + skoXC1) * -1) <
     skoXC1 * (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 |
     -1/12 + ln(1 + skoXC1) * -1 <= 0),
    inference(resolve, [], [refute_0_148, refute_0_162])).

cnf(refute_0_164, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     ln(1 + skoXC1) *
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (ln(1 + skoXC1) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) |
     0 <= skoXC1), inference(split, [], [refute_0_164])).

cnf(refute_0_166, plain,
    (ln(1 + skoXC1) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 |
     skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) <= -16),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (-16 < skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) |
     0 <= skoXC1), inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_168, plain,
    (ln(1 + skoXC1) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_166, refute_0_167])).

cnf(refute_0_169, plain,
    ((-1/7 +
      (1 + skoXC1) *
      (-273/10 +
       (1 + skoXC1) *
       (-1302/5 +
        (1 + skoXC1) *
        (-875/2 +
         (1 + skoXC1) *
         (175 +
          (1 + skoXC1) *
          (861/2 + (1 + skoXC1) * (574/5 + (1 + skoXC1) * 353/70))))))) /
     ((1 + skoXC1) *
      (7 +
       (1 + skoXC1) *
       (126 +
        (1 + skoXC1) *
        (525 +
         (1 + skoXC1) *
         (700 +
          (1 + skoXC1) * (315 + (1 + skoXC1) * (42 + (1 + skoXC1)))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     1 + skoXC1 <= 0 |
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) <=
     ln(1 + skoXC1)), inference(subst, [], [refute_0_133])).

cnf(refute_0_170, plain,
    ((-1/7 +
      (1 + skoXC1) *
      (-273/10 +
       (1 + skoXC1) *
       (-1302/5 +
        (1 + skoXC1) *
        (-875/2 +
         (1 + skoXC1) *
         (175 +
          (1 + skoXC1) *
          (861/2 + (1 + skoXC1) * (574/5 + (1 + skoXC1) * 353/70))))))) /
     ((1 + skoXC1) *
      (7 +
       (1 + skoXC1) *
       (126 +
        (1 + skoXC1) *
        (525 +
         (1 + skoXC1) *
         (700 +
          (1 + skoXC1) * (315 + (1 + skoXC1) * (42 + (1 + skoXC1)))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_169, refute_0_168])).

cnf(refute_0_171, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_170])).

cnf(refute_0_172, plain,
    ((4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-7392 +
      skoXC1 *
      (-35280 +
       skoXC1 *
       (-70840 +
        skoXC1 *
        (-77700 +
         skoXC1 *
         (-1011283/20 +
          skoXC1 *
          (-2380639/120 +
           skoXC1 *
           (-68317/15 +
            skoXC1 *
            (-135049/240 +
             skoXC1 * (-3763/120 + skoXC1 * -71/120))))))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0 |
     skoXC1 *
     (-7392 +
      skoXC1 *
      (-35280 +
       skoXC1 *
       (-70840 +
        skoXC1 *
        (-77700 +
         skoXC1 *
         (-1011283/20 +
          skoXC1 *
          (-2380639/120 +
           skoXC1 *
           (-68317/15 +
            skoXC1 *
            (-135049/240 +
             skoXC1 * (-3763/120 + skoXC1 * -71/120))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_173, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     0 <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) |
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_174, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     1 + skoXC1 <= 0 |
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) <=
     ln(1 + skoXC1)), inference(subst, [], [refute_0_46])).

cnf(refute_0_175, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_174, refute_0_168])).

cnf(refute_0_176, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 < 0 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(decision, [], [refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_178, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_176, refute_0_177])).

cnf(refute_0_179, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_178, refute_0_120])).

cnf(refute_0_180, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_173, refute_0_179])).

cnf(refute_0_181, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (-7392 +
      skoXC1 *
      (-35280 +
       skoXC1 *
       (-70840 +
        skoXC1 *
        (-77700 +
         skoXC1 *
         (-1011283/20 +
          skoXC1 *
          (-2380639/120 +
           skoXC1 *
           (-68317/15 +
            skoXC1 *
            (-135049/240 +
             skoXC1 * (-3763/120 + skoXC1 * -71/120))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) <= -16),
    inference(arithmetic, [], [refute_0_180])).

cnf(refute_0_182, plain,
    ((1/12 + skoXC1 * -23/24) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (-3857/2 +
        skoXC1 *
        (-2415/2 +
         skoXC1 * (-7623/20 + skoXC1 * (-1633/30 + skoXC1 * -71/30)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0 |
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (-3857/2 +
        skoXC1 *
        (-2415/2 +
         skoXC1 * (-7623/20 + skoXC1 * (-1633/30 + skoXC1 * -71/30)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     1/12 + skoXC1 * -23/24),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_183, plain,
    (1/12 + skoXC1 * -23/24 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-1 + skoXC1 * -1/2) | 0 <= -1 + skoXC1 * -1/2 |
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_184, plain,
    (skoXC1 * -1 < (1 + skoXC1 * 1/2) * (-1/12 + ln(1 + skoXC1) * -1) |
     -1/12 + ln(1 + skoXC1) * -1 <= 0 |
     1 + skoXC1 * 1/2 <= skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_185, plain,
    (~ lgen(0, sqrt(X_000036), X_000037) | sqrt(X_000036) <= X_000037),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_186, plain,
    (X_000036 < 0 | ~ lgen(0, (X_000036 + 1) / 2, X_000037) |
     lgen(0, sqrt(X_000036), X_000037)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_187, plain,
    (X_000036 < 0 | ~ lgen(0, (X_000036 + 1) / 2, X_000037) |
     sqrt(X_000036) <= X_000037),
    inference(resolve, [], [refute_0_186, refute_0_185])).

cnf(refute_0_188, plain,
    (X_000036 < 0 | X_000037 < 1/2 + X_000036 * 1/2 |
     sqrt(X_000036) <= X_000037),
    inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     1/2 + (1 + skoXC1) * 1/2 |
     sqrt(1 + skoXC1) <= skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1)),
    inference(subst, [], [refute_0_188])).

cnf(refute_0_190, plain,
    (1 + skoXC1 < 0 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     1/2 + (1 + skoXC1) * 1/2 | -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_189, refute_0_138])).

cnf(refute_0_191, plain,
    (skoXC1 < -1 |
     skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) < 1 + skoXC1 * 1/2 |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_190])).

cnf(refute_0_192, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) < 1 + skoXC1 * 1/2 |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_157, refute_0_191])).

cnf(refute_0_193, plain,
    (skoXC1 * -1 < (1 + skoXC1 * 1/2) * (-1/12 + ln(1 + skoXC1) * -1) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 |
     -1/12 + ln(1 + skoXC1) * -1 <= 0),
    inference(resolve, [], [refute_0_184, refute_0_192])).

cnf(refute_0_194, plain,
    (1/12 + skoXC1 * -23/24 < ln(1 + skoXC1) * (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= -1 + skoXC1 * -1/2 | 0 <= skoXC1),
    inference(split, [], [refute_0_194])).

cnf(refute_0_196, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_195])).

cnf(refute_0_197, plain, (-2 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_198, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_196, refute_0_197])).

cnf(refute_0_199, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) | 1 + skoXC1 <= 0 |
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) <= ln(1 + skoXC1)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_200, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_199, refute_0_198])).

cnf(refute_0_201, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_201, refute_0_177])).

cnf(refute_0_203, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (1/12 + skoXC1 * -23/24) / (-1 + skoXC1 * -1/2) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_202, refute_0_120])).

cnf(refute_0_204, plain,
    (1/12 + skoXC1 * -23/24 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-1 + skoXC1 * -1/2) | -1/12 <= ln(1 + skoXC1) |
     0 <= -1 + skoXC1 * -1/2 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_183, refute_0_203])).

cnf(refute_0_205, plain,
    (1/12 + skoXC1 * -23/24 <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (-3857/2 +
        skoXC1 *
        (-2415/2 +
         skoXC1 * (-7623/20 + skoXC1 * (-1633/30 + skoXC1 * -71/30)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <
     -1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_207, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     -1/12 | -1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_208, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     -1/12 | -1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     -1/12 | -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_208, refute_0_120])).

cnf(refute_0_210, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <
     -1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_206, refute_0_209])).

cnf(refute_0_211, plain,
    (77/2 <
     skoXC1 *
     (-588 +
      skoXC1 *
      (-2877/2 +
       skoXC1 *
       (-4144/3 +
        skoXC1 * (-2359/4 + skoXC1 * (-512/5 + skoXC1 * -289/60))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (77/2 <
     skoXC1 *
     (-588 +
      skoXC1 *
      (-2877/2 +
       skoXC1 *
       (-4144/3 +
        skoXC1 * (-2359/4 + skoXC1 * (-512/5 + skoXC1 * -289/60))))) |
     -462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(decision, [], [refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_213, plain,
    (77/2 <
     skoXC1 *
     (-588 +
      skoXC1 *
      (-2877/2 +
       skoXC1 *
       (-4144/3 +
        skoXC1 * (-2359/4 + skoXC1 * (-512/5 + skoXC1 * -289/60))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_211, refute_0_212])).

cnf(refute_0_214, plain,
    (-2 < skoXC1 | -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_213, refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_215, plain,
    (1/12 + skoXC1 * -23/24 <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (-3857/2 +
        skoXC1 *
        (-2415/2 +
         skoXC1 * (-7623/20 + skoXC1 * (-1633/30 + skoXC1 * -71/30)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_205, refute_0_214])).

cnf(refute_0_216, plain,
    ((1/12 + skoXC1 * -23/24) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (-3857/2 +
        skoXC1 *
        (-2415/2 +
         skoXC1 * (-7623/20 + skoXC1 * (-1633/30 + skoXC1 * -71/30)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_182, refute_0_215])).

cnf(refute_0_217, plain,
    (77/2 <
     skoXC1 *
     (-581/4 +
      skoXC1 *
      (-441/2 +
       skoXC1 *
       (-2527/12 +
        skoXC1 *
        (-1925/12 +
         skoXC1 *
         (-3291/40 + skoXC1 * (-1201/60 + skoXC1 * -169/120)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (77/2 <
     skoXC1 *
     (-581/4 +
      skoXC1 *
      (-441/2 +
       skoXC1 *
       (-2527/12 +
        skoXC1 *
        (-1925/12 +
         skoXC1 *
         (-3291/40 + skoXC1 * (-1201/60 + skoXC1 * -169/120)))))) |
     -462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_213, refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_219, plain,
    (77/2 <
     skoXC1 *
     (-581/4 +
      skoXC1 *
      (-441/2 +
       skoXC1 *
       (-2527/12 +
        skoXC1 *
        (-1925/12 +
         skoXC1 *
         (-3291/40 + skoXC1 * (-1201/60 + skoXC1 * -169/120)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_217, refute_0_218])).

cnf(refute_0_220, plain,
    (-16 < skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_219, refute_0_213, refute_0_119, refute_0_23,
               refute_0_24])).

cnf(refute_0_221, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (-7392 +
      skoXC1 *
      (-35280 +
       skoXC1 *
       (-70840 +
        skoXC1 *
        (-77700 +
         skoXC1 *
         (-1011283/20 +
          skoXC1 *
          (-2380639/120 +
           skoXC1 *
           (-68317/15 +
            skoXC1 *
            (-135049/240 +
             skoXC1 * (-3763/120 + skoXC1 * -71/120))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_181, refute_0_220])).

cnf(refute_0_222, plain,
    ((4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-7392 +
      skoXC1 *
      (-35280 +
       skoXC1 *
       (-70840 +
        skoXC1 *
        (-77700 +
         skoXC1 *
         (-1011283/20 +
          skoXC1 *
          (-2380639/120 +
           skoXC1 *
           (-68317/15 +
            skoXC1 *
            (-135049/240 +
             skoXC1 * (-3763/120 + skoXC1 * -71/120))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_172, refute_0_221])).

cnf(refute_0_223, plain,
    (616 <
     skoXC1 *
     (-3248 +
      skoXC1 *
      (-7322 +
       skoXC1 *
       (-26698/3 +
        skoXC1 *
        (-283367/48 +
         skoXC1 *
         (-48977/30 +
          skoXC1 *
          (29759/80 +
           skoXC1 *
           (2081/5 +
            skoXC1 *
            (57287/480 + skoXC1 * (209/15 + skoXC1 * 191/480))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (616 <
     skoXC1 *
     (-3248 +
      skoXC1 *
      (-7322 +
       skoXC1 *
       (-26698/3 +
        skoXC1 *
        (-283367/48 +
         skoXC1 *
         (-48977/30 +
          skoXC1 *
          (29759/80 +
           skoXC1 *
           (2081/5 +
            skoXC1 *
            (57287/480 + skoXC1 * (209/15 + skoXC1 * 191/480))))))))) |
     -462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_219, refute_0_213, refute_0_119, refute_0_23,
               refute_0_24])).

cnf(refute_0_225, plain,
    (616 <
     skoXC1 *
     (-3248 +
      skoXC1 *
      (-7322 +
       skoXC1 *
       (-26698/3 +
        skoXC1 *
        (-283367/48 +
         skoXC1 *
         (-48977/30 +
          skoXC1 *
          (29759/80 +
           skoXC1 *
           (2081/5 +
            skoXC1 *
            (57287/480 + skoXC1 * (209/15 + skoXC1 * 191/480))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (skoXC1 < 0 | -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 <= -1),
    inference(decision, [],
              [refute_0_225, refute_0_219, refute_0_213, refute_0_119,
               refute_0_23, refute_0_24])).

cnf(refute_0_227, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_171, refute_0_226])).

cnf(refute_0_228, plain,
    (-1 < skoXC1 | -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_225, refute_0_219, refute_0_213, refute_0_119,
               refute_0_23, refute_0_24])).

cnf(refute_0_229, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     (4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) /
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_227, refute_0_228])).

cnf(refute_0_230, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) *
     (-16 +
      skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -16 +
     skoXC1 * (-32 + skoXC1 * (-20 + skoXC1 * (-4 + skoXC1 * -1/8))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_147, refute_0_229])).

cnf(refute_0_231, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (-27456 +
      skoXC1 *
      (-144672 +
       skoXC1 *
       (-326480 +
        skoXC1 *
        (-411960 +
         skoXC1 *
         (-3182897/10 +
          skoXC1 *
          (-619957/4 +
           skoXC1 *
           (-1651016/35 +
            skoXC1 *
            (-1202939/140 +
             skoXC1 *
             (-119617/140 +
              skoXC1 * (-21803/560 + skoXC1 * -353/560)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8))) <= -16),
    inference(arithmetic, [], [refute_0_230])).

cnf(refute_0_232, plain,
    (4/3 +
     skoXC1 *
     (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96))) <
     skoXC1 *
     (-27456 +
      skoXC1 *
      (-144672 +
       skoXC1 *
       (-326480 +
        skoXC1 *
        (-411960 +
         skoXC1 *
         (-3182897/10 +
          skoXC1 *
          (-619957/4 +
           skoXC1 *
           (-1651016/35 +
            skoXC1 *
            (-1202939/140 +
             skoXC1 *
             (-119617/140 +
              skoXC1 * (-21803/560 + skoXC1 * -353/560)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_231, refute_0_220])).

cnf(refute_0_233, plain,
    ((4/3 +
      skoXC1 *
      (-40/3 + skoXC1 * (-67/3 + skoXC1 * (-29/3 + skoXC1 * -95/96)))) *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     skoXC1 *
     (-27456 +
      skoXC1 *
      (-144672 +
       skoXC1 *
       (-326480 +
        skoXC1 *
        (-411960 +
         skoXC1 *
         (-3182897/10 +
          skoXC1 *
          (-619957/4 +
           skoXC1 *
           (-1651016/35 +
            skoXC1 *
            (-1202939/140 +
             skoXC1 *
             (-119617/140 +
              skoXC1 * (-21803/560 + skoXC1 * -353/560)))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_232])).

cnf(refute_0_234, plain,
    (2288 <
     skoXC1 *
     (-13200 +
      skoXC1 *
      (-33044 +
       skoXC1 *
       (-45880 +
        skoXC1 *
        (-295839/8 +
         skoXC1 *
         (-621483/40 +
          skoXC1 *
          (-43375/48 +
           skoXC1 *
           (1291109/560 +
            skoXC1 *
            (931241/840 +
             skoXC1 *
             (62571/280 +
              skoXC1 * (21529/1120 + skoXC1 * 1207/3360)))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <=
     -1716), inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (2288 <
     skoXC1 *
     (-13200 +
      skoXC1 *
      (-33044 +
       skoXC1 *
       (-45880 +
        skoXC1 *
        (-295839/8 +
         skoXC1 *
         (-621483/40 +
          skoXC1 *
          (-43375/48 +
           skoXC1 *
           (1291109/560 +
            skoXC1 *
            (931241/840 +
             skoXC1 *
             (62571/280 +
              skoXC1 * (21529/1120 + skoXC1 * 1207/3360)))))))))) |
     -1716 <
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_225, refute_0_219, refute_0_213, refute_0_119,
               refute_0_23, refute_0_24])).

cnf(refute_0_236, plain,
    (2288 <
     skoXC1 *
     (-13200 +
      skoXC1 *
      (-33044 +
       skoXC1 *
       (-45880 +
        skoXC1 *
        (-295839/8 +
         skoXC1 *
         (-621483/40 +
          skoXC1 *
          (-43375/48 +
           skoXC1 *
           (1291109/560 +
            skoXC1 *
            (931241/840 +
             skoXC1 *
             (62571/280 +
              skoXC1 * (21529/1120 + skoXC1 * 1207/3360)))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain,
    (-1 <= skoXC1 | -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 |
     1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_236, refute_0_225, refute_0_219, refute_0_213,
               refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_238, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (198593/32 +
      skoXC1 *
      (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048)))) /
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) |
     -1/12 <= ln(1 + skoXC1) | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_237, refute_0_145])).

cnf(refute_0_239, plain,
    (skoXC1 < 0 | -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_236, refute_0_225, refute_0_219, refute_0_213,
               refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_240, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) <
     (198593/32 +
      skoXC1 *
      (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048)))) /
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_238, refute_0_239])).

cnf(refute_0_241, plain,
    (skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) <
     198593/32 +
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096)) <= 0),
    inference(resolve, [], [refute_0_136, refute_0_240])).

cnf(refute_0_242, plain,
    (-198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) <
     skoXC1 *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) /
     (-1/12 + ln(1 + skoXC1) * -1) | -1/12 <= ln(1 + skoXC1) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096)) <= -6001),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (-6001 < skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096)) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_236, refute_0_225, refute_0_219, refute_0_213,
               refute_0_119, refute_0_23, refute_0_24])).

cnf(refute_0_244, plain,
    (-198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) <
     skoXC1 *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) /
     (-1/12 + ln(1 + skoXC1) * -1) | -1/12 <= ln(1 + skoXC1) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_242, refute_0_243])).

cnf(refute_0_245, plain,
    ((-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) *
     (-1/12 + ln(1 + skoXC1) * -1) <
     skoXC1 *
     (6001 + skoXC1 * (15984 + skoXC1 * (14080 + skoXC1 * 4096))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     -1/12 + ln(1 + skoXC1) * -1 <= 0),
    inference(resolve, [], [refute_0_135, refute_0_244])).

cnf(refute_0_246, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     ln(1 + skoXC1) *
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (ln(1 + skoXC1) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) |
     1/12 <= ln(1 + skoXC1)), inference(split, [], [refute_0_246])).

cnf(refute_0_248, plain,
    (ln(1 + skoXC1) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) <=
     -198593/32), inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (-198593/32 <
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_24, refute_0_23, refute_0_119, refute_0_213,
               refute_0_219, refute_0_225, refute_0_236])).

cnf(refute_0_250, plain,
    (ln(1 + skoXC1) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    ((-1/7 +
      (1 + skoXC1) *
      (-273/10 +
       (1 + skoXC1) *
       (-1302/5 +
        (1 + skoXC1) *
        (-875/2 +
         (1 + skoXC1) *
         (175 +
          (1 + skoXC1) *
          (861/2 + (1 + skoXC1) * (574/5 + (1 + skoXC1) * 353/70))))))) /
     ((1 + skoXC1) *
      (7 +
       (1 + skoXC1) *
       (126 +
        (1 + skoXC1) *
        (525 +
         (1 + skoXC1) *
         (700 +
          (1 + skoXC1) * (315 + (1 + skoXC1) * (42 + (1 + skoXC1)))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_134, refute_0_250])).

cnf(refute_0_252, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_252, refute_0_228])).

cnf(refute_0_254, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) *
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_129, refute_0_253])).

cnf(refute_0_255, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (-85196397/8 +
      skoXC1 *
      (-1108314933/16 +
       skoXC1 *
       (-196825640 +
        skoXC1 *
        (-5122870785/16 +
         skoXC1 *
         (-26253998839/80 +
          skoXC1 *
          (-14058048697/64 +
           skoXC1 *
           (-215360364641/2240 +
            skoXC1 *
            (-935424479/35 +
             skoXC1 *
             (-153852952/35 +
              skoXC1 * (-12837632/35 + skoXC1 * -361472/35)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) <=
     -198593/32), inference(arithmetic, [], [refute_0_254])).

cnf(refute_0_256, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (-85196397/8 +
      skoXC1 *
      (-1108314933/16 +
       skoXC1 *
       (-196825640 +
        skoXC1 *
        (-5122870785/16 +
         skoXC1 *
         (-26253998839/80 +
          skoXC1 *
          (-14058048697/64 +
           skoXC1 *
           (-215360364641/2240 +
            skoXC1 *
            (-935424479/35 +
             skoXC1 *
             (-153852952/35 +
              skoXC1 * (-12837632/35 + skoXC1 * -361472/35)))))))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_255, refute_0_249])).

cnf(refute_0_257, plain,
    ((198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     skoXC1 *
     (-85196397/8 +
      skoXC1 *
      (-1108314933/16 +
       skoXC1 *
       (-196825640 +
        skoXC1 *
        (-5122870785/16 +
         skoXC1 *
         (-26253998839/80 +
          skoXC1 *
          (-14058048697/64 +
           skoXC1 *
           (-215360364641/2240 +
            skoXC1 *
            (-935424479/35 +
             skoXC1 *
             (-153852952/35 +
              skoXC1 * (-12837632/35 + skoXC1 * -361472/35)))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_128, refute_0_256])).

cnf(refute_0_258, plain,
    (28398799/32 <
     skoXC1 *
     (-210177231/32 +
      skoXC1 *
      (-1423429733/64 +
       skoXC1 *
       (-2911793353/64 +
        skoXC1 *
        (-1997955543/32 +
         skoXC1 *
         (-9640884963/160 +
          skoXC1 *
          (-15962051399/384 +
           skoXC1 *
           (-91069445667/4480 +
            skoXC1 *
            (-1426475509/210 +
             skoXC1 *
             (-50114772/35 +
              skoXC1 * (-5647232/35 + skoXC1 * -672256/105)))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <=
     -1716), inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain,
    ((198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-45874983/16 +
      skoXC1 *
      (-551230785/32 +
       skoXC1 *
       (-178465945/4 +
        skoXC1 *
        (-4161899175/64 +
         skoXC1 *
         (-9337418321/160 +
          skoXC1 *
          (-15895058119/480 +
           skoXC1 *
           (-175717138/15 +
            skoXC1 *
            (-36544336/15 +
             skoXC1 * (-3889664/15 + skoXC1 * -145408/15))))))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0 |
     skoXC1 *
     (-45874983/16 +
      skoXC1 *
      (-551230785/32 +
       skoXC1 *
       (-178465945/4 +
        skoXC1 *
        (-4161899175/64 +
         skoXC1 *
         (-9337418321/160 +
          skoXC1 *
          (-15895058119/480 +
           skoXC1 *
           (-175717138/15 +
            skoXC1 *
            (-36544336/15 +
             skoXC1 * (-3889664/15 + skoXC1 * -145408/15))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <=
     198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_260, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     0 <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) |
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_261, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     1 + skoXC1 <= 0 |
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) <=
     ln(1 + skoXC1)), inference(subst, [], [refute_0_46])).

cnf(refute_0_262, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_261, refute_0_250])).

cnf(refute_0_263, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) /
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_263, refute_0_228])).

cnf(refute_0_265, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-198593/32 +
      skoXC1 *
      (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048)))) |
     -1/12 <= ln(1 + skoXC1) |
     0 <=
     -198593/32 +
     skoXC1 *
     (-20078 + skoXC1 * (-23600 + skoXC1 * (-11776 + skoXC1 * -2048))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_260, refute_0_264])).

cnf(refute_0_266, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (-45874983/16 +
      skoXC1 *
      (-551230785/32 +
       skoXC1 *
       (-178465945/4 +
        skoXC1 *
        (-4161899175/64 +
         skoXC1 *
         (-9337418321/160 +
          skoXC1 *
          (-15895058119/480 +
           skoXC1 *
           (-175717138/15 +
            skoXC1 *
            (-36544336/15 +
             skoXC1 * (-3889664/15 + skoXC1 * -145408/15))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (20078 + skoXC1 * (23600 + skoXC1 * (11776 + skoXC1 * 2048))) <=
     -198593/32), inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (198593/384 +
     skoXC1 *
     (-25967/6 +
      skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3))) <
     skoXC1 *
     (-45874983/16 +
      skoXC1 *
      (-551230785/32 +
       skoXC1 *
       (-178465945/4 +
        skoXC1 *
        (-4161899175/64 +
         skoXC1 *
         (-9337418321/160 +
          skoXC1 *
          (-15895058119/480 +
           skoXC1 *
           (-175717138/15 +
            skoXC1 *
            (-36544336/15 +
             skoXC1 * (-3889664/15 + skoXC1 * -145408/15))))))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_266, refute_0_249])).

cnf(refute_0_268, plain,
    ((198593/384 +
      skoXC1 *
      (-25967/6 +
       skoXC1 * (-42052/3 + skoXC1 * (-39296/3 + skoXC1 * -11776/3)))) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-45874983/16 +
      skoXC1 *
      (-551230785/32 +
       skoXC1 *
       (-178465945/4 +
        skoXC1 *
        (-4161899175/64 +
         skoXC1 *
         (-9337418321/160 +
          skoXC1 *
          (-15895058119/480 +
           skoXC1 *
           (-175717138/15 +
            skoXC1 *
            (-36544336/15 +
             skoXC1 * (-3889664/15 + skoXC1 * -145408/15))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_259, refute_0_267])).

cnf(refute_0_269, plain,
    (15291661/64 <
     skoXC1 *
     (-13197499/8 +
      skoXC1 *
      (-331758077/64 +
       skoXC1 *
       (-58621913/6 +
        skoXC1 *
        (-4697841253/384 +
         skoXC1 *
         (-318696533/30 +
          skoXC1 *
          (-4130926667/640 +
           skoXC1 *
           (-26830707/10 +
            skoXC1 *
            (-10713596/15 +
             skoXC1 * (-1573504/15 + skoXC1 * -86528/15))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_268])).

cnf(refute_0_270, plain,
    (15291661/64 <
     skoXC1 *
     (-13197499/8 +
      skoXC1 *
      (-331758077/64 +
       skoXC1 *
       (-58621913/6 +
        skoXC1 *
        (-4697841253/384 +
         skoXC1 *
         (-318696533/30 +
          skoXC1 *
          (-4130926667/640 +
           skoXC1 *
           (-26830707/10 +
            skoXC1 *
            (-10713596/15 +
             skoXC1 * (-1573504/15 + skoXC1 * -86528/15))))))))) |
     -462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_24, refute_0_23, refute_0_119, refute_0_213,
               refute_0_219, refute_0_225, refute_0_236])).

cnf(refute_0_271, plain,
    (15291661/64 <
     skoXC1 *
     (-13197499/8 +
      skoXC1 *
      (-331758077/64 +
       skoXC1 *
       (-58621913/6 +
        skoXC1 *
        (-4697841253/384 +
         skoXC1 *
         (-318696533/30 +
          skoXC1 *
          (-4130926667/640 +
           skoXC1 *
           (-26830707/10 +
            skoXC1 *
            (-10713596/15 +
             skoXC1 * (-1573504/15 + skoXC1 * -86528/15))))))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_269, refute_0_270])).

cnf(refute_0_272, plain,
    (-1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (-210177231/32 +
      skoXC1 *
      (-1423429733/64 +
       skoXC1 *
       (-2911793353/64 +
        skoXC1 *
        (-1997955543/32 +
         skoXC1 *
         (-9640884963/160 +
          skoXC1 *
          (-15962051399/384 +
           skoXC1 *
           (-91069445667/4480 +
            skoXC1 *
            (-1426475509/210 +
             skoXC1 *
             (-50114772/35 +
              skoXC1 * (-5647232/35 + skoXC1 * -672256/105)))))))))) <=
     28398799/32 |
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <=
     -1716),
    inference(decision, [],
              [refute_0_236, refute_0_225, refute_0_219, refute_0_213,
               refute_0_119, refute_0_23, refute_0_24, refute_0_271])).

cnf(refute_0_273, plain,
    (-1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <=
     -1716), inference(resolve, [], [refute_0_272, refute_0_258])).

cnf(refute_0_274, plain,
    (-1716 <
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) |
     -1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_236, refute_0_225, refute_0_219, refute_0_213,
               refute_0_119, refute_0_23, refute_0_24, refute_0_271])).

cnf(refute_0_275, plain,
    (-1/12 <= ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_273, refute_0_274])).

cnf(refute_0_276, plain,
    (-1/12 < ln(1 + skoXC1) | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_275, refute_0_127])).

cnf(refute_0_277, plain,
    (-1/12 *
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) |
     6435 +
     skoXC1 *
     (24024 +
      skoXC1 *
      (36036 +
       skoXC1 *
       (27720 +
        skoXC1 *
        (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8)))))) <= 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <=
     -1/12), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_278, plain,
    (-1/12 <
     (-1487/280 +
      (1 + skoXC1) *
      (-852/5 +
       (1 + skoXC1) *
       (-5096/5 +
        (1 + skoXC1) *
        (-1372 +
         (1 + skoXC1) *
         (1225/2 +
          (1 + skoXC1) *
          (7252/5 +
           (1 + skoXC1) *
           (2352/5 + (1 + skoXC1) * (1172/35 + (1 + skoXC1) * 1/8)))))))) /
     (1 +
      (1 + skoXC1) *
      (56 +
       (1 + skoXC1) *
       (588 +
        (1 + skoXC1) *
        (1960 +
         (1 + skoXC1) *
         (2450 +
          (1 + skoXC1) *
          (1176 + (1 + skoXC1) * (196 + (1 + skoXC1) * 8))))))) |
     1 + skoXC1 <= 0 | ln(1 + skoXC1) <= -1/12),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_279, plain,
    (-1/12 <
     (-1487/280 +
      (1 + skoXC1) *
      (-852/5 +
       (1 + skoXC1) *
       (-5096/5 +
        (1 + skoXC1) *
        (-1372 +
         (1 + skoXC1) *
         (1225/2 +
          (1 + skoXC1) *
          (7252/5 +
           (1 + skoXC1) *
           (2352/5 + (1 + skoXC1) * (1172/35 + (1 + skoXC1) * 1/8)))))))) /
     (1 +
      (1 + skoXC1) *
      (56 +
       (1 + skoXC1) *
       (588 +
        (1 + skoXC1) *
        (1960 +
         (1 + skoXC1) *
         (2450 +
          (1 + skoXC1) *
          (1176 + (1 + skoXC1) * (196 + (1 + skoXC1) * 8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_278, refute_0_276])).

cnf(refute_0_280, plain,
    (-1/12 <
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (-1/12 <
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) /
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_280, refute_0_64])).

cnf(refute_0_282, plain,
    (-1/12 *
     (6435 +
      skoXC1 *
      (24024 +
       skoXC1 *
       (36036 +
        skoXC1 *
        (27720 +
         skoXC1 *
         (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (41613/2 +
       skoXC1 *
       (26169 +
        skoXC1 *
        (64405/4 +
         skoXC1 *
         (4983 +
          skoXC1 * (7083/10 + skoXC1 * (1207/35 + skoXC1 * 1/8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     6435 +
     skoXC1 *
     (24024 +
      skoXC1 *
      (36036 +
       skoXC1 *
       (27720 +
        skoXC1 *
        (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8)))))) <= 0),
    inference(resolve, [], [refute_0_277, refute_0_281])).

cnf(refute_0_283, plain,
    (-2145/4 <
     skoXC1 *
     (8437 +
      skoXC1 *
      (47619/2 +
       skoXC1 *
       (28479 +
        skoXC1 *
        (68255/4 +
         skoXC1 *
         (5193 +
          skoXC1 * (7293/10 + skoXC1 * (3691/105 + skoXC1 * 1/8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (24024 +
      skoXC1 *
      (36036 +
       skoXC1 *
       (27720 +
        skoXC1 *
        (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8)))))) <=
     -6435), inference(arithmetic, [], [refute_0_282])).

cnf(refute_0_284, plain,
    (-1/12 *
     (1716 +
      skoXC1 *
      (5544 +
       skoXC1 *
       (6930 +
        skoXC1 *
        (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7)))))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (4686 +
       skoXC1 *
       (4730 +
        skoXC1 *
        (2154 + skoXC1 * (2136/5 + skoXC1 * (283/10 + skoXC1 * 1/7)))))) |
     1716 +
     skoXC1 *
     (5544 +
      skoXC1 *
      (6930 +
       skoXC1 * (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7))))) <=
     0 |
     skoXC1 *
     (1716 +
      skoXC1 *
      (4686 +
       skoXC1 *
       (4730 +
        skoXC1 *
        (2154 + skoXC1 * (2136/5 + skoXC1 * (283/10 + skoXC1 * 1/7)))))) /
     (1716 +
      skoXC1 *
      (5544 +
       skoXC1 *
       (6930 +
        skoXC1 *
        (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7)))))) <=
     -1/12), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_285, plain,
    (~ lgen(0, ln(X_000092), X_000093) | ln(X_000092) <= X_000093),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_286, plain,
    (~ lgen(0,
            1/70 *
            (353 + 1921 * X_000092 ^ 5 + 8389 * X_000092 +
             38524 * X_000092 ^ 2 + 20149 * X_000092 ^ 4 +
             50774 * X_000092 ^ 3 + 10 * X_000092 ^ 6) * (X_000092 - 1) /
            (1 + 126 * X_000092 ^ 5 + 42 * X_000092 + 315 * X_000092 ^ 2 +
             525 * X_000092 ^ 4 + 700 * X_000092 ^ 3 + 7 * X_000092 ^ 6),
            X_000093) | X_000092 <= 0 | lgen(0, ln(X_000092), X_000093)),
    inference(subst, [], [ln_upper_bound_cf13])).

cnf(refute_0_287, plain,
    (~ lgen(0,
            1/70 *
            (353 + 1921 * X_000092 ^ 5 + 8389 * X_000092 +
             38524 * X_000092 ^ 2 + 20149 * X_000092 ^ 4 +
             50774 * X_000092 ^ 3 + 10 * X_000092 ^ 6) * (X_000092 - 1) /
            (1 + 126 * X_000092 ^ 5 + 42 * X_000092 + 315 * X_000092 ^ 2 +
             525 * X_000092 ^ 4 + 700 * X_000092 ^ 3 + 7 * X_000092 ^ 6),
            X_000093) | X_000092 <= 0 | ln(X_000092) <= X_000093),
    inference(resolve, [], [refute_0_286, refute_0_285])).

cnf(refute_0_288, plain,
    (X_000093 <
     (-353/70 +
      X_000092 *
      (-574/5 +
       X_000092 *
       (-861/2 +
        X_000092 *
        (-175 +
         X_000092 *
         (875/2 +
          X_000092 * (1302/5 + X_000092 * (273/10 + X_000092 * 1/7))))))) /
     (1 +
      X_000092 *
      (42 +
       X_000092 *
       (315 +
        X_000092 *
        (700 + X_000092 * (525 + X_000092 * (126 + X_000092 * 7)))))) |
     X_000092 <= 0 | ln(X_000092) <= X_000093),
    inference(arithmetic, [], [refute_0_287])).

cnf(refute_0_289, plain,
    (-1/12 <
     (-353/70 +
      (1 + skoXC1) *
      (-574/5 +
       (1 + skoXC1) *
       (-861/2 +
        (1 + skoXC1) *
        (-175 +
         (1 + skoXC1) *
         (875/2 +
          (1 + skoXC1) *
          (1302/5 + (1 + skoXC1) * (273/10 + (1 + skoXC1) * 1/7))))))) /
     (1 +
      (1 + skoXC1) *
      (42 +
       (1 + skoXC1) *
       (315 +
        (1 + skoXC1) *
        (700 +
         (1 + skoXC1) *
         (525 + (1 + skoXC1) * (126 + (1 + skoXC1) * 7)))))) |
     1 + skoXC1 <= 0 | ln(1 + skoXC1) <= -1/12),
    inference(subst, [], [refute_0_288])).

cnf(refute_0_290, plain,
    (-1/12 <
     (-353/70 +
      (1 + skoXC1) *
      (-574/5 +
       (1 + skoXC1) *
       (-861/2 +
        (1 + skoXC1) *
        (-175 +
         (1 + skoXC1) *
         (875/2 +
          (1 + skoXC1) *
          (1302/5 + (1 + skoXC1) * (273/10 + (1 + skoXC1) * 1/7))))))) /
     (1 +
      (1 + skoXC1) *
      (42 +
       (1 + skoXC1) *
       (315 +
        (1 + skoXC1) *
        (700 +
         (1 + skoXC1) *
         (525 + (1 + skoXC1) * (126 + (1 + skoXC1) * 7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_289, refute_0_276])).

cnf(refute_0_291, plain,
    (-1/12 <
     skoXC1 *
     (1716 +
      skoXC1 *
      (4686 +
       skoXC1 *
       (4730 +
        skoXC1 *
        (2154 + skoXC1 * (2136/5 + skoXC1 * (283/10 + skoXC1 * 1/7)))))) /
     (1716 +
      skoXC1 *
      (5544 +
       skoXC1 *
       (6930 +
        skoXC1 *
        (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (-1/12 <
     skoXC1 *
     (1716 +
      skoXC1 *
      (4686 +
       skoXC1 *
       (4730 +
        skoXC1 *
        (2154 + skoXC1 * (2136/5 + skoXC1 * (283/10 + skoXC1 * 1/7)))))) /
     (1716 +
      skoXC1 *
      (5544 +
       skoXC1 *
       (6930 +
        skoXC1 *
        (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_291, refute_0_64])).

cnf(refute_0_293, plain,
    (-1/12 *
     (1716 +
      skoXC1 *
      (5544 +
       skoXC1 *
       (6930 +
        skoXC1 *
        (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7)))))) <
     skoXC1 *
     (1716 +
      skoXC1 *
      (4686 +
       skoXC1 *
       (4730 +
        skoXC1 *
        (2154 + skoXC1 * (2136/5 + skoXC1 * (283/10 + skoXC1 * 1/7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     1716 +
     skoXC1 *
     (5544 +
      skoXC1 *
      (6930 +
       skoXC1 * (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7))))) <=
     0), inference(resolve, [], [refute_0_284, refute_0_292])).

cnf(refute_0_294, plain,
    (-143 <
     skoXC1 *
     (2178 +
      skoXC1 *
      (10527/2 +
       skoXC1 *
       (5080 +
        skoXC1 *
        (2259 + skoXC1 * (2206/5 + skoXC1 * (1733/60 + skoXC1 * 1/7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (5544 +
      skoXC1 *
      (6930 +
       skoXC1 * (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7))))) <=
     -1716), inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (7451/35 + skoXC1 * 1487/280))))))) <
     1/12 *
     (6435 +
      skoXC1 *
      (27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (44352 +
         skoXC1 *
         (23100 +
          skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1)))))))) |
     1/12 <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (7451/35 + skoXC1 * 1487/280))))))) /
     (6435 +
      skoXC1 *
      (27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (44352 +
         skoXC1 *
         (23100 +
          skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1)))))))) |
     6435 +
     skoXC1 *
     (27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (44352 +
        skoXC1 *
        (23100 +
         skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_296, plain,
    (~ lgen(0, X_000121, ln(X_000120)) | X_000121 <= ln(X_000120)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_297, plain,
    (~ lgen(0, X_000121,
            1/280 *
            (35 + 9411 * X_000120 + 141123 * X_000120 ^ 2 +
             547235 * X_000120 ^ 3 + 718735 * X_000120 ^ 4 +
             334575 * X_000120 ^ 5 + 49199 * X_000120 ^ 6 +
             1487 * X_000120 ^ 7) * (X_000120 - 1) /
            (X_000120 *
             (8 + 196 * X_000120 + 1176 * X_000120 ^ 2 +
              2450 * X_000120 ^ 3 + 1960 * X_000120 ^ 4 +
              588 * X_000120 ^ 5 + 56 * X_000120 ^ 6 + X_000120 ^ 7))) |
     X_000120 <= 0 | lgen(0, X_000121, ln(X_000120))),
    inference(subst, [], [ln_lower_bound_cf15])).

cnf(refute_0_298, plain,
    (~ lgen(0, X_000121,
            1/280 *
            (35 + 9411 * X_000120 + 141123 * X_000120 ^ 2 +
             547235 * X_000120 ^ 3 + 718735 * X_000120 ^ 4 +
             334575 * X_000120 ^ 5 + 49199 * X_000120 ^ 6 +
             1487 * X_000120 ^ 7) * (X_000120 - 1) /
            (X_000120 *
             (8 + 196 * X_000120 + 1176 * X_000120 ^ 2 +
              2450 * X_000120 ^ 3 + 1960 * X_000120 ^ 4 +
              588 * X_000120 ^ 5 + 56 * X_000120 ^ 6 + X_000120 ^ 7))) |
     X_000120 <= 0 | X_000121 <= ln(X_000120)),
    inference(resolve, [], [refute_0_297, refute_0_296])).

cnf(refute_0_299, plain,
    ((-1/8 +
      X_000120 *
      (-1172/35 +
       X_000120 *
       (-2352/5 +
        X_000120 *
        (-7252/5 +
         X_000120 *
         (-1225/2 +
          X_000120 *
          (1372 +
           X_000120 *
           (5096/5 + X_000120 * (852/5 + X_000120 * 1487/280)))))))) /
     (X_000120 *
      (8 +
       X_000120 *
       (196 +
        X_000120 *
        (1176 +
         X_000120 *
         (2450 +
          X_000120 *
          (1960 + X_000120 * (588 + X_000120 * (56 + X_000120)))))))) <
     X_000121 | X_000120 <= 0 | X_000121 <= ln(X_000120)),
    inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    ((-1/8 +
      (1 + skoXC1) *
      (-1172/35 +
       (1 + skoXC1) *
       (-2352/5 +
        (1 + skoXC1) *
        (-7252/5 +
         (1 + skoXC1) *
         (-1225/2 +
          (1 + skoXC1) *
          (1372 +
           (1 + skoXC1) *
           (5096/5 +
            (1 + skoXC1) * (852/5 + (1 + skoXC1) * 1487/280)))))))) /
     ((1 + skoXC1) *
      (8 +
       (1 + skoXC1) *
       (196 +
        (1 + skoXC1) *
        (1176 +
         (1 + skoXC1) *
         (2450 +
          (1 + skoXC1) *
          (1960 +
           (1 + skoXC1) *
           (588 + (1 + skoXC1) * (56 + (1 + skoXC1))))))))) < 1/12 |
     1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_299])).

cnf(refute_0_301, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (7451/35 + skoXC1 * 1487/280))))))) /
     (6435 +
      skoXC1 *
      (27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (44352 +
         skoXC1 *
         (23100 +
          skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1)))))))) <
     1/12 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_300])).

cnf(refute_0_302, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (7451/35 + skoXC1 * 1487/280))))))) <
     1/12 *
     (6435 +
      skoXC1 *
      (27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (44352 +
         skoXC1 *
         (23100 +
          skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1)))))))) |
     1/12 <= ln(1 + skoXC1) |
     6435 +
     skoXC1 *
     (27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (44352 +
        skoXC1 *
        (23100 +
         skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1))))))) <=
     0 | skoXC1 <= -1),
    inference(resolve, [], [refute_0_295, refute_0_301])).

cnf(refute_0_303, plain,
    (-2145/4 <
     skoXC1 *
     (-4147 +
      skoXC1 *
      (-40469/2 +
       skoXC1 *
       (-32769 +
        skoXC1 *
        (-103785/4 +
         skoXC1 *
         (-10803 +
          skoXC1 *
          (-22767/10 + skoXC1 * (-21793/105 + skoXC1 * -4391/840))))))) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (44352 +
        skoXC1 *
        (23100 +
         skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1))))))) <=
     -6435 | skoXC1 <= -1), inference(arithmetic, [], [refute_0_302])).

cnf(refute_0_304, plain,
    (-1/12 *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0 |
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <=
     -1/12), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_305, plain,
    (-1/12 <
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     1 + skoXC1 <= 0 | ln(1 + skoXC1) <= -1/12),
    inference(subst, [], [refute_0_106])).

cnf(refute_0_306, plain,
    (-1/12 <
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_305, refute_0_276])).

cnf(refute_0_307, plain,
    (-1/12 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (-1/12 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_307, refute_0_64])).

cnf(refute_0_309, plain,
    (-1/12 *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0), inference(resolve, [], [refute_0_304, refute_0_308])).

cnf(refute_0_310, plain,
    (-77/2 <
     skoXC1 *
     (567 +
      skoXC1 *
      (1134 +
       skoXC1 *
       (2492/3 + skoXC1 * (973/4 + skoXC1 * (229/10 + skoXC1 * 1/6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     -462), inference(arithmetic, [], [refute_0_309])).

cnf(refute_0_311, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) <
     1/12 *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     1/12 <=
     skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_312, plain,
    ((-1/7 +
      (1 + skoXC1) *
      (-273/10 +
       (1 + skoXC1) *
       (-1302/5 +
        (1 + skoXC1) *
        (-875/2 +
         (1 + skoXC1) *
         (175 +
          (1 + skoXC1) *
          (861/2 + (1 + skoXC1) * (574/5 + (1 + skoXC1) * 353/70))))))) /
     ((1 + skoXC1) *
      (7 +
       (1 + skoXC1) *
       (126 +
        (1 + skoXC1) *
        (525 +
         (1 + skoXC1) *
         (700 +
          (1 + skoXC1) * (315 + (1 + skoXC1) * (42 + (1 + skoXC1)))))))) <
     1/12 | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_133])).

cnf(refute_0_313, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) /
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) <
     1/12 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_312])).

cnf(refute_0_314, plain,
    (skoXC1 *
     (1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (6126/5 + skoXC1 * (1501/10 + skoXC1 * 353/70)))))) <
     1/12 *
     (1716 +
      skoXC1 *
      (6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (7350 +
         skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1))))))) |
     1/12 <= ln(1 + skoXC1) |
     1716 +
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <= 0 |
     skoXC1 <= -1), inference(resolve, [], [refute_0_311, refute_0_313])).

cnf(refute_0_315, plain,
    (-143 <
     skoXC1 *
     (-1177 +
      skoXC1 *
      (-9603/2 +
       skoXC1 *
       (-12855/2 +
        skoXC1 *
        (-3981 +
         skoXC1 *
         (-5881/5 + skoXC1 * (-8761/60 + skoXC1 * -2083/420)))))) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) <=
     -1716 | skoXC1 <= -1), inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <
     1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     1/12 <=
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_317, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     1/12 | 1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_318, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     1/12 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_317])).

cnf(refute_0_319, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) <
     1/12 *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0 |
     skoXC1 <= -1), inference(resolve, [], [refute_0_316, refute_0_318])).

cnf(refute_0_320, plain,
    (-77/2 <
     skoXC1 *
     (-336 +
      skoXC1 *
      (-2247/2 +
       skoXC1 *
       (-3584/3 +
        skoXC1 * (-2149/4 + skoXC1 * (-482/5 + skoXC1 * -93/20))))) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462 | skoXC1 <= -1), inference(arithmetic, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (-77/2 <
     skoXC1 *
     (-336 +
      skoXC1 *
      (-2247/2 +
       skoXC1 *
       (-3584/3 +
        skoXC1 * (-2149/4 + skoXC1 * (-482/5 + skoXC1 * -93/20))))) |
     -462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     skoXC1 <= -1), inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_322, plain,
    (-77/2 <
     skoXC1 *
     (-336 +
      skoXC1 *
      (-2247/2 +
       skoXC1 *
       (-3584/3 +
        skoXC1 * (-2149/4 + skoXC1 * (-482/5 + skoXC1 * -93/20))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_320, refute_0_321])).

cnf(refute_0_323, plain,
    (-77/2 <
     skoXC1 *
     (-336 +
      skoXC1 *
      (-2247/2 +
       skoXC1 *
       (-3584/3 +
        skoXC1 * (-2149/4 + skoXC1 * (-482/5 + skoXC1 * -93/20))))) |
     -1 < skoXC1), inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_324, plain,
    (-77/2 <
     skoXC1 *
     (-336 +
      skoXC1 *
      (-2247/2 +
       skoXC1 *
       (-3584/3 +
        skoXC1 * (-2149/4 + skoXC1 * (-482/5 + skoXC1 * -93/20))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_322, refute_0_323])).

cnf(refute_0_325, plain,
    (-143 <
     skoXC1 *
     (-1177 +
      skoXC1 *
      (-9603/2 +
       skoXC1 *
       (-12855/2 +
        skoXC1 *
        (-3981 +
         skoXC1 *
         (-5881/5 + skoXC1 * (-8761/60 + skoXC1 * -2083/420)))))) |
     -1716 <
     skoXC1 *
     (6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (7350 +
        skoXC1 * (2940 + skoXC1 * (588 + skoXC1 * (49 + skoXC1)))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(decision, [], [refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_326, plain,
    (-143 <
     skoXC1 *
     (-1177 +
      skoXC1 *
      (-9603/2 +
       skoXC1 *
       (-12855/2 +
        skoXC1 *
        (-3981 +
         skoXC1 *
         (-5881/5 + skoXC1 * (-8761/60 + skoXC1 * -2083/420)))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_315, refute_0_325])).

cnf(refute_0_327, plain,
    (-143 <
     skoXC1 *
     (-1177 +
      skoXC1 *
      (-9603/2 +
       skoXC1 *
       (-12855/2 +
        skoXC1 *
        (-3981 +
         skoXC1 *
         (-5881/5 + skoXC1 * (-8761/60 + skoXC1 * -2083/420)))))) |
     -1 < skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [], [refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_328, plain,
    (-143 <
     skoXC1 *
     (-1177 +
      skoXC1 *
      (-9603/2 +
       skoXC1 *
       (-12855/2 +
        skoXC1 *
        (-3981 +
         skoXC1 *
         (-5881/5 + skoXC1 * (-8761/60 + skoXC1 * -2083/420)))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_326, refute_0_327])).

cnf(refute_0_329, plain,
    (-77/2 <
     skoXC1 *
     (567 +
      skoXC1 *
      (1134 +
       skoXC1 *
       (2492/3 + skoXC1 * (973/4 + skoXC1 * (229/10 + skoXC1 * 1/6))))) |
     -462 <
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_328, refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_330, plain,
    (-77/2 <
     skoXC1 *
     (567 +
      skoXC1 *
      (1134 +
       skoXC1 *
       (2492/3 + skoXC1 * (973/4 + skoXC1 * (229/10 + skoXC1 * 1/6))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_310, refute_0_329])).

cnf(refute_0_331, plain,
    (skoXC1 < 0 |
     -77/2 <
     skoXC1 *
     (567 +
      skoXC1 *
      (1134 +
       skoXC1 *
       (2492/3 + skoXC1 * (973/4 + skoXC1 * (229/10 + skoXC1 * 1/6))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_328, refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_332, plain,
    (-77/2 <
     skoXC1 *
     (567 +
      skoXC1 *
      (1134 +
       skoXC1 *
       (2492/3 + skoXC1 * (973/4 + skoXC1 * (229/10 + skoXC1 * 1/6))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_330, refute_0_331])).

cnf(refute_0_333, plain,
    (-2145/4 <
     skoXC1 *
     (-4147 +
      skoXC1 *
      (-40469/2 +
       skoXC1 *
       (-32769 +
        skoXC1 *
        (-103785/4 +
         skoXC1 *
         (-10803 +
          skoXC1 *
          (-22767/10 + skoXC1 * (-21793/105 + skoXC1 * -4391/840))))))) |
     -6435 <
     skoXC1 *
     (27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (44352 +
        skoXC1 *
        (23100 +
         skoXC1 * (6720 + skoXC1 * (1008 + skoXC1 * (64 + skoXC1))))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(decision, [],
              [refute_0_332, refute_0_328, refute_0_324, refute_0_24,
               refute_0_23])).

cnf(refute_0_334, plain,
    (-2145/4 <
     skoXC1 *
     (-4147 +
      skoXC1 *
      (-40469/2 +
       skoXC1 *
       (-32769 +
        skoXC1 *
        (-103785/4 +
         skoXC1 *
         (-10803 +
          skoXC1 *
          (-22767/10 + skoXC1 * (-21793/105 + skoXC1 * -4391/840))))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_303, refute_0_333])).

cnf(refute_0_335, plain,
    (-2145/4 <
     skoXC1 *
     (-4147 +
      skoXC1 *
      (-40469/2 +
       skoXC1 *
       (-32769 +
        skoXC1 *
        (-103785/4 +
         skoXC1 *
         (-10803 +
          skoXC1 *
          (-22767/10 + skoXC1 * (-21793/105 + skoXC1 * -4391/840))))))) |
     -1 < skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_332, refute_0_328, refute_0_324, refute_0_24,
               refute_0_23])).

cnf(refute_0_336, plain,
    (-2145/4 <
     skoXC1 *
     (-4147 +
      skoXC1 *
      (-40469/2 +
       skoXC1 *
       (-32769 +
        skoXC1 *
        (-103785/4 +
         skoXC1 *
         (-10803 +
          skoXC1 *
          (-22767/10 + skoXC1 * (-21793/105 + skoXC1 * -4391/840))))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_334, refute_0_335])).

cnf(refute_0_337, plain,
    (-143 <
     skoXC1 *
     (2178 +
      skoXC1 *
      (10527/2 +
       skoXC1 *
       (5080 +
        skoXC1 *
        (2259 + skoXC1 * (2206/5 + skoXC1 * (1733/60 + skoXC1 * 1/7)))))) |
     -1716 <
     skoXC1 *
     (5544 +
      skoXC1 *
      (6930 +
       skoXC1 * (4200 + skoXC1 * (1260 + skoXC1 * (168 + skoXC1 * 7))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_336, refute_0_332, refute_0_328, refute_0_324,
               refute_0_24, refute_0_23])).

cnf(refute_0_338, plain,
    (-143 <
     skoXC1 *
     (2178 +
      skoXC1 *
      (10527/2 +
       skoXC1 *
       (5080 +
        skoXC1 *
        (2259 + skoXC1 * (2206/5 + skoXC1 * (1733/60 + skoXC1 * 1/7)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_294, refute_0_337])).

cnf(refute_0_339, plain,
    (skoXC1 < 0 |
     -143 <
     skoXC1 *
     (2178 +
      skoXC1 *
      (10527/2 +
       skoXC1 *
       (5080 +
        skoXC1 *
        (2259 + skoXC1 * (2206/5 + skoXC1 * (1733/60 + skoXC1 * 1/7)))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_336, refute_0_332, refute_0_328, refute_0_324,
               refute_0_24, refute_0_23])).

cnf(refute_0_340, plain,
    (-143 <
     skoXC1 *
     (2178 +
      skoXC1 *
      (10527/2 +
       skoXC1 *
       (5080 +
        skoXC1 *
        (2259 + skoXC1 * (2206/5 + skoXC1 * (1733/60 + skoXC1 * 1/7)))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_338, refute_0_339])).

cnf(refute_0_341, plain,
    (-2145/4 <
     skoXC1 *
     (8437 +
      skoXC1 *
      (47619/2 +
       skoXC1 *
       (28479 +
        skoXC1 *
        (68255/4 +
         skoXC1 *
         (5193 +
          skoXC1 * (7293/10 + skoXC1 * (3691/105 + skoXC1 * 1/8))))))) |
     -6435 <
     skoXC1 *
     (24024 +
      skoXC1 *
      (36036 +
       skoXC1 *
       (27720 +
        skoXC1 *
        (11550 + skoXC1 * (2520 + skoXC1 * (252 + skoXC1 * 8)))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_340, refute_0_336, refute_0_332, refute_0_328,
               refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_342, plain,
    (-2145/4 <
     skoXC1 *
     (8437 +
      skoXC1 *
      (47619/2 +
       skoXC1 *
       (28479 +
        skoXC1 *
        (68255/4 +
         skoXC1 *
         (5193 +
          skoXC1 * (7293/10 + skoXC1 * (3691/105 + skoXC1 * 1/8))))))) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_283, refute_0_341])).

cnf(refute_0_343, plain,
    (skoXC1 < 0 |
     -2145/4 <
     skoXC1 *
     (8437 +
      skoXC1 *
      (47619/2 +
       skoXC1 *
       (28479 +
        skoXC1 *
        (68255/4 +
         skoXC1 *
         (5193 +
          skoXC1 * (7293/10 + skoXC1 * (3691/105 + skoXC1 * 1/8))))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_340, refute_0_336, refute_0_332, refute_0_328,
               refute_0_324, refute_0_24, refute_0_23])).

cnf(refute_0_344, plain,
    (-2145/4 <
     skoXC1 *
     (8437 +
      skoXC1 *
      (47619/2 +
       skoXC1 *
       (28479 +
        skoXC1 *
        (68255/4 +
         skoXC1 *
         (5193 +
          skoXC1 * (7293/10 + skoXC1 * (3691/105 + skoXC1 * 1/8))))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_342, refute_0_343])).

cnf(refute_0_345, plain,
    (-77/2 <
     skoXC1 *
     (588 +
      skoXC1 *
      (2877/2 +
       skoXC1 *
       (4144/3 +
        skoXC1 * (2359/4 + skoXC1 * (512/5 + skoXC1 * 289/60))))) |
     -1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462),
    inference(decision, [],
              [refute_0_276, refute_0_344, refute_0_340, refute_0_336,
               refute_0_332, refute_0_328, refute_0_324, refute_0_24,
               refute_0_23])).

cnf(refute_0_346, plain,
    (-1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(resolve, [], [refute_0_123, refute_0_345])).

cnf(refute_0_347, plain,
    (-462 <
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 + skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) |
     -1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_276, refute_0_344, refute_0_340, refute_0_336,
               refute_0_332, refute_0_328, refute_0_324, refute_0_24,
               refute_0_23])).

cnf(refute_0_348, plain, (-1/12 < ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_346, refute_0_347])).

cnf(refute_0_349, plain,
    (sqrt(1 + skoXC1) < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_71, refute_0_348])).

cnf(refute_0_350, plain,
    (1 + skoXC1 < 0 | 1 < 1 + skoXC1 |
     1 + skoXC1 < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_53, refute_0_349])).

cnf(refute_0_351, plain,
    (skoXC1 < -1 |
     1 + skoXC1 < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 < skoXC1 | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(arithmetic, [], [refute_0_350])).

cnf(refute_0_352, plain,
    (0 < skoXC1 | -1 <= skoXC1 | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_23, refute_0_24, refute_0_324, refute_0_328,
               refute_0_332, refute_0_336, refute_0_340, refute_0_344])).
SZS status Timeout for abs-problem-5-sqrt-weak.tptp
Wall time limit exceeded.

cnf(refute_0_353, plain,
    (1 + skoXC1 < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 < skoXC1 | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_352, refute_0_351])).

cnf(refute_0_354, plain,
    (0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_23, refute_0_24, refute_0_324, refute_0_328,
               refute_0_332, refute_0_336, refute_0_340, refute_0_344])).

cnf(refute_0_355, plain,
    (1 + skoXC1 < skoXC1 * -1 / (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_354, refute_0_353])).

cnf(refute_0_356, plain,
    (skoXC1 * -1 < (1 + skoXC1) * (-1/12 + ln(1 + skoXC1) * -1) |
     0 <= -1/12 + ln(1 + skoXC1) * -1 | 0 <= skoXC1 |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_48, refute_0_355])).

cnf(refute_0_357, plain,
    (1/12 + skoXC1 * -11/12 < ln(1 + skoXC1) * (-1 + skoXC1 * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | ln(1 + skoXC1) <= -1/12),
    inference(arithmetic, [], [refute_0_356])).

cnf(refute_0_358, plain,
    (1/12 + skoXC1 * -11/12 < ln(1 + skoXC1) * (-1 + skoXC1 * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_357, refute_0_348])).

cnf(refute_0_359, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     0 <= -1 + skoXC1 * -1 | 0 <= skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(split, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     0 <= skoXC1 | 1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_359])).

cnf(refute_0_361, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_360, refute_0_177])).

cnf(refute_0_362, plain,
    (ln(1 + skoXC1) < (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_361, refute_0_120])).

cnf(refute_0_363, plain,
    ((-1/6 +
      (1 + skoXC1) *
      (-107/5 +
       (1 + skoXC1) *
       (-125 +
        (1 + skoXC1) *
        (-200/3 +
         (1 + skoXC1) *
         (275/2 + (1 + skoXC1) * (71 + (1 + skoXC1) * 71/15)))))) /
     ((1 + skoXC1) *
      (6 +
       (1 + skoXC1) *
       (75 +
        (1 + skoXC1) *
        (200 +
         (1 + skoXC1) * (150 + (1 + skoXC1) * (30 + (1 + skoXC1))))))) <
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     1/12 <= ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_47, refute_0_362])).

cnf(refute_0_364, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_363])).

cnf(refute_0_365, plain,
    (skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     (1/12 + skoXC1 * -11/12) / (-1 + skoXC1 * -1) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_364, refute_0_120])).

cnf(refute_0_366, plain,
    (1/12 + skoXC1 * -11/12 <
     skoXC1 *
     (462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (1288 + skoXC1 * (1127/2 + skoXC1 * (497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) *
     (-1 + skoXC1 * -1) | 0 <= -1 + skoXC1 * -1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_42, refute_0_365])).

cnf(refute_0_367, plain,
    (1/12 + skoXC1 * -11/12 <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1743 +
       skoXC1 *
       (-2569 +
        skoXC1 *
        (-3703/2 +
         skoXC1 * (-6629/10 + skoXC1 * (-1562/15 + skoXC1 * -71/15)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     1/12 <= ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (1/12 + skoXC1 * -11/12 <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1743 +
       skoXC1 *
       (-2569 +
        skoXC1 *
        (-3703/2 +
         skoXC1 * (-6629/10 + skoXC1 * (-1562/15 + skoXC1 * -71/15)))))) /
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_367, refute_0_120])).

cnf(refute_0_369, plain,
    ((1/12 + skoXC1 * -11/12) *
     (462 +
      skoXC1 *
      (1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1)))))) <
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1743 +
       skoXC1 *
       (-2569 +
        skoXC1 *
        (-3703/2 +
         skoXC1 * (-6629/10 + skoXC1 * (-1562/15 + skoXC1 * -71/15)))))) |
     1/12 <= ln(1 + skoXC1) |
     462 +
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_41, refute_0_368])).

cnf(refute_0_370, plain,
    (77/2 <
     skoXC1 *
     (-329/2 +
      skoXC1 *
      (-1029/2 +
       skoXC1 *
       (-5579/6 +
        skoXC1 *
        (-10213/12 +
         skoXC1 * (-7543/20 + skoXC1 * (-4273/60 + skoXC1 * -229/60)))))) |
     1/12 <= ln(1 + skoXC1) |
     skoXC1 *
     (1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (1120 + skoXC1 * (315 + skoXC1 * (36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (skoXC1 *
     (-462 +
      skoXC1 *
      (-1029 +
       skoXC1 *
       (-784 + skoXC1 * (-469/2 + skoXC1 * (-112/5 + skoXC1 * -1/6))))) *
     (1 + skoXC1) <
     -77/2 +
     skoXC1 *
     (-1211/2 +
      skoXC1 *
      (-1470 +
       skoXC1 *
       (-4235/3 +
        skoXC1 * (-7385/12 + skoXC1 * (-457/4 + skoXC1 * -13/2))))) |
     1 + skoXC1 <= 0 |
     (-77/2 +
      skoXC1 *
      (-1211/2 +
       skoXC1 *
       (-1470 +
        skoXC1 *
        (-4235/3 +
         skoXC1 * (-7385/12 + skoXC1 * (-457/4 + skoXC1 * -13/2)))))) /
     (1 + skoXC1) <=
     skoXC1 *
     (-462 +
      skoXC1 *
      (-1029 +
       skoXC1 *
       (-784 + skoXC1 * (-469/2 + skoXC1 * (-112/5 + skoXC1 * -1/6)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_372, plain,
    ((1/12 + skoXC1 * 13/12) / (1 + skoXC1) *
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) |
     462 +
     skoXC1 *
     (1260 +
      skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6)))) <=
     0 |
     skoXC1 *
     (462 +
      skoXC1 *
      (1029 +
       skoXC1 *
       (784 + skoXC1 * (469/2 + skoXC1 * (112/5 + skoXC1 * 1/6))))) /
     (462 +
      skoXC1 *
      (1260 +
       skoXC1 * (1260 + skoXC1 * (560 + skoXC1 * (105 + skoXC1 * 6))))) <=
     (1/12 + skoXC1 * 13/12) / (1 + skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_373, plain,
    ((1/12 + skoXC1 * 13/12) / (1 + skoXC1) <
     (-71/15 +
      (1 + skoXC1) *
      (-71 +
       (1 + skoXC1) *
       (-275/2 +
        (1 + skoXC1) *
        (200/3 +
         (1 + skoXC1) *
         (125 + (1 + skoXC1) * (107/5 + (1 + skoXC1) * 1/6)))))) /
     (1 +
      (1 + skoXC1) *
      (30 +
       (1 + skoXC1) *
       (150 +
        (1 + skoXC1) * (200 + (1 + skoXC1) * (75 + (1 + skoXC1) * 6))))) |
     1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <= (1/12 + skoXC1 * 13/12) / (1 + skoXC1)),
    inference(subst, [], [refute_0_106])).

cnf(refute_0_374, plain,
    (skoXC1 < (1 + skoXC1) * (-1/12 + ln(1 + skoXC1)) |
     0 <= -1/12 + ln(1 + skoXC1) |
     skoXC1 / (-1/12 + ln(1 + skoXC1)) <= 1 + skoXC1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_375, plain,
    (1 + skoXC1 < 0 | 1 < 1 + skoXC1 |
     1 + skoXC1 < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     skoXC1 / (-1/12 + ln(1 + skoXC1)) <= sqrt(1 + skoXC1)),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_376, plain,
    (1 + skoXC1 < 0 | skoXC1 < 0 | 1 < 1 + skoXC1 |
     1 + skoXC1 < skoXC1 / (-1/12 + ln(1 + skoXC1)) |
     1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_375, refute_0_84])).

cnf(refute_0_377, plain,
    (1 + skoXC1 < 0 | skoXC1 < 0 |
     skoXC1 < (1 + skoXC1) * (-1/12 + ln(1 + skoXC1)) | 1 < 1 + skoXC1 |
     0 <= -1/12 + ln(1 + skoXC1) | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_374, refute_0_376])).

cnf(refute_0_378, plain,
    (skoXC1 < -1 | skoXC1 < 0 |
     1/12 + skoXC1 * 13/12 < ln(1 + skoXC1) * (1 + skoXC1) | 0 < skoXC1 |
     1/12 <= ln(1 + skoXC1)), inference(arithmetic, [], [refute_0_377])).

cnf(refute_0_379, plain, (skoXC1 < 0 | 0 < skoXC1 | -1 <= skoXC1),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_380, plain,
    (skoXC1 < 0 | 1/12 + skoXC1 * 13/12 < ln(1 + skoXC1) * (1 + skoXC1) |
     0 < skoXC1 | 1/12 <= ln(1 + skoXC1)),
    inference(resolve, [], [refute_0_379, refute_0_378])).
Processor time: 2.105 = 0.972 (Metis) + 1.133 (RCF)
