% SZS output start CNFRefutation for abs-problem-3.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/3 * (10 * X ^ 2 + 19 * X + 1) * (X - 1) /
            (X * (X ^ 2 + 6 * X + 3))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/12 * (47 * X ^ 3 + 239 * X ^ 2 + 131 * X + 3) * (X - 1) /
            (X * (X ^ 3 + 12 * X ^ 2 + 18 * X + 4))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (142 * X ^ 5 + 2272 * X ^ 4 + 6397 * X ^ 3 + 4397 * X ^ 2 +
             647 * X + 5) * (X - 1) /
            (X *
             (X ^ 5 + 30 * X ^ 4 + 150 * X ^ 3 + 200 * X ^ 2 + 75 * X +
              6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/70 *
            (10 + 1921 * X + 20149 * X ^ 2 + 50774 * X ^ 3 +
             38524 * X ^ 4 + 8389 * X ^ 5 + 353 * X ^ 6) * (X - 1) /
            (X *
             (7 + 126 * X + 525 * X ^ 2 + 700 * X ^ 3 + 315 * X ^ 4 +
              42 * X ^ 5 + X ^ 6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf15, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/280 *
            (35 + 9411 * X + 141123 * X ^ 2 + 547235 * X ^ 3 +
             718735 * X ^ 4 + 334575 * X ^ 5 + 49199 * X ^ 6 +
             1487 * X ^ 7) * (X - 1) /
            (X *
             (8 + 196 * X + 1176 * X ^ 2 + 2450 * X ^ 3 + 1960 * X ^ 4 +
              588 * X ^ 5 + 56 * X ^ 6 + X ^ 7))) | lgen(R, Y, ln(X)))).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(abs_problem_3, conjecture,
    (! [X] :
       ((0 < X & X <= 582811643/1000000000) =>
        abs(ln(1 - X)) <= 3 * X / 2))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X <= 582811643/1000000000) =>
        abs(ln(1 - X)) <= 3 * X / 2)),
    inference(strip, [], [abs_problem_3])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X <= 582811643/1000000000) =>
          abs(ln(1 - X)) <= 3 * X / 2)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & 3 * X / 2 < abs(ln(1 - X)) & X <= 582811643/1000000000)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (3 * skoXC1 / 2 < abs(ln(1 - skoXC1)) & 0 < skoXC1 &
     skoXC1 <= 582811643/1000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (3 * skoXC1 / 2 < abs(ln(1 - skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 582811643/1000000000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (-36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (-11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (-7451/35 + skoXC1 * 1487/280))))))) <
     skoXC1 * -3/2 *
     (6435 +
      skoXC1 *
      (-27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (-44352 +
         skoXC1 *
         (23100 +
          skoXC1 *
          (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))) |
     skoXC1 * -3/2 <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (-36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (-11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (-7451/35 + skoXC1 * 1487/280))))))) /
     (6435 +
      skoXC1 *
      (-27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (-44352 +
         skoXC1 *
         (23100 +
          skoXC1 *
          (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))) |
     6435 +
     skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 *
         (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain, (3 * skoXC1 / 2 < abs(ln(1 - skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain, (skoXC1 * 3/2 < abs(ln(1 + skoXC1 * -1))),
    inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (0 <= ln(1 + skoXC1 * -1) |
     abs(ln(1 + skoXC1 * -1)) = -ln(1 + skoXC1 * -1)),
    inference(subst, [], [abs_negative])).

cnf(refute_0_4, plain,
    (skoXC1 * 3/2 < -ln(1 + skoXC1 * -1) |
     abs(ln(1 + skoXC1 * -1)) != -ln(1 + skoXC1 * -1) |
     abs(ln(1 + skoXC1 * -1)) <= skoXC1 * 3/2),
    introduced(tautology, [equality])).

cnf(refute_0_5, plain,
    (skoXC1 * 3/2 < -ln(1 + skoXC1 * -1) | 0 <= ln(1 + skoXC1 * -1) |
     abs(ln(1 + skoXC1 * -1)) <= skoXC1 * 3/2),
    inference(resolve, [], [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 * 3/2 < -ln(1 + skoXC1 * -1) | 0 <= ln(1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_5, refute_0_2])).

cnf(refute_0_7, plain,
    (ln(1 + skoXC1 * -1) < skoXC1 * -3/2 | 1 <= skoXC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain, (skoXC1 <= 582811643/1000000000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_9, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_10, plain, (skoXC1 < 1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (ln(1 + skoXC1 * -1) < skoXC1 * -3/2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_7, refute_0_10])).

cnf(refute_0_12, plain, (ln(1 + skoXC1 * -1) < skoXC1 * -3/2),
    inference(resolve, [], [refute_0_11, refute_0_9])).

cnf(refute_0_13, plain,
    (~ lgen(0, X_000141, ln(X_000140)) | X_000141 <= ln(X_000140)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_14, plain,
    (~ lgen(0, X_000141,
            1/280 *
            (35 + 9411 * X_000140 + 141123 * X_000140 ^ 2 +
             547235 * X_000140 ^ 3 + 718735 * X_000140 ^ 4 +
             334575 * X_000140 ^ 5 + 49199 * X_000140 ^ 6 +
             1487 * X_000140 ^ 7) * (X_000140 - 1) /
            (X_000140 *
             (8 + 196 * X_000140 + 1176 * X_000140 ^ 2 +
              2450 * X_000140 ^ 3 + 1960 * X_000140 ^ 4 +
              588 * X_000140 ^ 5 + 56 * X_000140 ^ 6 + X_000140 ^ 7))) |
     X_000140 <= 0 | lgen(0, X_000141, ln(X_000140))),
    inference(subst, [], [ln_lower_bound_cf15])).

cnf(refute_0_15, plain,
    (~ lgen(0, X_000141,
            1/280 *
            (35 + 9411 * X_000140 + 141123 * X_000140 ^ 2 +
             547235 * X_000140 ^ 3 + 718735 * X_000140 ^ 4 +
             334575 * X_000140 ^ 5 + 49199 * X_000140 ^ 6 +
             1487 * X_000140 ^ 7) * (X_000140 - 1) /
            (X_000140 *
             (8 + 196 * X_000140 + 1176 * X_000140 ^ 2 +
              2450 * X_000140 ^ 3 + 1960 * X_000140 ^ 4 +
              588 * X_000140 ^ 5 + 56 * X_000140 ^ 6 + X_000140 ^ 7))) |
     X_000140 <= 0 | X_000141 <= ln(X_000140)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    ((-1/8 +
      X_000140 *
      (-1172/35 +
       X_000140 *
       (-2352/5 +
        X_000140 *
        (-7252/5 +
         X_000140 *
         (-1225/2 +
          X_000140 *
          (1372 +
           X_000140 *
           (5096/5 + X_000140 * (852/5 + X_000140 * 1487/280)))))))) /
     (X_000140 *
      (8 +
       X_000140 *
       (196 +
        X_000140 *
        (1176 +
         X_000140 *
         (2450 +
          X_000140 *
          (1960 + X_000140 * (588 + X_000140 * (56 + X_000140)))))))) <
     X_000141 | X_000140 <= 0 | X_000141 <= ln(X_000140)),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((-1/8 +
      (1 + skoXC1 * -1) *
      (-1172/35 +
       (1 + skoXC1 * -1) *
       (-2352/5 +
        (1 + skoXC1 * -1) *
        (-7252/5 +
         (1 + skoXC1 * -1) *
         (-1225/2 +
          (1 + skoXC1 * -1) *
          (1372 +
           (1 + skoXC1 * -1) *
           (5096/5 +
            (1 + skoXC1 * -1) *
            (852/5 + (1 + skoXC1 * -1) * 1487/280)))))))) /
     ((1 + skoXC1 * -1) *
      (8 +
       (1 + skoXC1 * -1) *
       (196 +
        (1 + skoXC1 * -1) *
        (1176 +
         (1 + skoXC1 * -1) *
         (2450 +
          (1 + skoXC1 * -1) *
          (1960 +
           (1 + skoXC1 * -1) *
           (588 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1))))))))) <
     skoXC1 * -3/2 | skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0), inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((-1/8 +
      (1 + skoXC1 * -1) *
      (-1172/35 +
       (1 + skoXC1 * -1) *
       (-2352/5 +
        (1 + skoXC1 * -1) *
        (-7252/5 +
         (1 + skoXC1 * -1) *
         (-1225/2 +
          (1 + skoXC1 * -1) *
          (1372 +
           (1 + skoXC1 * -1) *
           (5096/5 +
            (1 + skoXC1 * -1) *
            (852/5 + (1 + skoXC1 * -1) * 1487/280)))))))) /
     ((1 + skoXC1 * -1) *
      (8 +
       (1 + skoXC1 * -1) *
       (196 +
        (1 + skoXC1 * -1) *
        (1176 +
         (1 + skoXC1 * -1) *
         (2450 +
          (1 + skoXC1 * -1) *
          (1960 +
           (1 + skoXC1 * -1) *
           (588 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1))))))))) <
     skoXC1 * -3/2 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_17, refute_0_12])).

cnf(refute_0_19, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (-36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (-11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (-7451/35 + skoXC1 * 1487/280))))))) /
     (6435 +
      skoXC1 *
      (-27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (-44352 +
         skoXC1 *
         (23100 +
          skoXC1 *
          (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 < 1),
    inference(decision, [], [refute_0_8, refute_0_9])).

cnf(refute_0_21, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (-36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (-11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (-7451/35 + skoXC1 * 1487/280))))))) /
     (6435 +
      skoXC1 *
      (-27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (-44352 +
         skoXC1 *
         (23100 +
          skoXC1 *
          (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (48477/2 +
       skoXC1 *
       (-36465 +
        skoXC1 *
        (111485/4 +
         skoXC1 *
         (-11363 +
          skoXC1 *
          (23607/10 + skoXC1 * (-7451/35 + skoXC1 * 1487/280))))))) <
     skoXC1 * -3/2 *
     (6435 +
      skoXC1 *
      (-27456 +
       skoXC1 *
       (48048 +
        skoXC1 *
        (-44352 +
         skoXC1 *
         (23100 +
          skoXC1 *
          (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))) |
     6435 +
     skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 *
         (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_21])).

cnf(refute_0_23, plain,
    (0 <
     skoXC1 *
     (-6435/2 +
      skoXC1 *
      (33891/2 +
       skoXC1 *
       (-35607 +
        skoXC1 *
        (154627/4 +
         skoXC1 *
         (-23287 +
          skoXC1 *
          (77193/10 +
           skoXC1 *
           (-45469/35 + skoXC1 * (25393/280 + skoXC1 * -3/2)))))))) |
     skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 *
         (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1))))))) <= -6435),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 *
     (-1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (-7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (-6126/5 + skoXC1 * (1501/10 + skoXC1 * -353/70)))))) <
     skoXC1 * -3/2 *
     (1716 +
      skoXC1 *
      (-6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (-7350 +
         skoXC1 *
         (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1))))))) |
     skoXC1 * -3/2 <=
     skoXC1 *
     (-1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (-7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (-6126/5 + skoXC1 * (1501/10 + skoXC1 * -353/70)))))) /
     (1716 +
      skoXC1 *
      (-6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (-7350 +
         skoXC1 *
         (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1))))))) |
     1716 +
     skoXC1 *
     (-6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (-7350 +
        skoXC1 *
        (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_25, plain,
    (~ lgen(0, X_000127, ln(X_000126)) | X_000127 <= ln(X_000126)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_26, plain,
    (~ lgen(0, X_000127,
            1/70 *
            (10 + 1921 * X_000126 + 20149 * X_000126 ^ 2 +
             50774 * X_000126 ^ 3 + 38524 * X_000126 ^ 4 +
             8389 * X_000126 ^ 5 + 353 * X_000126 ^ 6) * (X_000126 - 1) /
            (X_000126 *
             (7 + 126 * X_000126 + 525 * X_000126 ^ 2 +
              700 * X_000126 ^ 3 + 315 * X_000126 ^ 4 + 42 * X_000126 ^ 5 +
              X_000126 ^ 6))) | X_000126 <= 0 |
     lgen(0, X_000127, ln(X_000126))),
    inference(subst, [], [ln_lower_bound_cf13])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000127,
            1/70 *
            (10 + 1921 * X_000126 + 20149 * X_000126 ^ 2 +
             50774 * X_000126 ^ 3 + 38524 * X_000126 ^ 4 +
             8389 * X_000126 ^ 5 + 353 * X_000126 ^ 6) * (X_000126 - 1) /
            (X_000126 *
             (7 + 126 * X_000126 + 525 * X_000126 ^ 2 +
              700 * X_000126 ^ 3 + 315 * X_000126 ^ 4 + 42 * X_000126 ^ 5 +
              X_000126 ^ 6))) | X_000126 <= 0 | X_000127 <= ln(X_000126)),
    inference(resolve, [], [refute_0_26, refute_0_25])).

cnf(refute_0_28, plain,
    ((-1/7 +
      X_000126 *
      (-273/10 +
       X_000126 *
       (-1302/5 +
        X_000126 *
        (-875/2 +
         X_000126 *
         (175 +
          X_000126 *
          (861/2 + X_000126 * (574/5 + X_000126 * 353/70))))))) /
     (X_000126 *
      (7 +
       X_000126 *
       (126 +
        X_000126 *
        (525 +
         X_000126 *
         (700 + X_000126 * (315 + X_000126 * (42 + X_000126))))))) <
     X_000127 | X_000126 <= 0 | X_000127 <= ln(X_000126)),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    ((-1/7 +
      (1 + skoXC1 * -1) *
      (-273/10 +
       (1 + skoXC1 * -1) *
       (-1302/5 +
        (1 + skoXC1 * -1) *
        (-875/2 +
         (1 + skoXC1 * -1) *
         (175 +
          (1 + skoXC1 * -1) *
          (861/2 +
           (1 + skoXC1 * -1) * (574/5 + (1 + skoXC1 * -1) * 353/70))))))) /
     ((1 + skoXC1 * -1) *
      (7 +
       (1 + skoXC1 * -1) *
       (126 +
        (1 + skoXC1 * -1) *
        (525 +
         (1 + skoXC1 * -1) *
         (700 +
          (1 + skoXC1 * -1) *
          (315 + (1 + skoXC1 * -1) * (42 + (1 + skoXC1 * -1)))))))) <
     skoXC1 * -3/2 | skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0), inference(subst, [], [refute_0_28])).

cnf(refute_0_30, plain,
    ((-1/7 +
      (1 + skoXC1 * -1) *
      (-273/10 +
       (1 + skoXC1 * -1) *
       (-1302/5 +
        (1 + skoXC1 * -1) *
        (-875/2 +
         (1 + skoXC1 * -1) *
         (175 +
          (1 + skoXC1 * -1) *
          (861/2 +
           (1 + skoXC1 * -1) * (574/5 + (1 + skoXC1 * -1) * 353/70))))))) /
     ((1 + skoXC1 * -1) *
      (7 +
       (1 + skoXC1 * -1) *
       (126 +
        (1 + skoXC1 * -1) *
        (525 +
         (1 + skoXC1 * -1) *
         (700 +
          (1 + skoXC1 * -1) *
          (315 + (1 + skoXC1 * -1) * (42 + (1 + skoXC1 * -1)))))))) <
     skoXC1 * -3/2 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_29, refute_0_12])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (-1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (-7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (-6126/5 + skoXC1 * (1501/10 + skoXC1 * -353/70)))))) /
     (1716 +
      skoXC1 *
      (-6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (-7350 +
         skoXC1 *
         (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1))))))) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (skoXC1 *
     (-1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (-7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (-6126/5 + skoXC1 * (1501/10 + skoXC1 * -353/70)))))) /
     (1716 +
      skoXC1 *
      (-6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (-7350 +
         skoXC1 *
         (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1))))))) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_31, refute_0_20])).

cnf(refute_0_33, plain,
    (skoXC1 *
     (-1716 +
      skoXC1 *
      (5610 +
       skoXC1 *
       (-7040 +
        skoXC1 *
        (4226 +
         skoXC1 * (-6126/5 + skoXC1 * (1501/10 + skoXC1 * -353/70)))))) <
     skoXC1 * -3/2 *
     (1716 +
      skoXC1 *
      (-6468 +
       skoXC1 *
       (9702 +
        skoXC1 *
        (-7350 +
         skoXC1 *
         (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1))))))) |
     1716 +
     skoXC1 *
     (-6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (-7350 +
        skoXC1 *
        (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1)))))) <= 0),
    inference(resolve, [], [refute_0_24, refute_0_32])).

cnf(refute_0_34, plain,
    (0 <
     skoXC1 *
     (-858 +
      skoXC1 *
      (4092 +
       skoXC1 *
       (-7513 +
        skoXC1 *
        (6799 +
         skoXC1 *
         (-15924/5 +
          skoXC1 * (7319/10 + skoXC1 * (-2396/35 + skoXC1 * 3/2))))))) |
     skoXC1 *
     (-6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (-7350 +
        skoXC1 *
        (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1)))))) <=
     -1716), inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 *
     (-462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (-1288 + skoXC1 * (1127/2 + skoXC1 * (-497/5 + skoXC1 * 71/15))))) <
     skoXC1 * -3/2 *
     (462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1)))))) |
     skoXC1 * -3/2 <=
     skoXC1 *
     (-462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (-1288 + skoXC1 * (1127/2 + skoXC1 * (-497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (-1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000113, ln(X_000112)) | X_000113 <= ln(X_000112)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000113,
            1/30 *
            (142 * X_000112 ^ 5 + 2272 * X_000112 ^ 4 +
             6397 * X_000112 ^ 3 + 4397 * X_000112 ^ 2 + 647 * X_000112 +
             5) * (X_000112 - 1) /
            (X_000112 *
             (X_000112 ^ 5 + 30 * X_000112 ^ 4 + 150 * X_000112 ^ 3 +
              200 * X_000112 ^ 2 + 75 * X_000112 + 6))) | X_000112 <= 0 |
     lgen(0, X_000113, ln(X_000112))),
    inference(subst, [], [ln_lower_bound_cf11])).

cnf(refute_0_38, plain,
    (~ lgen(0, X_000113,
            1/30 *
            (142 * X_000112 ^ 5 + 2272 * X_000112 ^ 4 +
             6397 * X_000112 ^ 3 + 4397 * X_000112 ^ 2 + 647 * X_000112 +
             5) * (X_000112 - 1) /
            (X_000112 *
             (X_000112 ^ 5 + 30 * X_000112 ^ 4 + 150 * X_000112 ^ 3 +
              200 * X_000112 ^ 2 + 75 * X_000112 + 6))) | X_000112 <= 0 |
     X_000113 <= ln(X_000112)),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    ((-1/6 +
      X_000112 *
      (-107/5 +
       X_000112 *
       (-125 +
        X_000112 *
        (-200/3 +
         X_000112 * (275/2 + X_000112 * (71 + X_000112 * 71/15)))))) /
     (X_000112 *
      (6 +
       X_000112 *
       (75 +
        X_000112 *
        (200 + X_000112 * (150 + X_000112 * (30 + X_000112)))))) <
     X_000113 | X_000112 <= 0 | X_000113 <= ln(X_000112)),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    ((-1/6 +
      (1 + skoXC1 * -1) *
      (-107/5 +
       (1 + skoXC1 * -1) *
       (-125 +
        (1 + skoXC1 * -1) *
        (-200/3 +
         (1 + skoXC1 * -1) *
         (275/2 +
          (1 + skoXC1 * -1) * (71 + (1 + skoXC1 * -1) * 71/15)))))) /
     ((1 + skoXC1 * -1) *
      (6 +
       (1 + skoXC1 * -1) *
       (75 +
        (1 + skoXC1 * -1) *
        (200 +
         (1 + skoXC1 * -1) *
         (150 + (1 + skoXC1 * -1) * (30 + (1 + skoXC1 * -1))))))) <
     skoXC1 * -3/2 | skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0), inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    ((-1/6 +
      (1 + skoXC1 * -1) *
      (-107/5 +
       (1 + skoXC1 * -1) *
       (-125 +
        (1 + skoXC1 * -1) *
        (-200/3 +
         (1 + skoXC1 * -1) *
         (275/2 +
          (1 + skoXC1 * -1) * (71 + (1 + skoXC1 * -1) * 71/15)))))) /
     ((1 + skoXC1 * -1) *
      (6 +
       (1 + skoXC1 * -1) *
       (75 +
        (1 + skoXC1 * -1) *
        (200 +
         (1 + skoXC1 * -1) *
         (150 + (1 + skoXC1 * -1) * (30 + (1 + skoXC1 * -1))))))) <
     skoXC1 * -3/2 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_40, refute_0_12])).

cnf(refute_0_42, plain,
    (skoXC1 *
     (-462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (-1288 + skoXC1 * (1127/2 + skoXC1 * (-497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1)))))) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (skoXC1 *
     (-462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (-1288 + skoXC1 * (1127/2 + skoXC1 * (-497/5 + skoXC1 * 71/15))))) /
     (462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1)))))) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_42, refute_0_20])).

cnf(refute_0_44, plain,
    (skoXC1 *
     (-462 +
      skoXC1 *
      (1281 +
       skoXC1 *
       (-1288 + skoXC1 * (1127/2 + skoXC1 * (-497/5 + skoXC1 * 71/15))))) <
     skoXC1 * -3/2 *
     (462 +
      skoXC1 *
      (-1512 +
       skoXC1 *
       (1890 +
        skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1)))))) |
     462 +
     skoXC1 *
     (-1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_35, refute_0_43])).

cnf(refute_0_45, plain,
    (0 <
     skoXC1 *
     (-231 +
      skoXC1 *
      (987 +
       skoXC1 *
       (-1547 +
        skoXC1 *
        (2233/2 +
         skoXC1 * (-3731/10 + skoXC1 * (739/15 + skoXC1 * -3/2)))))) |
     skoXC1 *
     (-1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1))))) <=
     -462), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (skoXC1 *
     (-126 +
      skoXC1 *
      (287 + skoXC1 * (-217 + skoXC1 * (361/6 + skoXC1 * -131/30)))) <
     skoXC1 * -3/2 *
     (126 +
      skoXC1 *
      (-350 +
       skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1))))) |
     skoXC1 * -3/2 <=
     skoXC1 *
     (-126 +
      skoXC1 *
      (287 + skoXC1 * (-217 + skoXC1 * (361/6 + skoXC1 * -131/30)))) /
     (126 +
      skoXC1 *
      (-350 +
       skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1))))) |
     126 +
     skoXC1 *
     (-350 +
      skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_47, plain,
    (~ lgen(0, X_000099, ln(X_000098)) | X_000099 <= ln(X_000098)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_48, plain,
    (~ lgen(0, X_000099,
            1/30 *
            (6 + 481 * X_000098 + 1881 * X_000098 ^ 2 +
             1281 * X_000098 ^ 3 + 131 * X_000098 ^ 4) * (X_000098 - 1) /
            (X_000098 *
             (5 + 40 * X_000098 + 60 * X_000098 ^ 2 + 20 * X_000098 ^ 3 +
              X_000098 ^ 4))) | X_000098 <= 0 |
     lgen(0, X_000099, ln(X_000098))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_49, plain,
    (~ lgen(0, X_000099,
            1/30 *
            (6 + 481 * X_000098 + 1881 * X_000098 ^ 2 +
             1281 * X_000098 ^ 3 + 131 * X_000098 ^ 4) * (X_000098 - 1) /
            (X_000098 *
             (5 + 40 * X_000098 + 60 * X_000098 ^ 2 + 20 * X_000098 ^ 3 +
              X_000098 ^ 4))) | X_000098 <= 0 | X_000099 <= ln(X_000098)),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    ((-1/5 +
      X_000098 *
      (-95/6 +
       X_000098 *
       (-140/3 +
        X_000098 * (20 + X_000098 * (115/3 + X_000098 * 131/30))))) /
     (X_000098 *
      (5 +
       X_000098 * (40 + X_000098 * (60 + X_000098 * (20 + X_000098))))) <
     X_000099 | X_000098 <= 0 | X_000099 <= ln(X_000098)),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    ((-1/5 +
      (1 + skoXC1 * -1) *
      (-95/6 +
       (1 + skoXC1 * -1) *
       (-140/3 +
        (1 + skoXC1 * -1) *
        (20 +
         (1 + skoXC1 * -1) * (115/3 + (1 + skoXC1 * -1) * 131/30))))) /
     ((1 + skoXC1 * -1) *
      (5 +
       (1 + skoXC1 * -1) *
       (40 +
        (1 + skoXC1 * -1) *
        (60 + (1 + skoXC1 * -1) * (20 + (1 + skoXC1 * -1)))))) <
     skoXC1 * -3/2 | skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0), inference(subst, [], [refute_0_50])).

cnf(refute_0_52, plain,
    ((-1/5 +
      (1 + skoXC1 * -1) *
      (-95/6 +
       (1 + skoXC1 * -1) *
       (-140/3 +
        (1 + skoXC1 * -1) *
        (20 +
         (1 + skoXC1 * -1) * (115/3 + (1 + skoXC1 * -1) * 131/30))))) /
     ((1 + skoXC1 * -1) *
      (5 +
       (1 + skoXC1 * -1) *
       (40 +
        (1 + skoXC1 * -1) *
        (60 + (1 + skoXC1 * -1) * (20 + (1 + skoXC1 * -1)))))) <
     skoXC1 * -3/2 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_51, refute_0_12])).

cnf(refute_0_53, plain,
    (skoXC1 *
     (-126 +
      skoXC1 *
      (287 + skoXC1 * (-217 + skoXC1 * (361/6 + skoXC1 * -131/30)))) /
     (126 +
      skoXC1 *
      (-350 +
       skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1))))) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 *
     (-126 +
      skoXC1 *
      (287 + skoXC1 * (-217 + skoXC1 * (361/6 + skoXC1 * -131/30)))) /
     (126 +
      skoXC1 *
      (-350 +
       skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1))))) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_53, refute_0_20])).

cnf(refute_0_55, plain,
    (skoXC1 *
     (-126 +
      skoXC1 *
      (287 + skoXC1 * (-217 + skoXC1 * (361/6 + skoXC1 * -131/30)))) <
     skoXC1 * -3/2 *
     (126 +
      skoXC1 *
      (-350 +
       skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1))))) |
     126 +
     skoXC1 *
     (-350 +
      skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1)))) <=
     0), inference(resolve, [], [refute_0_46, refute_0_54])).

cnf(refute_0_56, plain,
    (0 <
     skoXC1 *
     (-63 +
      skoXC1 *
      (238 +
       skoXC1 *
       (-308 + skoXC1 * (989/6 + skoXC1 * (-497/15 + skoXC1 * 3/2))))) |
     skoXC1 *
     (-350 +
      skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1)))) <=
     -126), inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 *
     (-35 + skoXC1 * (125/2 + skoXC1 * (-95/3 + skoXC1 * 47/12))) <
     skoXC1 * -3/2 *
     (35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1)))) |
     skoXC1 * -3/2 <=
     skoXC1 *
     (-35 + skoXC1 * (125/2 + skoXC1 * (-95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1)))) |
     35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_58, plain,
    (~ lgen(0, X_000085, ln(X_000084)) | X_000085 <= ln(X_000084)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_59, plain,
    (~ lgen(0, X_000085,
            1/12 *
            (47 * X_000084 ^ 3 + 239 * X_000084 ^ 2 + 131 * X_000084 + 3) *
            (X_000084 - 1) /
            (X_000084 *
             (X_000084 ^ 3 + 12 * X_000084 ^ 2 + 18 * X_000084 + 4))) |
     X_000084 <= 0 | lgen(0, X_000085, ln(X_000084))),
    inference(subst, [], [ln_lower_bound_cf7])).

cnf(refute_0_60, plain,
    (~ lgen(0, X_000085,
            1/12 *
            (47 * X_000084 ^ 3 + 239 * X_000084 ^ 2 + 131 * X_000084 + 3) *
            (X_000084 - 1) /
            (X_000084 *
             (X_000084 ^ 3 + 12 * X_000084 ^ 2 + 18 * X_000084 + 4))) |
     X_000084 <= 0 | X_000085 <= ln(X_000084)),
    inference(resolve, [], [refute_0_59, refute_0_58])).

cnf(refute_0_61, plain,
    ((-1/4 +
      X_000084 *
      (-32/3 + X_000084 * (-9 + X_000084 * (16 + X_000084 * 47/12)))) /
     (X_000084 * (4 + X_000084 * (18 + X_000084 * (12 + X_000084)))) <
     X_000085 | X_000084 <= 0 | X_000085 <= ln(X_000084)),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    ((-1/4 +
      (1 + skoXC1 * -1) *
      (-32/3 +
       (1 + skoXC1 * -1) *
       (-9 + (1 + skoXC1 * -1) * (16 + (1 + skoXC1 * -1) * 47/12)))) /
     ((1 + skoXC1 * -1) *
      (4 +
       (1 + skoXC1 * -1) *
       (18 + (1 + skoXC1 * -1) * (12 + (1 + skoXC1 * -1))))) <
     skoXC1 * -3/2 | skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0), inference(subst, [], [refute_0_61])).

cnf(refute_0_63, plain,
    ((-1/4 +
      (1 + skoXC1 * -1) *
      (-32/3 +
       (1 + skoXC1 * -1) *
       (-9 + (1 + skoXC1 * -1) * (16 + (1 + skoXC1 * -1) * 47/12)))) /
     ((1 + skoXC1 * -1) *
      (4 +
       (1 + skoXC1 * -1) *
       (18 + (1 + skoXC1 * -1) * (12 + (1 + skoXC1 * -1))))) <
     skoXC1 * -3/2 | 1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_62, refute_0_12])).

cnf(refute_0_64, plain,
    (skoXC1 *
     (-35 + skoXC1 * (125/2 + skoXC1 * (-95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1)))) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 *
     (-35 + skoXC1 * (125/2 + skoXC1 * (-95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1)))) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_64, refute_0_20])).

cnf(refute_0_66, plain,
    (skoXC1 *
     (-35 + skoXC1 * (125/2 + skoXC1 * (-95/3 + skoXC1 * 47/12))) <
     skoXC1 * -3/2 *
     (35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1)))) |
     35 + skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1))) <= 0),
    inference(resolve, [], [refute_0_57, refute_0_65])).

cnf(refute_0_67, plain,
    (0 <
     skoXC1 *
     (-35/2 +
      skoXC1 *
      (115/2 + skoXC1 * (-175/3 + skoXC1 * (241/12 + skoXC1 * -3/2)))) |
     skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1))) <= -35),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * (-10 + skoXC1 * (13 + skoXC1 * -10/3)) <
     skoXC1 * -3/2 * (10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1))) |
     skoXC1 * -3/2 <=
     skoXC1 * (-10 + skoXC1 * (13 + skoXC1 * -10/3)) /
     (10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1))) |
     10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_69, plain,
    (~ lgen(0, X_000071, ln(X_000070)) | X_000071 <= ln(X_000070)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_70, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | lgen(0, X_000071, ln(X_000070))),
    inference(subst, [], [ln_lower_bound_cf5])).

cnf(refute_0_71, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(resolve, [], [refute_0_70, refute_0_69])).

cnf(refute_0_72, plain,
    ((-1/3 + X_000070 * (-6 + X_000070 * (3 + X_000070 * 10/3))) /
     (X_000070 * (3 + X_000070 * (6 + X_000070))) < X_000071 |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    ((-1/3 +
      (1 + skoXC1 * -1) *
      (-6 + (1 + skoXC1 * -1) * (3 + (1 + skoXC1 * -1) * 10/3))) /
     ((1 + skoXC1 * -1) *
      (3 + (1 + skoXC1 * -1) * (6 + (1 + skoXC1 * -1)))) < skoXC1 * -3/2 |
     skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_74, plain,
    ((-1/3 +
      (1 + skoXC1 * -1) *
      (-6 + (1 + skoXC1 * -1) * (3 + (1 + skoXC1 * -1) * 10/3))) /
     ((1 + skoXC1 * -1) *
      (3 + (1 + skoXC1 * -1) * (6 + (1 + skoXC1 * -1)))) < skoXC1 * -3/2 |
     1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_73, refute_0_12])).

cnf(refute_0_75, plain,
    (skoXC1 * (-10 + skoXC1 * (13 + skoXC1 * -10/3)) /
     (10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1))) < skoXC1 * -3/2 |
     1 <= skoXC1), inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * (-10 + skoXC1 * (13 + skoXC1 * -10/3)) /
     (10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1))) < skoXC1 * -3/2),
    inference(resolve, [], [refute_0_75, refute_0_20])).

cnf(refute_0_77, plain,
    (skoXC1 * (-10 + skoXC1 * (13 + skoXC1 * -10/3)) <
     skoXC1 * -3/2 * (10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1))) |
     10 + skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_68, refute_0_76])).

cnf(refute_0_78, plain,
    (0 < skoXC1 * (-5 + skoXC1 * (14 + skoXC1 * (-61/6 + skoXC1 * 3/2))) |
     skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1)) <= -10),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 * (-3 + skoXC1 * 5/2) <
     skoXC1 * -3/2 * (3 + skoXC1 * (-4 + skoXC1)) |
     skoXC1 * -3/2 <=
     skoXC1 * (-3 + skoXC1 * 5/2) / (3 + skoXC1 * (-4 + skoXC1)) |
     3 + skoXC1 * (-4 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_80, plain,
    (~ lgen(0, X_000057, ln(X_000056)) | X_000057 <= ln(X_000056)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_81, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     lgen(0, X_000057, ln(X_000056))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_82, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     X_000057 <= ln(X_000056)),
    inference(resolve, [], [refute_0_81, refute_0_80])).

cnf(refute_0_83, plain,
    ((-1/2 + X_000056 * (-2 + X_000056 * 5/2)) /
     (X_000056 * (2 + X_000056)) < X_000057 | X_000056 <= 0 |
     X_000057 <= ln(X_000056)), inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    ((-1/2 + (1 + skoXC1 * -1) * (-2 + (1 + skoXC1 * -1) * 5/2)) /
     ((1 + skoXC1 * -1) * (2 + (1 + skoXC1 * -1))) < skoXC1 * -3/2 |
     skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0),
    inference(subst, [], [refute_0_83])).

cnf(refute_0_85, plain,
    ((-1/2 + (1 + skoXC1 * -1) * (-2 + (1 + skoXC1 * -1) * 5/2)) /
     ((1 + skoXC1 * -1) * (2 + (1 + skoXC1 * -1))) < skoXC1 * -3/2 |
     1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_84, refute_0_12])).

cnf(refute_0_86, plain,
    (skoXC1 * (-3 + skoXC1 * 5/2) / (3 + skoXC1 * (-4 + skoXC1)) <
     skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (skoXC1 * (-3 + skoXC1 * 5/2) / (3 + skoXC1 * (-4 + skoXC1)) <
     skoXC1 * -3/2), inference(resolve, [], [refute_0_86, refute_0_20])).

cnf(refute_0_88, plain,
    (skoXC1 * (-3 + skoXC1 * 5/2) <
     skoXC1 * -3/2 * (3 + skoXC1 * (-4 + skoXC1)) |
     3 + skoXC1 * (-4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_79, refute_0_87])).

cnf(refute_0_89, plain,
    (0 < skoXC1 * (-3/2 + skoXC1 * (7/2 + skoXC1 * -3/2)) |
     skoXC1 * (-4 + skoXC1) <= -3),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (skoXC1 * -1 < skoXC1 * -3/2 * (1 + skoXC1 * -1) |
     skoXC1 * -3/2 <= skoXC1 * -1 / (1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_91, plain,
    (~ lgen(0, X_000043, ln(X_000042)) | X_000043 <= ln(X_000042)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_92, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     lgen(0, X_000043, ln(X_000042))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_93, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     X_000043 <= ln(X_000042)),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    ((-1 + X_000042) / X_000042 < X_000043 | X_000042 <= 0 |
     X_000043 <= ln(X_000042)), inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    ((-1 + (1 + skoXC1 * -1)) / (1 + skoXC1 * -1) < skoXC1 * -3/2 |
     skoXC1 * -3/2 <= ln(1 + skoXC1 * -1) | 1 + skoXC1 * -1 <= 0),
    inference(subst, [], [refute_0_94])).

cnf(refute_0_96, plain,
    ((-1 + (1 + skoXC1 * -1)) / (1 + skoXC1 * -1) < skoXC1 * -3/2 |
     1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_95, refute_0_12])).

cnf(refute_0_97, plain,
    (skoXC1 * -1 / (1 + skoXC1 * -1) < skoXC1 * -3/2 | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain, (skoXC1 * -1 / (1 + skoXC1 * -1) < skoXC1 * -3/2),
    inference(resolve, [], [refute_0_97, refute_0_20])).

cnf(refute_0_99, plain,
    (skoXC1 * -1 < skoXC1 * -3/2 * (1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_90, refute_0_98])).

cnf(refute_0_100, plain,
    (0 < skoXC1 * (-1/2 + skoXC1 * 3/2) | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain, (skoXC1 < 1 | 0 < skoXC1 * (-1/2 + skoXC1 * 3/2)),
    inference(decision, [], [refute_0_8, refute_0_9])).

cnf(refute_0_102, plain, (0 < skoXC1 * (-1/2 + skoXC1 * 3/2)),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (-3 < skoXC1 * (-4 + skoXC1) |
     0 < skoXC1 * (-3/2 + skoXC1 * (7/2 + skoXC1 * -3/2))),
    inference(decision, [], [refute_0_102, refute_0_8, refute_0_9])).

cnf(refute_0_104, plain,
    (0 < skoXC1 * (-3/2 + skoXC1 * (7/2 + skoXC1 * -3/2))),
    inference(resolve, [], [refute_0_89, refute_0_103])).

cnf(refute_0_105, plain,
    (-10 < skoXC1 * (-18 + skoXC1 * (9 + skoXC1 * -1)) |
     0 < skoXC1 * (-5 + skoXC1 * (14 + skoXC1 * (-61/6 + skoXC1 * 3/2)))),
    inference(decision, [],
              [refute_0_104, refute_0_102, refute_0_8, refute_0_9])).

cnf(refute_0_106, plain,
    (0 < skoXC1 * (-5 + skoXC1 * (14 + skoXC1 * (-61/6 + skoXC1 * 3/2)))),
    inference(resolve, [], [refute_0_78, refute_0_105])).

cnf(refute_0_107, plain,
    (-35 < skoXC1 * (-80 + skoXC1 * (60 + skoXC1 * (-16 + skoXC1))) |
     0 <
     skoXC1 *
     (-35/2 +
      skoXC1 *
      (115/2 + skoXC1 * (-175/3 + skoXC1 * (241/12 + skoXC1 * -3/2))))),
    inference(decision, [],
              [refute_0_106, refute_0_104, refute_0_102, refute_0_8,
               refute_0_9])).

cnf(refute_0_108, plain,
    (0 <
     skoXC1 *
     (-35/2 +
      skoXC1 *
      (115/2 + skoXC1 * (-175/3 + skoXC1 * (241/12 + skoXC1 * -3/2))))),
    inference(resolve, [], [refute_0_67, refute_0_107])).

cnf(refute_0_109, plain,
    (-126 <
     skoXC1 *
     (-350 +
      skoXC1 * (350 + skoXC1 * (-150 + skoXC1 * (25 + skoXC1 * -1)))) |
     0 <
     skoXC1 *
     (-63 +
      skoXC1 *
      (238 +
       skoXC1 *
       (-308 + skoXC1 * (989/6 + skoXC1 * (-497/15 + skoXC1 * 3/2)))))),
    inference(decision, [],
              [refute_0_108, refute_0_106, refute_0_104, refute_0_102,
               refute_0_8, refute_0_9])).

cnf(refute_0_110, plain,
    (0 <
     skoXC1 *
     (-63 +
      skoXC1 *
      (238 +
       skoXC1 *
       (-308 + skoXC1 * (989/6 + skoXC1 * (-497/15 + skoXC1 * 3/2)))))),
    inference(resolve, [], [refute_0_56, refute_0_109])).

cnf(refute_0_111, plain,
    (-462 <
     skoXC1 *
     (-1512 +
      skoXC1 *
      (1890 +
       skoXC1 * (-1120 + skoXC1 * (315 + skoXC1 * (-36 + skoXC1))))) |
     0 <
     skoXC1 *
     (-231 +
      skoXC1 *
      (987 +
       skoXC1 *
       (-1547 +
        skoXC1 *
        (2233/2 +
         skoXC1 * (-3731/10 + skoXC1 * (739/15 + skoXC1 * -3/2))))))),
    inference(decision, [],
              [refute_0_110, refute_0_108, refute_0_106, refute_0_104,
               refute_0_102, refute_0_8, refute_0_9])).

cnf(refute_0_112, plain,
    (0 <
     skoXC1 *
     (-231 +
      skoXC1 *
      (987 +
       skoXC1 *
       (-1547 +
        skoXC1 *
        (2233/2 +
         skoXC1 * (-3731/10 + skoXC1 * (739/15 + skoXC1 * -3/2))))))),
    inference(resolve, [], [refute_0_45, refute_0_111])).

cnf(refute_0_113, plain,
    (-1716 <
     skoXC1 *
     (-6468 +
      skoXC1 *
      (9702 +
       skoXC1 *
       (-7350 +
        skoXC1 *
        (2940 + skoXC1 * (-588 + skoXC1 * (49 + skoXC1 * -1)))))) |
     0 <
     skoXC1 *
     (-858 +
      skoXC1 *
      (4092 +
       skoXC1 *
       (-7513 +
        skoXC1 *
        (6799 +
         skoXC1 *
         (-15924/5 +
          skoXC1 * (7319/10 + skoXC1 * (-2396/35 + skoXC1 * 3/2)))))))),
    inference(decision, [],
              [refute_0_112, refute_0_110, refute_0_108, refute_0_106,
               refute_0_104, refute_0_102, refute_0_8, refute_0_9])).

cnf(refute_0_114, plain,
    (0 <
     skoXC1 *
     (-858 +
      skoXC1 *
      (4092 +
       skoXC1 *
       (-7513 +
        skoXC1 *
        (6799 +
         skoXC1 *
         (-15924/5 +
          skoXC1 * (7319/10 + skoXC1 * (-2396/35 + skoXC1 * 3/2)))))))),
    inference(resolve, [], [refute_0_34, refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 *
         (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1))))))) <= -6435 |
     skoXC1 *
     (-6435/2 +
      skoXC1 *
      (33891/2 +
       skoXC1 *
       (-35607 +
        skoXC1 *
        (154627/4 +
         skoXC1 *
         (-23287 +
          skoXC1 *
          (77193/10 +
           skoXC1 *
           (-45469/35 + skoXC1 * (25393/280 + skoXC1 * -3/2)))))))) <= 0),
    inference(decision, [],
              [refute_0_114, refute_0_112, refute_0_110, refute_0_108,
               refute_0_106, refute_0_104, refute_0_102, refute_0_8,
               refute_0_9])).

cnf(refute_0_116, plain,
    (skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 *
         (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1))))))) <= -6435),
    inference(resolve, [], [refute_0_115, refute_0_23])).

cnf(refute_0_117, plain,
    (-6435 <
     skoXC1 *
     (-27456 +
      skoXC1 *
      (48048 +
       skoXC1 *
       (-44352 +
        skoXC1 *
        (23100 +
         skoXC1 * (-6720 + skoXC1 * (1008 + skoXC1 * (-64 + skoXC1)))))))),
    inference(decision, [],
              [refute_0_114, refute_0_112, refute_0_110, refute_0_108,
               refute_0_106, refute_0_104, refute_0_102, refute_0_8,
               refute_0_9])).

cnf(refute_0_118, plain, ($false),
    inference(resolve, [], [refute_0_116, refute_0_117])).
% SZS output end CNFRefutation for abs-problem-3.tptp
