% SZS output start CNFRefutation for Arthan1.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof('Rob_Arthan_problem_1', conjecture,
    (! [S] :
       ((0 <= S & S <= 3) =>
        (1 + S + S ^ 2 + sin(S)) *
        ((1 + S) * (2 - sin(S)) + 2 + 4 * S + 2 * cos(S)) -
        2 * (1 + S) * (1 + 2 * S + cos(S)) ^ 2 < 0))).

fof(subgoal_0, plain,
    (! [S] :
       ((0 <= S & S <= 3) =>
        (1 + S + S ^ 2 + sin(S)) *
        ((1 + S) * (2 - sin(S)) + 2 + 4 * S + 2 * cos(S)) -
        2 * (1 + S) * (1 + 2 * S + cos(S)) ^ 2 < 0)),
    inference(strip, [], ['Rob_Arthan_problem_1'])).

fof(negate_0_0, plain,
    (~ ! [S] :
         ((0 <= S & S <= 3) =>
          (1 + S + S ^ 2 + sin(S)) *
          ((1 + S) * (2 - sin(S)) + 2 + 4 * S + 2 * cos(S)) -
          2 * (1 + S) * (1 + 2 * S + cos(S)) ^ 2 < 0)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [S] :
       (S <= 3 & 0 <= S &
        0 <=
        (1 + S + S ^ 2 + sin(S)) *
        ((1 + S) * (2 - sin(S)) + 2 + 4 * S + 2 * cos(S)) -
        2 * (1 + S) * (1 + 2 * S + cos(S)) ^ 2)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <=
     (1 + skoSC1 + skoSC1 ^ 2 + sin(skoSC1)) *
     ((1 + skoSC1) * (2 - sin(skoSC1)) + 2 + 4 * skoSC1 +
      2 * cos(skoSC1)) -
     2 * (1 + skoSC1) * (1 + 2 * skoSC1 + cos(skoSC1)) ^ 2 & 0 <= skoSC1 &
     skoSC1 <= 3), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (0 <=
     (1 + skoSC1 + skoSC1 ^ 2 + sin(skoSC1)) *
     ((1 + skoSC1) * (2 - sin(skoSC1)) + 2 + 4 * skoSC1 +
      2 * cos(skoSC1)) -
     2 * (1 + skoSC1) * (1 + 2 * skoSC1 + cos(skoSC1)) ^ 2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoSC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoSC1 <= 3),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (cos(X_000016) < X_000017 | ~ lgen(1, cos(X_000016), X_000017)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain,
    (~ lgen(1, 1 - X_000016 ^ 2 / 2 + X_000016 ^ 4 / 24, X_000017) |
     lgen(1, cos(X_000016), X_000017)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_3, plain,
    (cos(X_000016) < X_000017 |
     ~ lgen(1, 1 - X_000016 ^ 2 / 2 + X_000016 ^ 4 / 24, X_000017)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (cos(X_000016) < X_000017 |
     X_000017 <=
     1 + X_000016 * (X_000016 * (-1/2 + X_000016 * (X_000016 * 1/24)))),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (cos(skoSC1) <
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) |
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) |
     2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_7, plain,
    (X_000024 < cos(X_000023) | ~ lgen(1, X_000024, cos(X_000023))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_8, plain,
    (~ lgen(1, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800) |
     lgen(1, X_000024, cos(X_000023))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_9, plain,
    (X_000024 < cos(X_000023) |
     ~ lgen(1, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800)),
    inference(resolve, [], [refute_0_8, refute_0_7])).

cnf(refute_0_10, plain,
    (X_000024 < cos(X_000023) |
     1 +
     X_000023 *
     (X_000023 *
      (-1/2 +
       X_000023 *
       (X_000023 *
        (1/24 +
         X_000023 *
         (X_000023 *
          (-1/720 +
           X_000023 *
           (X_000023 *
            (1/40320 + X_000023 * (X_000023 * -1/3628800))))))))) <=
     X_000024), inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     cos(skoSC1) |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 <=
     (1 + skoSC1 + skoSC1 ^ 2 + sin(skoSC1)) *
     ((1 + skoSC1) * (2 - sin(skoSC1)) + 2 + 4 * skoSC1 +
      2 * cos(skoSC1)) -
     2 * (1 + skoSC1) * (1 + 2 * skoSC1 + cos(skoSC1)) ^ 2),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_13, plain,
    (sin(skoSC1) *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2))),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) < 0 |
     0 <
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) |
     0 <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2))),
    inference(split, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 <
     sin(skoSC1) * (-1 + skoSC1 * -1) |
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 <
     sin(skoSC1) * (1 + skoSC1) |
     cos(skoSC1) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000))))))))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_17, plain,
    (cos(skoSC1) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_18, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) |
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 +
         skoSC1 *
         (83/120 +
          skoSC1 *
          (59/360 +
           skoSC1 *
           (-269/5040 +
            skoSC1 *
            (-19/2520 +
             skoSC1 *
             (241/120960 +
              skoSC1 *
              (1/5760 +
               skoSC1 *
               (-551/13305600 +
                skoSC1 *
                (-71/29937600 +
                 skoSC1 *
                 (1/1871100 +
                  skoSC1 *
                  (1/47900160 + skoSC1 * -1/239500800)))))))))))))) |
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_19, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) < cos(skoSC1) |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_20, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800)))))))))))) <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800)))))))))))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 +
        skoSC1 *
        (-5/3 +
         skoSC1 *
         (43/120 +
          skoSC1 *
          (1/5 +
           skoSC1 *
           (-29/1680 +
            skoSC1 *
            (-23/2520 +
             skoSC1 *
             (7/17280 +
              skoSC1 *
              (19/90720 +
               skoSC1 *
               (-223/39916800 +
                skoSC1 *
                (-19/6652800 +
                 skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
     cos(skoSC1) *
     (-2 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-6 +
        skoSC1 *
        (-1/3 +
         skoSC1 *
         (skoSC1 *
          (1/60 +
           skoSC1 *
           (skoSC1 *
            (-1/2520 +
             skoSC1 *
             (skoSC1 *
              (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     skoSC1 *
     (1 +
      skoSC1 *
      (1 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (1/120 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (-1/5040 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (1/362880 +
              skoSC1 *
              (1/362880 +
               skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_21, plain,
    (X_000089 < sin(X_000088) | ~ lgen(1, X_000089, sin(X_000088))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_22, plain,
    (X_000088 < 0 | ~ lgen(1, X_000089, X_000088 - X_000088 ^ 3 / 6) |
     lgen(1, X_000089, sin(X_000088))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_23, plain,
    (X_000088 < 0 | X_000089 < sin(X_000088) |
     ~ lgen(1, X_000089, X_000088 - X_000088 ^ 3 / 6)),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (X_000088 < 0 | X_000089 < sin(X_000088) |
     X_000088 * (1 + X_000088 * (X_000088 * -1/6)) <= X_000089),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (skoSC1 < 0 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800)))))))))))) <
     sin(skoSC1) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))))),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <
     skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) |
     skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_27, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) <= 0 |
     sin(skoSC1) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))), inference(split, [], [refute_0_13])).

cnf(refute_0_28, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     sin(skoSC1) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (X_000121 < sin(X_000120) | ~ lgen(1, X_000121, sin(X_000120))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_30, plain,
    (X_000120 < 0 |
     ~ lgen(1, X_000121,
            X_000120 - X_000120 ^ 3 / 6 + X_000120 ^ 5 / 120 -
            X_000120 ^ 7 / 5040 + X_000120 ^ 9 / 362880 -
            X_000120 ^ 11 / 39916800) | lgen(1, X_000121, sin(X_000120))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_31, plain,
    (X_000120 < 0 | X_000121 < sin(X_000120) |
     ~ lgen(1, X_000121,
            X_000120 - X_000120 ^ 3 / 6 + X_000120 ^ 5 / 120 -
            X_000120 ^ 7 / 5040 + X_000120 ^ 9 / 362880 -
            X_000120 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_30, refute_0_29])).

cnf(refute_0_32, plain,
    (X_000120 < 0 | X_000121 < sin(X_000120) |
     X_000120 *
     (1 +
      X_000120 *
      (X_000120 *
       (-1/6 +
        X_000120 *
        (X_000120 *
         (1/120 +
          X_000120 *
          (X_000120 *
           (-1/5040 +
            X_000120 *
            (X_000120 *
             (1/362880 + X_000120 * (X_000120 * -1/39916800)))))))))) <=
     X_000121), inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (skoSC1 < 0 |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) < sin(skoSC1) |
     skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))), inference(subst, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))),
    inference(resolve, [], [refute_0_28, refute_0_33])).

cnf(refute_0_35, plain, (0 <= skoSC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_36, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))),
    inference(resolve, [], [refute_0_35, refute_0_34])).

cnf(refute_0_37, plain,
    (skoSC1 *
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/6 +
        skoSC1 *
        (skoSC1 *
         (1/120 +
          skoSC1 *
          (skoSC1 *
           (-1/5040 +
            skoSC1 *
            (skoSC1 *
             (1/362880 + skoSC1 * (skoSC1 * -1/39916800)))))))))) *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) <= 0),
    inference(resolve, [], [refute_0_36, refute_0_26])).

cnf(refute_0_38, plain,
    (sin(skoSC1) *
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800)))))))))))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 +
        skoSC1 *
        (-5/3 +
         skoSC1 *
         (43/120 +
          skoSC1 *
          (1/5 +
           skoSC1 *
           (-29/1680 +
            skoSC1 *
            (-23/2520 +
             skoSC1 *
             (7/17280 +
              skoSC1 *
              (19/90720 +
               skoSC1 *
               (-223/39916800 +
                skoSC1 *
                (-19/6652800 +
                 skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
     cos(skoSC1) *
     (-2 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-6 +
        skoSC1 *
        (-1/3 +
         skoSC1 *
         (skoSC1 *
          (1/60 +
           skoSC1 *
           (skoSC1 *
            (-1/2520 +
             skoSC1 *
             (skoSC1 *
              (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (1 +
      skoSC1 *
      (1 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (1/120 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (-1/5040 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (1/362880 +
              skoSC1 *
              (1/362880 +
               skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))) <=
     0 |
     sin(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))))),
    inference(split, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) |
     0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_41, plain,
    (cos(skoSC1) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_42, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) < 1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_43, plain,
    (X_000010 < cos(X_000009) | ~ lgen(1, X_000010, cos(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain,
    (~ lgen(1, X_000010, 1 - X_000009 ^ 2 / 2) |
     lgen(1, X_000010, cos(X_000009))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_45, plain,
    (X_000010 < cos(X_000009) | ~ lgen(1, X_000010, 1 - X_000009 ^ 2 / 2)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (X_000010 < cos(X_000009) |
     1 + X_000009 * (X_000009 * -1/2) <= X_000010),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_49, plain,
    (skoSC1 < 0 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <
     sin(skoSC1) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))))),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_50, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_51, plain,
    (skoSC1 < 0 |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) < sin(skoSC1) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))), inference(subst, [], [refute_0_24])).

cnf(refute_0_52, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))),
    inference(resolve, [], [refute_0_28, refute_0_51])).

cnf(refute_0_53, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1))),
    inference(resolve, [], [refute_0_35, refute_0_52])).

cnf(refute_0_54, plain,
    (skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) <= 0),
    inference(resolve, [], [refute_0_53, refute_0_50])).

cnf(refute_0_55, plain,
    (sin(skoSC1) *
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     sin(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))))),
    inference(split, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_56, refute_0_49])).

cnf(refute_0_58, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_35, refute_0_57])).

cnf(refute_0_59, plain,
    (skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_58, refute_0_48])).

cnf(refute_0_60, plain,
    (cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) < 0 |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     0 <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(split, [], [refute_0_55])).

cnf(refute_0_64, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) < 0 |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     0 <
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(split, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-3 +
      skoSC1 *
      (2 +
       skoSC1 *
       (9/2 + skoSC1 * (5/3 + skoSC1 * (-1/3 + skoSC1 * -1/6))))) <= 2),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain, (skoSC1 <= 3),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_68, plain,
    (~ lgen(0, X_000010, cos(X_000009)) | X_000010 <= cos(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_69, plain,
    (~ lgen(0, X_000010, 1 - X_000009 ^ 2 / 2) |
     lgen(0, X_000010, cos(X_000009))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_70, plain,
    (~ lgen(0, X_000010, 1 - X_000009 ^ 2 / 2) |
     X_000010 <= cos(X_000009)),
    inference(resolve, [], [refute_0_69, refute_0_68])).

cnf(refute_0_71, plain,
    (1 + X_000009 * (X_000009 * -1/2) < X_000010 |
     X_000010 <= cos(X_000009)), inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) <=
     cos(skoSC1)), inference(subst, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 <
     skoSC1 * (1 + skoSC1) | 1 + skoSC1 <= 0 |
     skoSC1 <=
     (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_74, plain,
    (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 <
     sin(skoSC1) * (1 + skoSC1) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    introduced(tautology, [assume])).

cnf(refute_0_75, plain,
    ((3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < sin(skoSC1) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     1 + skoSC1 <= 0), inference(split, [], [refute_0_74])).

cnf(refute_0_76, plain,
    ((3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < sin(skoSC1) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain, (-1 < skoSC1),
    inference(decision, [], [refute_0_67, refute_0_35])).

cnf(refute_0_78, plain,
    ((3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < sin(skoSC1) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (~ lgen(0, sin(X_000067), X_000068) | sin(X_000067) <= X_000068),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_80, plain,
    (X_000067 < 0 | ~ lgen(0, X_000067, X_000068) |
     lgen(0, sin(X_000067), X_000068)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_81, plain,
    (X_000067 < 0 | ~ lgen(0, X_000067, X_000068) |
     sin(X_000067) <= X_000068),
    inference(resolve, [], [refute_0_80, refute_0_79])).

cnf(refute_0_82, plain,
    (X_000068 < X_000067 | X_000067 < 0 | sin(X_000067) <= X_000068),
    inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (skoSC1 < 0 |
     (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < skoSC1 |
     sin(skoSC1) <=
     (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1)), inference(subst, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoSC1 < 0 |
     (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < skoSC1 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_83, refute_0_78])).

cnf(refute_0_85, plain,
    ((3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2) /
     (1 + skoSC1) < skoSC1 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_35, refute_0_84])).

cnf(refute_0_86, plain,
    (3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 <
     skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     1 + skoSC1 <= 0), inference(resolve, [], [refute_0_73, refute_0_85])).

cnf(refute_0_87, plain,
    (cos(skoSC1) < -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (cos(skoSC1) < -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_87, refute_0_77])).

cnf(refute_0_89, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_72, refute_0_88])).

cnf(refute_0_90, plain,
    (5/2 < skoSC1 * (-3/2 + skoSC1 * (2 + skoSC1 * 1/2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (~ lgen(0, X_000024, cos(X_000023)) | X_000024 <= cos(X_000023)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_92, plain,
    (~ lgen(0, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800) |
     lgen(0, X_000024, cos(X_000023))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_93, plain,
    (~ lgen(0, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800) |
     X_000024 <= cos(X_000023)),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (1 +
     X_000023 *
     (X_000023 *
      (-1/2 +
       X_000023 *
       (X_000023 *
        (1/24 +
         X_000023 *
         (X_000023 *
          (-1/720 +
           X_000023 *
           (X_000023 *
            (1/40320 + X_000023 * (X_000023 * -1/3628800))))))))) <
     X_000024 | X_000024 <= cos(X_000023)),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) <=
     cos(skoSC1)), inference(subst, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <
     -3/2 + skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * 1/2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_95, refute_0_88])).

cnf(refute_0_97, plain,
    (5/2 <
     skoSC1 *
     (-3/2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (1/2 +
        skoSC1 *
        (-1/24 +
         skoSC1 *
         (skoSC1 *
          (1/720 +
           skoSC1 *
           (skoSC1 * (-1/40320 + skoSC1 * (skoSC1 * 1/3628800))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (2 <
     skoSC1 *
     (-3 +
      skoSC1 *
      (2 +
       skoSC1 *
       (9/2 + skoSC1 * (5/3 + skoSC1 * (-1/3 + skoSC1 * -1/6))))) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97])).

cnf(refute_0_99, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_66, refute_0_98])).

cnf(refute_0_100, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) <= 1 + skoSC1 * (skoSC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_101, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) |
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) <= cos(skoSC1)),
    inference(subst, [], [refute_0_71])).

cnf(refute_0_102, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))),
    introduced(tautology, [assume])).

cnf(refute_0_103, plain,
    (cos(skoSC1) <
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) | 0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))),
    inference(split, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (cos(skoSC1) <
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) | skoSC1 <= -1),
    inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (-1 < skoSC1 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97])).

cnf(refute_0_106, plain,
    (cos(skoSC1) <
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_101, refute_0_106])).

cnf(refute_0_108, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_100, refute_0_107])).

cnf(refute_0_109, plain,
    (4 < skoSC1 * (-10 + skoSC1 * (-5 + skoSC1 * 2/3)) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-10 + skoSC1 * (-5 + skoSC1 * 2/3)) <= 4 | skoSC1 <= -1),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97])).

cnf(refute_0_111, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(resolve, [], [refute_0_110, refute_0_109])).

cnf(refute_0_112, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_111, refute_0_105])).

cnf(refute_0_113, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_112, refute_0_99])).

cnf(refute_0_114, plain,
    (skoSC1 <
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) |
     0 <= skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     skoSC1 *
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_115, plain,
    (sin(X_000067) < X_000068 | ~ lgen(1, sin(X_000067), X_000068)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_116, plain,
    (X_000067 < 0 | ~ lgen(1, X_000067, X_000068) |
     lgen(1, sin(X_000067), X_000068)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_117, plain,
    (sin(X_000067) < X_000068 | X_000067 < 0 |
     ~ lgen(1, X_000067, X_000068)),
    inference(resolve, [], [refute_0_116, refute_0_115])).

cnf(refute_0_118, plain,
    (sin(X_000067) < X_000068 | X_000067 < 0 | X_000068 <= X_000067),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (skoSC1 < 0 |
     sin(skoSC1) <
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     skoSC1), inference(subst, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (0 <= skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     sin(skoSC1)), inference(split, [], [refute_0_55])).

cnf(refute_0_121, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     sin(skoSC1)), inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     skoSC1), inference(resolve, [], [refute_0_121, refute_0_119])).

cnf(refute_0_123, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0 |
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     skoSC1), inference(resolve, [], [refute_0_35, refute_0_122])).

cnf(refute_0_124, plain,
    (0 <= skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6)))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_114])).

cnf(refute_0_125, plain,
    (cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 + skoSC1 * (-3/2 + skoSC1 * (1/2 + skoSC1 * 1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) < 0 |
     0 <
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 + skoSC1 * (-3/2 + skoSC1 * (1/2 + skoSC1 * 1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(split, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 + skoSC1 * (3/2 + skoSC1 * (-1/2 + skoSC1 * -1/6))))) <= 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    (2 <
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 + skoSC1 * (3/2 + skoSC1 * (-1/2 + skoSC1 * -1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97])).

cnf(refute_0_129, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) <= 0),
    inference(resolve, [], [refute_0_112, refute_0_129])).

cnf(refute_0_131, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_130, refute_0_113])).

cnf(refute_0_132, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) < 0 |
     0 <
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(split, [], [refute_0_60])).

cnf(refute_0_133, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 +
        skoSC1 *
        (4/3 + skoSC1 * (-2/3 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) <=
     2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (2 <
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 +
        skoSC1 *
        (4/3 + skoSC1 * (-2/3 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(decision, [],
              [refute_0_97, refute_0_90, refute_0_35, refute_0_67])).

cnf(refute_0_135, plain,
    (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_112, refute_0_135])).

cnf(refute_0_137, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_136, refute_0_131])).

cnf(refute_0_138, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_62, refute_0_137])).

cnf(refute_0_139, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_138, refute_0_47])).

cnf(refute_0_140, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_139, refute_0_42])).

cnf(refute_0_141, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) |
     cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_141, refute_0_137])).

cnf(refute_0_143, plain,
    (0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <= cos(skoSC1)),
    inference(split, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <= cos(skoSC1)),
    inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) < 0 |
     0 < 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     0 <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(split, [], [refute_0_142])).

cnf(refute_0_146, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (5 +
      skoSC1 *
      (8 +
       skoSC1 *
       (11/6 +
        skoSC1 *
        (-5/3 + skoSC1 * (-5/6 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) <=
     0), inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     0 <
     skoSC1 *
     (5 +
      skoSC1 *
      (8 +
       skoSC1 *
       (11/6 +
        skoSC1 *
        (-5/3 + skoSC1 * (-5/6 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(decision, [],
              [refute_0_97, refute_0_90, refute_0_35, refute_0_67])).

cnf(refute_0_148, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_97, refute_0_90, refute_0_35, refute_0_67])).

cnf(refute_0_150, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_152, plain,
    (skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_153, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(split, [], [refute_0_142])).

cnf(refute_0_154, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(decision, [],
              [refute_0_97, refute_0_90, refute_0_35, refute_0_67])).

cnf(refute_0_156, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_156, refute_0_152])).

cnf(refute_0_158, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_157, refute_0_151])).

cnf(refute_0_159, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-7/6 + skoSC1 * (-1/3 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) <=
     -2), inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (-2 <
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-7/6 + skoSC1 * (-1/3 + skoSC1 * (-5/36 + skoSC1 * 1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(decision, [],
              [refute_0_150, refute_0_97, refute_0_90, refute_0_35,
               refute_0_67])).

cnf(refute_0_161, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(resolve, [], [refute_0_159, refute_0_160])).

cnf(refute_0_162, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_161, refute_0_150])).

cnf(refute_0_163, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <= cos(skoSC1)),
    inference(resolve, [], [refute_0_144, refute_0_162])).

cnf(refute_0_164, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_163, refute_0_41])).

cnf(refute_0_165, plain,
    (0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (5/3 + skoSC1 * (5/6 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_164, refute_0_40])).

cnf(refute_0_166, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-13/12 +
         skoSC1 * (-1/4 + skoSC1 * (-13/72 + skoSC1 * -1/72)))))) <= -2),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-13/12 +
         skoSC1 * (-1/4 + skoSC1 * (-13/72 + skoSC1 * -1/72)))))) <= -2),
    inference(decision, [],
              [refute_0_161, refute_0_150, refute_0_97, refute_0_90,
               refute_0_35, refute_0_67])).

cnf(refute_0_168, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-13/12 +
         skoSC1 * (-1/4 + skoSC1 * (-13/72 + skoSC1 * -1/72)))))) <= -2),
    inference(resolve, [], [refute_0_166, refute_0_167])).

cnf(refute_0_169, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-13/12 +
         skoSC1 * (-1/4 + skoSC1 * (-13/72 + skoSC1 * -1/72)))))) <= -2),
    inference(decision, [],
              [refute_0_161, refute_0_150, refute_0_97, refute_0_90,
               refute_0_35, refute_0_67])).

cnf(refute_0_170, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-13/12 +
         skoSC1 * (-1/4 + skoSC1 * (-13/72 + skoSC1 * -1/72)))))) <= -2),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_172, plain,
    (cos(skoSC1) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_173, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) |
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_174, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) < cos(skoSC1) |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_175, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36))))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_138, refute_0_174])).

cnf(refute_0_176, plain,
    ((1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 + skoSC1 * (2/3 + skoSC1 * (5/36 + skoSC1 * -1/36)))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_175, refute_0_173])).

cnf(refute_0_177, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) |
     cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_177, refute_0_137])).

cnf(refute_0_179, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(split, [], [refute_0_178])).

cnf(refute_0_180, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (2 <
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97,
               refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_182, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_182, refute_0_172])).

cnf(refute_0_184, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-13/120 +
           skoSC1 *
           (-11/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_183, refute_0_171])).

cnf(refute_0_185, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-11/12 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (1/45 +
           skoSC1 *
           (-1/18 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (281/60480 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-7/51840 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (7/3110400 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) <= -2),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (-2 <
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-11/12 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (1/45 +
           skoSC1 *
           (-1/18 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (281/60480 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-7/51840 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (7/3110400 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(decision, [],
              [refute_0_67, refute_0_35, refute_0_90, refute_0_97,
               refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_187, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_185, refute_0_186])).

cnf(refute_0_188, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(resolve, [], [refute_0_187, refute_0_181])).

cnf(refute_0_189, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_190, plain,
    (cos(skoSC1) <
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_191, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) |
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_192, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) < cos(skoSC1) |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_193, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_64])).

cnf(refute_0_194, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_194, refute_0_137])).

cnf(refute_0_196, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6)))))) /
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_195, refute_0_192])).

cnf(refute_0_197, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     (1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 + skoSC1 * (-5/3 + skoSC1 * (1/3 + skoSC1 * 1/6))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     2 + skoSC1 * (8 + skoSC1 * (6 + skoSC1 * 1/3)) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_196, refute_0_191])).

cnf(refute_0_198, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * (-6 + skoSC1 * -1/3)) |
     cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_198, refute_0_137])).

cnf(refute_0_200, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(split, [], [refute_0_199])).

cnf(refute_0_201, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (2 <
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_188, refute_0_67, refute_0_35, refute_0_90,
               refute_0_97, refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_203, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 < skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(resolve, [], [refute_0_201, refute_0_202])).

cnf(refute_0_204, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (1 + skoSC1 * (-1/6 + skoSC1 * -1/6))) <= 0),
    inference(decision, [],
              [refute_0_188, refute_0_67, refute_0_35, refute_0_90,
               refute_0_97, refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_205, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(resolve, [], [refute_0_204, refute_0_203])).

cnf(refute_0_206, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_205, refute_0_190])).

cnf(refute_0_207, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-7 +
       skoSC1 *
       (-5/6 +
        skoSC1 *
        (5/4 +
         skoSC1 *
         (1/6 +
          skoSC1 *
          (-29/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (1/3780 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 *
               (-11/1814400 +
                skoSC1 * (1/604800 + skoSC1 * 1/10886400)))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_206, refute_0_189])).

cnf(refute_0_208, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (5 +
       skoSC1 *
       (-7/6 +
        skoSC1 *
        (-7/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-1/180 +
           skoSC1 *
           (-1/12 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (281/60480 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-7/51840 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (7/3110400 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) <= -2),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     -2 <
     skoSC1 *
     (7 +
      skoSC1 *
      (5 +
       skoSC1 *
       (-7/6 +
        skoSC1 *
        (-7/12 +
         skoSC1 *
         (1/2 +
          skoSC1 *
          (-1/180 +
           skoSC1 *
           (-1/12 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (281/60480 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-7/51840 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (7/3110400 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(decision, [],
              [refute_0_188, refute_0_67, refute_0_35, refute_0_90,
               refute_0_97, refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_210, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (2 <
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) |
     0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_188, refute_0_67, refute_0_35, refute_0_90,
               refute_0_97, refute_0_150, refute_0_161, refute_0_170])).

cnf(refute_0_212, plain,
    (0 < skoSC1 * (-1 + skoSC1 * (-1 + skoSC1 * (1/6 + skoSC1 * 1/6))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain,
    (0 <
     skoSC1 *
     (1 +
      skoSC1 *
      (1 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (1/120 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (-1/5040 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (1/362880 +
              skoSC1 *
              (1/362880 +
               skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_170, refute_0_161, refute_0_150, refute_0_97,
               refute_0_90, refute_0_35, refute_0_67, refute_0_188,
               refute_0_212])).

cnf(refute_0_214, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     sin(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))))),
    inference(resolve, [], [refute_0_39, refute_0_213])).

cnf(refute_0_215, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))))),
    inference(resolve, [], [refute_0_214, refute_0_25])).

cnf(refute_0_216, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-2 +
        skoSC1 *
        (-9/2 +
         skoSC1 *
         (-5/3 +
          skoSC1 *
          (43/120 +
           skoSC1 *
           (1/5 +
            skoSC1 *
            (-29/1680 +
             skoSC1 *
             (-23/2520 +
              skoSC1 *
              (7/17280 +
               skoSC1 *
               (19/90720 +
                skoSC1 *
                (-223/39916800 +
                 skoSC1 *
                 (-19/6652800 +
                  skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
      cos(skoSC1) *
      (-2 +
       skoSC1 *
       (-8 +
        skoSC1 *
        (-6 +
         skoSC1 *
         (-1/3 +
          skoSC1 *
          (skoSC1 *
           (1/60 +
            skoSC1 *
            (skoSC1 *
             (-1/2520 +
              skoSC1 *
              (skoSC1 *
               (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))))),
    inference(resolve, [], [refute_0_35, refute_0_215])).

cnf(refute_0_217, plain,
    (skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) *
     (skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (-1/6 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (1/120 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (-1/5040 +
              skoSC1 *
              (1/362880 +
               skoSC1 *
               (1/362880 +
                skoSC1 * (-1/39916800 + skoSC1 * -1/39916800)))))))))))) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-9/2 +
        skoSC1 *
        (-5/3 +
         skoSC1 *
         (43/120 +
          skoSC1 *
          (1/5 +
           skoSC1 *
           (-29/1680 +
            skoSC1 *
            (-23/2520 +
             skoSC1 *
             (7/17280 +
              skoSC1 *
              (19/90720 +
               skoSC1 *
               (-223/39916800 +
                skoSC1 *
                (-19/6652800 +
                 skoSC1 * (1/19958400 + skoSC1 * 1/39916800))))))))))))) +
     cos(skoSC1) *
     (-2 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-6 +
        skoSC1 *
        (-1/3 +
         skoSC1 *
         (skoSC1 *
          (1/60 +
           skoSC1 *
           (skoSC1 *
            (-1/2520 +
             skoSC1 *
             (skoSC1 *
              (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (1 +
      skoSC1 *
      (1 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (1/120 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (-1/5040 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (1/362880 +
              skoSC1 *
              (1/362880 +
               skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))) <=
     0), inference(resolve, [], [refute_0_216, refute_0_20])).

cnf(refute_0_218, plain,
    (cos(skoSC1) *
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 +
         skoSC1 *
         (83/120 +
          skoSC1 *
          (59/360 +
           skoSC1 *
           (-269/5040 +
            skoSC1 *
            (-19/2520 +
             skoSC1 *
             (241/120960 +
              skoSC1 *
              (1/5760 +
               skoSC1 *
               (-551/13305600 +
                skoSC1 *
                (-71/29937600 +
                 skoSC1 *
                 (1/1871100 +
                  skoSC1 *
                  (1/47900160 + skoSC1 * -1/239500800)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (1 +
      skoSC1 *
      (1 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-1/6 +
         skoSC1 *
         (1/120 +
          skoSC1 *
          (1/120 +
           skoSC1 *
           (-1/5040 +
            skoSC1 *
            (-1/5040 +
             skoSC1 *
             (1/362880 +
              skoSC1 *
              (1/362880 +
               skoSC1 * (-1/39916800 + skoSC1 * -1/39916800))))))))))) <=
     0), inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (cos(skoSC1) *
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 +
         skoSC1 *
         (83/120 +
          skoSC1 *
          (59/360 +
           skoSC1 *
           (-269/5040 +
            skoSC1 *
            (-19/2520 +
             skoSC1 *
             (241/120960 +
              skoSC1 *
              (1/5760 +
               skoSC1 *
               (-551/13305600 +
                skoSC1 *
                (-71/29937600 +
                 skoSC1 *
                 (1/1871100 +
                  skoSC1 *
                  (1/47900160 + skoSC1 * -1/239500800)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_218, refute_0_213])).

cnf(refute_0_220, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 +
     skoSC1 *
     (-8 +
      skoSC1 *
      (-6 +
       skoSC1 *
       (-1/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
     cos(skoSC1) * (2 + skoSC1 * 2) < 0 |
     0 <
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
     cos(skoSC1) * (2 + skoSC1 * 2) |
     0 <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 +
         skoSC1 *
         (83/120 +
          skoSC1 *
          (59/360 +
           skoSC1 *
           (-269/5040 +
            skoSC1 *
            (-19/2520 +
             skoSC1 *
             (241/120960 +
              skoSC1 *
              (1/5760 +
               skoSC1 *
               (-551/13305600 +
                skoSC1 *
                (-71/29937600 +
                 skoSC1 *
                 (1/1871100 +
                  skoSC1 *
                  (1/47900160 + skoSC1 * -1/239500800)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(split, [], [refute_0_219])).

cnf(refute_0_223, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 +
     skoSC1 *
     (-8 +
      skoSC1 *
      (-6 +
       skoSC1 *
       (-1/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 +
        skoSC1 *
        (4/3 +
         skoSC1 *
         (-83/120 +
          skoSC1 *
          (-59/360 +
           skoSC1 *
           (269/5040 +
            skoSC1 *
            (19/2520 +
             skoSC1 *
             (-241/120960 +
              skoSC1 *
              (-1/5760 +
               skoSC1 *
               (551/13305600 +
                skoSC1 *
                (71/29937600 +
                 skoSC1 *
                 (-1/1871100 +
                  skoSC1 *
                  (-1/47900160 + skoSC1 * 1/239500800)))))))))))))) <= 2),
    inference(arithmetic, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (2 <
     skoSC1 *
     (-3 +
      skoSC1 *
      (3 +
       skoSC1 *
       (11/2 +
        skoSC1 *
        (4/3 +
         skoSC1 *
         (-83/120 +
          skoSC1 *
          (-59/360 +
           skoSC1 *
           (269/5040 +
            skoSC1 *
            (19/2520 +
             skoSC1 *
             (-241/120960 +
              skoSC1 *
              (-1/5760 +
               skoSC1 *
               (551/13305600 +
                skoSC1 *
                (71/29937600 +
                 skoSC1 *
                 (-1/1871100 +
                  skoSC1 *
                  (-1/47900160 + skoSC1 * 1/239500800)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_170, refute_0_161, refute_0_150, refute_0_97,
               refute_0_90, refute_0_35, refute_0_67, refute_0_188,
               refute_0_212])).

cnf(refute_0_225, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 +
     skoSC1 *
     (-8 +
      skoSC1 *
      (-6 +
       skoSC1 *
       (-1/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) <= 1 + skoSC1 * (skoSC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_227, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) |
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) <= cos(skoSC1)),
    inference(subst, [], [refute_0_71])).

cnf(refute_0_228, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) <
     cos(skoSC1) * (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))),
    introduced(tautology, [assume])).

cnf(refute_0_229, plain,
    (cos(skoSC1) <
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) | 0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))),
    inference(split, [], [refute_0_228])).

cnf(refute_0_230, plain,
    (cos(skoSC1) <
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (cos(skoSC1) <
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_230, refute_0_105])).

cnf(refute_0_232, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400))))))))))) /
     (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_227, refute_0_231])).

cnf(refute_0_233, plain,
    (2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_226, refute_0_232])).

cnf(refute_0_234, plain,
    (4 <
     skoSC1 *
     (-10 +
      skoSC1 *
      (-5 +
       skoSC1 *
       (2/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-10 +
      skoSC1 *
      (-5 +
       skoSC1 *
       (2/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) <= 4 |
     skoSC1 <= -1),
    inference(decision, [],
              [refute_0_170, refute_0_161, refute_0_150, refute_0_97,
               refute_0_90, refute_0_35, refute_0_67, refute_0_188,
               refute_0_212])).

cnf(refute_0_236, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 <= -1), inference(resolve, [], [refute_0_235, refute_0_234])).

cnf(refute_0_237, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <=
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_236, refute_0_105])).

cnf(refute_0_238, plain,
    (-2 +
     skoSC1 *
     (-8 +
      skoSC1 *
      (-6 +
       skoSC1 *
       (-1/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) <
     cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_237, refute_0_225])).

cnf(refute_0_239, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_221, refute_0_238])).

cnf(refute_0_240, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 *
       (-3 +
        skoSC1 *
        (-11/2 +
         skoSC1 *
         (-4/3 +
          skoSC1 *
          (83/120 +
           skoSC1 *
           (59/360 +
            skoSC1 *
            (-269/5040 +
             skoSC1 *
             (-19/2520 +
              skoSC1 *
              (241/120960 +
               skoSC1 *
               (1/5760 +
                skoSC1 *
                (-551/13305600 +
                 skoSC1 *
                 (-71/29937600 +
                  skoSC1 *
                  (1/1871100 +
                   skoSC1 *
                   (1/47900160 + skoSC1 * -1/239500800))))))))))))))) /
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_239, refute_0_19])).

cnf(refute_0_241, plain,
    ((1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 +
      skoSC1 *
      (8 +
       skoSC1 *
       (6 +
        skoSC1 *
        (1/3 +
         skoSC1 *
         (skoSC1 *
          (-1/60 +
           skoSC1 *
           (skoSC1 *
            (1/2520 +
             skoSC1 *
             (skoSC1 *
              (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
      cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 +
      skoSC1 *
      (-3 +
       skoSC1 *
       (-11/2 +
        skoSC1 *
        (-4/3 +
         skoSC1 *
         (83/120 +
          skoSC1 *
          (59/360 +
           skoSC1 *
           (-269/5040 +
            skoSC1 *
            (-19/2520 +
             skoSC1 *
             (241/120960 +
              skoSC1 *
              (1/5760 +
               skoSC1 *
               (-551/13305600 +
                skoSC1 *
                (-71/29937600 +
                 skoSC1 *
                 (1/1871100 +
                  skoSC1 *
                  (1/47900160 + skoSC1 * -1/239500800)))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     2 +
     skoSC1 *
     (8 +
      skoSC1 *
      (6 +
       skoSC1 *
       (1/3 +
        skoSC1 *
        (skoSC1 *
         (-1/60 +
          skoSC1 *
          (skoSC1 *
           (1/2520 +
            skoSC1 *
            (skoSC1 *
             (-1/181440 + skoSC1 * (skoSC1 * 1/19958400)))))))))) +
     cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_240, refute_0_18])).

cnf(refute_0_242, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <=
     -2 +
     skoSC1 *
     (-8 +
      skoSC1 *
      (-6 +
       skoSC1 *
       (-1/3 +
        skoSC1 *
        (skoSC1 *
         (1/60 +
          skoSC1 *
          (skoSC1 *
           (-1/2520 +
            skoSC1 *
            (skoSC1 *
             (1/181440 + skoSC1 * (skoSC1 * -1/19958400)))))))))) |
     cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_242, refute_0_238])).

cnf(refute_0_244, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(split, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(arithmetic, [], [refute_0_244])).

cnf(refute_0_246, plain,
    (2 <
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(decision, [],
              [refute_0_170, refute_0_161, refute_0_150, refute_0_97,
               refute_0_90, refute_0_35, refute_0_67, refute_0_188,
               refute_0_212])).

cnf(refute_0_247, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(resolve, [], [refute_0_245, refute_0_246])).

cnf(refute_0_248, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_247, refute_0_17])).

cnf(refute_0_249, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-5 +
      skoSC1 *
      (-8 +
       skoSC1 *
       (-11/6 +
        skoSC1 *
        (19/12 +
         skoSC1 *
         (13/24 +
          skoSC1 *
          (-1/12 +
           skoSC1 *
           (-109/1680 +
            skoSC1 *
            (1/1344 +
             skoSC1 *
             (229/72576 +
              skoSC1 *
              (23/907200 +
               skoSC1 *
               (-19/211200 +
                skoSC1 *
                (-43/59875200 +
                 skoSC1 *
                 (221/119750400 +
                  skoSC1 *
                  (1/47900160 +
                   skoSC1 *
                   (-713/25147584000 +
                    skoSC1 *
                    (skoSC1 *
                     (127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (-1/362125209600 +
                        skoSC1 *
                        (skoSC1 * 1/72425041920000)))))))))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_248, refute_0_16])).

cnf(refute_0_250, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-11/12 +
         skoSC1 *
         (1/8 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-107/5040 +
            skoSC1 *
            (1/240 +
             skoSC1 *
             (91/51840 +
              skoSC1 *
              (-151/907200 +
               skoSC1 *
               (-2041/39916800 +
                skoSC1 *
                (733/239500800 +
                 skoSC1 *
                 (17/34214400 +
                  skoSC1 *
                  (-1/22809600 +
                   skoSC1 *
                   (271/50295168000 +
                    skoSC1 *
                    (skoSC1 *
                     (-127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (1/362125209600 +
                        skoSC1 *
                        (skoSC1 * -1/72425041920000)))))))))))))))))))) <=
     -2), inference(arithmetic, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (-2 <
     skoSC1 *
     (7 +
      skoSC1 *
      (6 +
       skoSC1 *
       (-1/6 +
        skoSC1 *
        (-11/12 +
         skoSC1 *
         (1/8 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-107/5040 +
            skoSC1 *
            (1/240 +
             skoSC1 *
             (91/51840 +
              skoSC1 *
              (-151/907200 +
               skoSC1 *
               (-2041/39916800 +
                skoSC1 *
                (733/239500800 +
                 skoSC1 *
                 (17/34214400 +
                  skoSC1 *
                  (-1/22809600 +
                   skoSC1 *
                   (271/50295168000 +
                    skoSC1 *
                    (skoSC1 *
                     (-127/402361344000 +
                      skoSC1 *
                      (skoSC1 *
                       (1/362125209600 +
                        skoSC1 *
                        (skoSC1 * -1/72425041920000)))))))))))))))))))) |
     sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(decision, [],
              [refute_0_170, refute_0_161, refute_0_150, refute_0_97,
               refute_0_90, refute_0_35, refute_0_67, refute_0_188,
               refute_0_212])).

cnf(refute_0_252, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2 |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(resolve, [], [refute_0_250, refute_0_251])).

cnf(refute_0_253, plain,
    (sin(skoSC1) * (1 + skoSC1) <=
     3 + skoSC1 * (4 + skoSC1 * (-2 + skoSC1 * -1)) + cos(skoSC1) * 2),
    inference(resolve, [], [refute_0_252, refute_0_246])).

cnf(refute_0_254, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 <
     sin(skoSC1) * (-1 + skoSC1 * -1) |
     cos(skoSC1) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))),
    inference(resolve, [], [refute_0_253, refute_0_15])).

cnf(refute_0_255, plain,
    (skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_256, plain,
    (skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_257, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 + skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6)))) |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_258, plain,
    ((2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_259, plain,
    ((2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1)) <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) * (skoSC1 * (1 + skoSC1)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) | skoSC1 * (1 + skoSC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_260, plain,
    (skoSC1 < 0 |
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1)) <
     sin(skoSC1) |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1))),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_261, plain,
    (skoSC1 <
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) |
     0 <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) |
     skoSC1 *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_262, plain,
    (skoSC1 < 0 |
     sin(skoSC1) <
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <= skoSC1),
    inference(subst, [], [refute_0_118])).

cnf(refute_0_263, plain,
    (0 <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <= sin(skoSC1)),
    inference(split, [], [refute_0_13])).

cnf(refute_0_264, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <= sin(skoSC1)),
    inference(arithmetic, [], [refute_0_263])).

cnf(refute_0_265, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <= skoSC1),
    inference(resolve, [], [refute_0_264, refute_0_262])).

cnf(refute_0_266, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-10 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) /
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <= skoSC1),
    inference(resolve, [], [refute_0_35, refute_0_265])).

cnf(refute_0_267, plain,
    (0 <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
     sin(skoSC1) * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 *
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 +
      sin(skoSC1) * (1 + skoSC1)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-10 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2))),
    inference(resolve, [], [refute_0_266, refute_0_261])).

cnf(refute_0_268, plain,
    (sin(skoSC1) * (skoSC1 * (1 + skoSC1)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_267])).

cnf(refute_0_269, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0 |
     sin(skoSC1) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1))),
    inference(split, [], [refute_0_268])).

cnf(refute_0_270, plain,
    (skoSC1 < 0 |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1)) |
     skoSC1 * (1 + skoSC1) <= 0),
    inference(resolve, [], [refute_0_269, refute_0_260])).

cnf(refute_0_271, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
      cos(skoSC1) *
      (-2 + skoSC1 * (-8 + skoSC1 * -6) +
       cos(skoSC1) * (-2 + skoSC1 * -2))) / (skoSC1 * (1 + skoSC1)) |
     skoSC1 * (1 + skoSC1) <= 0),
    inference(resolve, [], [refute_0_35, refute_0_270])).

cnf(refute_0_272, plain,
    (skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) * (skoSC1 * (1 + skoSC1)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0),
    inference(resolve, [], [refute_0_271, refute_0_259])).

cnf(refute_0_273, plain,
    (cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 + skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0),
    inference(arithmetic, [], [refute_0_272])).

cnf(refute_0_274, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0 |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_273])).

cnf(refute_0_275, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-8 + skoSC1 * -6) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_274])).

cnf(refute_0_276, plain,
    ((-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * -6) | 2 + skoSC1 * 2 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_277, plain,
    ((-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_278, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-8 + skoSC1 * -6)),
    introduced(tautology, [assume])).

cnf(refute_0_279, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     2 + skoSC1 * 2 <= 0 |
     cos(skoSC1) <= (-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(split, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     cos(skoSC1) <= (-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2) |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (~ lgen(0, cos(X_000030), X_000031) | cos(X_000030) <= X_000031),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_282, plain,
    (~ lgen(0,
            1 - X_000030 ^ 2 / 2 + X_000030 ^ 4 / 24 - X_000030 ^ 6 / 720 +
            X_000030 ^ 8 / 40320 - X_000030 ^ 10 / 3628800 +
            X_000030 ^ 12 / 479001600, X_000031) |
     lgen(0, cos(X_000030), X_000031)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_283, plain,
    (~ lgen(0,
            1 - X_000030 ^ 2 / 2 + X_000030 ^ 4 / 24 - X_000030 ^ 6 / 720 +
            X_000030 ^ 8 / 40320 - X_000030 ^ 10 / 3628800 +
            X_000030 ^ 12 / 479001600, X_000031) |
     cos(X_000030) <= X_000031),
    inference(resolve, [], [refute_0_282, refute_0_281])).

cnf(refute_0_284, plain,
    (X_000031 <
     1 +
     X_000030 *
     (X_000030 *
      (-1/2 +
       X_000030 *
       (X_000030 *
        (1/24 +
         X_000030 *
         (X_000030 *
          (-1/720 +
           X_000030 *
           (X_000030 *
            (1/40320 +
             X_000030 *
             (X_000030 *
              (-1/3628800 + X_000030 * (X_000030 * 1/479001600))))))))))) |
     cos(X_000030) <= X_000031),
    inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 *
            (1/40320 +
             skoSC1 *
             (skoSC1 *
              (-1/3628800 + skoSC1 * (skoSC1 * 1/479001600))))))))))) |
     cos(skoSC1) <=
     -3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12)))),
    inference(subst, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) * (-1 + skoSC1 * -1) |
     0 <= -1 + skoSC1 * -1 |
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) <= skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_287, plain,
    (~ lgen(0, X_000089, sin(X_000088)) | X_000089 <= sin(X_000088)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_288, plain,
    (X_000088 < 0 | ~ lgen(0, X_000089, X_000088 - X_000088 ^ 3 / 6) |
     lgen(0, X_000089, sin(X_000088))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_289, plain,
    (X_000088 < 0 | ~ lgen(0, X_000089, X_000088 - X_000088 ^ 3 / 6) |
     X_000089 <= sin(X_000088)),
    inference(resolve, [], [refute_0_288, refute_0_287])).

cnf(refute_0_290, plain,
    (X_000088 * (1 + X_000088 * (X_000088 * -1/6)) < X_000089 |
     X_000088 < 0 | X_000089 <= sin(X_000088)),
    inference(arithmetic, [], [refute_0_289])).

cnf(refute_0_291, plain,
    (skoSC1 < 0 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) |
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) <= sin(skoSC1)),
    inference(subst, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 <
     sin(skoSC1) * (-1 + skoSC1 * -1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    introduced(tautology, [assume])).

cnf(refute_0_293, plain,
    (sin(skoSC1) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) | 0 <= -1 + skoSC1 * -1 |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(split, [], [refute_0_292])).

cnf(refute_0_294, plain,
    (sin(skoSC1) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (sin(skoSC1) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_294, refute_0_77])).

cnf(refute_0_296, plain,
    (skoSC1 < 0 |
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_291, refute_0_295])).

cnf(refute_0_297, plain,
    (skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) <
     (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2) /
     (-1 + skoSC1 * -1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_35, refute_0_296])).

cnf(refute_0_298, plain,
    (-3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 <
     skoSC1 * (1 + skoSC1 * (skoSC1 * -1/6)) * (-1 + skoSC1 * -1) |
     0 <= -1 + skoSC1 * -1 |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_286, refute_0_297])).

cnf(refute_0_299, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     cos(skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     cos(skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_299, refute_0_77])).

cnf(refute_0_301, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 *
            (1/40320 +
             skoSC1 *
             (skoSC1 *
              (-1/3628800 + skoSC1 * (skoSC1 * 1/479001600))))))))))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_285, refute_0_300])).

cnf(refute_0_302, plain,
    (-5/2 <
     skoSC1 *
     (3/2 +
      skoSC1 *
      (-2 +
       skoSC1 *
       (-5/12 +
        skoSC1 *
        (1/8 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 *
            (1/40320 +
             skoSC1 *
             (skoSC1 *
              (-1/3628800 + skoSC1 * (skoSC1 * 1/479001600))))))))))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_301])).

cnf(refute_0_303, plain,
    (~ lgen(0, cos(X_000016), X_000017) | cos(X_000016) <= X_000017),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_304, plain,
    (~ lgen(0, 1 - X_000016 ^ 2 / 2 + X_000016 ^ 4 / 24, X_000017) |
     lgen(0, cos(X_000016), X_000017)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_305, plain,
    (~ lgen(0, 1 - X_000016 ^ 2 / 2 + X_000016 ^ 4 / 24, X_000017) |
     cos(X_000016) <= X_000017),
    inference(resolve, [], [refute_0_304, refute_0_303])).

cnf(refute_0_306, plain,
    (X_000017 <
     1 + X_000016 * (X_000016 * (-1/2 + X_000016 * (X_000016 * 1/24))) |
     cos(X_000016) <= X_000017),
    inference(arithmetic, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) |
     cos(skoSC1) <=
     -3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12)))),
    inference(subst, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (-3/2 +
     skoSC1 * (-3/2 + skoSC1 * (3/2 + skoSC1 * (5/12 + skoSC1 * -1/12))) <
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_307, refute_0_300])).

cnf(refute_0_309, plain,
    (-5/2 <
     skoSC1 * (3/2 + skoSC1 * (-2 + skoSC1 * (-5/12 + skoSC1 * 1/8))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_308])).

cnf(refute_0_310, plain,
    (-1 < skoSC1 |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(decision, [],
              [refute_0_302, refute_0_309, refute_0_67, refute_0_35])).

cnf(refute_0_311, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     cos(skoSC1) <= (-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(resolve, [], [refute_0_280, refute_0_310])).

cnf(refute_0_312, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (-2 + skoSC1 * (-8 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(resolve, [], [refute_0_311, refute_0_277])).

cnf(refute_0_313, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-8 + skoSC1 * -6) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * 2 <= 0),
    inference(resolve, [], [refute_0_312, refute_0_276])).

cnf(refute_0_314, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (10 + skoSC1 * (5 + skoSC1 * -1)) <= -4 | skoSC1 <= -1),
    inference(arithmetic, [], [refute_0_313])).

cnf(refute_0_315, plain,
    (-4 < skoSC1 * (10 + skoSC1 * (5 + skoSC1 * -1)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 <= -1),
    inference(decision, [],
              [refute_0_302, refute_0_309, refute_0_67, refute_0_35])).

cnf(refute_0_316, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 <= -1), inference(resolve, [], [refute_0_314, refute_0_315])).

cnf(refute_0_317, plain,
    (-2 + skoSC1 * (-8 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_316, refute_0_310])).

cnf(refute_0_318, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0 |
     cos(skoSC1) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_275, refute_0_317])).

cnf(refute_0_319, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (1 + skoSC1) <= 0 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_318, refute_0_258])).

cnf(refute_0_320, plain,
    (skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_321, plain,
    (skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_322, plain,
    ((2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_323, plain,
    ((2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_324, plain,
    (skoSC1 * (1 + skoSC1) < 0 | 0 < skoSC1 * (1 + skoSC1) |
     0 <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) +
     cos(skoSC1) *
     (-2 + skoSC1 * (-8 + skoSC1 * -6) +
      cos(skoSC1) * (-2 + skoSC1 * -2)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(split, [], [refute_0_268])).

cnf(refute_0_325, plain,
    (0 < skoSC1 * (-1 + skoSC1 * -1) | 0 < skoSC1 * (1 + skoSC1) |
     cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_324])).

cnf(refute_0_326, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (-1 + skoSC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_302, refute_0_309, refute_0_67, refute_0_35])).

cnf(refute_0_327, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     cos(skoSC1) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_326, refute_0_325])).

cnf(refute_0_328, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_327])).

cnf(refute_0_329, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-8 + skoSC1 * -6) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     cos(skoSC1) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_328])).

cnf(refute_0_330, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     cos(skoSC1) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_329, refute_0_317])).

cnf(refute_0_331, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1)))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_330, refute_0_323])).

cnf(refute_0_332, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (3 + skoSC1 * (-2 + skoSC1 * (-4 + skoSC1 * -1))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_331, refute_0_322])).

cnf(refute_0_333, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-8 + skoSC1 * -6) |
     cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_332])).

cnf(refute_0_334, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_333, refute_0_317])).

cnf(refute_0_335, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0 |
     cos(skoSC1) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(split, [], [refute_0_334])).

cnf(refute_0_336, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     cos(skoSC1) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(arithmetic, [], [refute_0_335])).

cnf(refute_0_337, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(decision, [],
              [refute_0_302, refute_0_309, refute_0_67, refute_0_35])).

cnf(refute_0_338, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     cos(skoSC1) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_336, refute_0_337])).

cnf(refute_0_339, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_338, refute_0_321])).

cnf(refute_0_340, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-5 + skoSC1 * (-7 + skoSC1 * (skoSC1 * 2))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_339, refute_0_320])).

cnf(refute_0_341, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     skoSC1 *
     (7 + skoSC1 * (5 + skoSC1 * (-2 + skoSC1 * (-3/2 + skoSC1 * 1/2)))) <=
     -2), inference(arithmetic, [], [refute_0_340])).

cnf(refute_0_342, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     -2 <
     skoSC1 *
     (7 + skoSC1 * (5 + skoSC1 * (-2 + skoSC1 * (-3/2 + skoSC1 * 1/2)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(decision, [],
              [refute_0_302, refute_0_309, refute_0_67, refute_0_35])).

cnf(refute_0_343, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(resolve, [], [refute_0_341, refute_0_342])).

cnf(refute_0_344, plain,
    (0 < skoSC1 * (1 + skoSC1) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_343, refute_0_337])).

cnf(refute_0_345, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 +
      skoSC1 *
      (3 +
       skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6))))) /
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_319, refute_0_344])).

cnf(refute_0_346, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 +
     skoSC1 *
     (3 + skoSC1 * (-3 + skoSC1 * (-5 + skoSC1 * (-5/6 + skoSC1 * 1/6)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (8 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_345, refute_0_257])).

cnf(refute_0_347, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-8 + skoSC1 * -6) |
     cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(arithmetic, [], [refute_0_346])).

cnf(refute_0_348, plain,
    (cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_347, refute_0_317])).

cnf(refute_0_349, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(split, [], [refute_0_348])).

cnf(refute_0_350, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(arithmetic, [], [refute_0_349])).

cnf(refute_0_351, plain,
    (-2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(decision, [],
              [refute_0_344, refute_0_302, refute_0_309, refute_0_67,
               refute_0_35])).

cnf(refute_0_352, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     cos(skoSC1) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_350, refute_0_351])).

cnf(refute_0_353, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_352, refute_0_256])).

cnf(refute_0_354, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 *
     (-5 +
      skoSC1 * (-8 + skoSC1 * (-1 + skoSC1 * (13/6 + skoSC1 * 1/6)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_353, refute_0_255])).

cnf(refute_0_355, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     skoSC1 *
     (7 + skoSC1 * (6 + skoSC1 * (-1 + skoSC1 * (-5/3 + skoSC1 * 1/3)))) <=
     -2), inference(arithmetic, [], [refute_0_354])).

cnf(refute_0_356, plain,
    (-2 <
     skoSC1 *
     (7 + skoSC1 * (6 + skoSC1 * (-1 + skoSC1 * (-5/3 + skoSC1 * 1/3)))) |
     sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(decision, [],
              [refute_0_344, refute_0_302, refute_0_309, refute_0_67,
               refute_0_35])).

cnf(refute_0_357, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2 |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(resolve, [], [refute_0_355, refute_0_356])).

cnf(refute_0_358, plain,
    (sin(skoSC1) * (-1 + skoSC1 * -1) <=
     -3 + skoSC1 * (-4 + skoSC1 * (2 + skoSC1)) + cos(skoSC1) * -2),
    inference(resolve, [], [refute_0_357, refute_0_351])).

cnf(refute_0_359, plain,
    (cos(skoSC1) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))),
    inference(resolve, [], [refute_0_358, refute_0_254])).

cnf(refute_0_360, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0 |
     cos(skoSC1) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(split, [], [refute_0_359])).

cnf(refute_0_361, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-10 + skoSC1 * -6) |
     cos(skoSC1) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(arithmetic, [], [refute_0_360])).

cnf(refute_0_362, plain,
    ((-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-10 + skoSC1 * -6) | 2 + skoSC1 * 2 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_363, plain,
    ((-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_364, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-10 + skoSC1 * -6)),
    introduced(tautology, [assume])).

cnf(refute_0_365, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     2 + skoSC1 * 2 <= 0 |
     cos(skoSC1) <=
     (-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(split, [], [refute_0_364])).

cnf(refute_0_366, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     cos(skoSC1) <=
     (-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2) |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_365])).

cnf(refute_0_367, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     cos(skoSC1) <=
     (-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(resolve, [], [refute_0_366, refute_0_77])).

cnf(refute_0_368, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (-2 + skoSC1 * (-10 + skoSC1 * -6)) / (2 + skoSC1 * 2)),
    inference(resolve, [], [refute_0_367, refute_0_363])).

cnf(refute_0_369, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) * (2 + skoSC1 * 2) <=
     -2 + skoSC1 * (-10 + skoSC1 * -6) | 2 + skoSC1 * 2 <= 0),
    inference(resolve, [], [refute_0_368, refute_0_362])).

cnf(refute_0_370, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     skoSC1 * (12 + skoSC1 * (5 + skoSC1 * -1)) <= -4 | skoSC1 <= -1),
    inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) < 0 |
     0 < 2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) |
     0 <= 2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))),
    inference(split, [], [refute_0_359])).

cnf(refute_0_372, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) < cos(skoSC1) * (-2 + skoSC1 * -2) |
     -2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     skoSC1 * (skoSC1 * (6 + skoSC1 * 2)) <= 2),
    inference(arithmetic, [], [refute_0_371])).

cnf(refute_0_373, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) <=
     1 + skoSC1 * (skoSC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_374, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) |
     (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) <= cos(skoSC1)),
    inference(subst, [], [refute_0_71])).

cnf(refute_0_375, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) < cos(skoSC1) * (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    introduced(tautology, [assume])).

cnf(refute_0_376, plain,
    (cos(skoSC1) < (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    inference(split, [], [refute_0_375])).

cnf(refute_0_377, plain,
    (cos(skoSC1) < (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6) |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_376])).

cnf(refute_0_378, plain,
    (cos(skoSC1) < (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    inference(resolve, [], [refute_0_377, refute_0_77])).

cnf(refute_0_379, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <
     (2 + skoSC1 * (10 + skoSC1 * 6)) / (-2 + skoSC1 * -2) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    inference(resolve, [], [refute_0_374, refute_0_378])).

cnf(refute_0_380, plain,
    (2 + skoSC1 * (10 + skoSC1 * 6) <
     (1 + skoSC1 * (skoSC1 * -1/2)) * (-2 + skoSC1 * -2) |
     0 <= -2 + skoSC1 * -2 |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    inference(resolve, [], [refute_0_373, refute_0_379])).

cnf(refute_0_381, plain,
    (4 < skoSC1 * (-12 + skoSC1 * (-5 + skoSC1)) |
     cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6) |
     skoSC1 <= -1), inference(arithmetic, [], [refute_0_380])).

cnf(refute_0_382, plain,
    (skoSC1 * (-12 + skoSC1 * (-5 + skoSC1)) <= 4 | skoSC1 <= -1),
    inference(decision, [], [refute_0_35, refute_0_67])).

cnf(refute_0_383, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6) |
     skoSC1 <= -1), inference(resolve, [], [refute_0_382, refute_0_381])).

cnf(refute_0_384, plain,
    (cos(skoSC1) * (-2 + skoSC1 * -2) <= 2 + skoSC1 * (10 + skoSC1 * 6)),
    inference(resolve, [], [refute_0_383, refute_0_77])).

cnf(refute_0_385, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     skoSC1 * (skoSC1 * (6 + skoSC1 * 2)) <= 2),
    inference(resolve, [], [refute_0_384, refute_0_372])).

cnf(refute_0_386, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     -4 < skoSC1 * (12 + skoSC1 * (5 + skoSC1 * -1)) | skoSC1 <= -1),
    inference(decision, [], [refute_0_385, refute_0_67, refute_0_35])).

cnf(refute_0_387, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     skoSC1 <= -1), inference(resolve, [], [refute_0_370, refute_0_386])).

cnf(refute_0_388, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2) |
     -1 < skoSC1),
    inference(decision, [], [refute_0_385, refute_0_67, refute_0_35])).

cnf(refute_0_389, plain,
    (-2 + skoSC1 * (-10 + skoSC1 * -6) < cos(skoSC1) * (2 + skoSC1 * 2)),
    inference(resolve, [], [refute_0_387, refute_0_388])).

cnf(refute_0_390, plain,
    (cos(skoSC1) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_361, refute_0_389])).

cnf(refute_0_391, plain,
    (1 +
     skoSC1 *
     (skoSC1 *
      (-1/2 +
       skoSC1 *
       (skoSC1 *
        (1/24 +
         skoSC1 *
         (skoSC1 *
          (-1/720 +
           skoSC1 *
           (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800))))))))) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_390, refute_0_11])).

cnf(refute_0_392, plain,
    ((1 +
      skoSC1 *
      (skoSC1 *
       (-1/2 +
        skoSC1 *
        (skoSC1 *
         (1/24 +
          skoSC1 *
          (skoSC1 *
           (-1/720 +
            skoSC1 *
            (skoSC1 * (1/40320 + skoSC1 * (skoSC1 * -1/3628800)))))))))) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) |
     2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_391, refute_0_6])).

cnf(refute_0_393, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-10 + skoSC1 * -6) |
     cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800)))))))))))),
    inference(arithmetic, [], [refute_0_392])).

cnf(refute_0_394, plain,
    (cos(skoSC1) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800)))))))))))),
    inference(resolve, [], [refute_0_393, refute_0_389])).

cnf(refute_0_395, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(split, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(arithmetic, [], [refute_0_395])).

cnf(refute_0_397, plain,
    (skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_398, plain,
    (skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) < cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_399, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     1 + skoSC1 * (skoSC1 * -1/2) |
     (1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) |
     2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_400, plain,
    ((2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <
     cos(skoSC1) |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_401, plain,
    (1 + skoSC1 * (skoSC1 * -1/2) <=
     (2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2))) /
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2))),
    inference(resolve, [], [refute_0_390, refute_0_400])).

cnf(refute_0_402, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2)) <=
     2 + skoSC1 * (skoSC1 * (-6 + skoSC1 * -2)) |
     2 + skoSC1 * (10 + skoSC1 * 6) + cos(skoSC1) * (2 + skoSC1 * 2) <= 0),
    inference(resolve, [], [refute_0_401, refute_0_399])).

cnf(refute_0_403, plain,
    (cos(skoSC1) * (2 + skoSC1 * 2) <= -2 + skoSC1 * (-10 + skoSC1 * -6) |
     cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3)))),
    inference(arithmetic, [], [refute_0_402])).

cnf(refute_0_404, plain,
    (cos(skoSC1) * (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3)))),
    inference(resolve, [], [refute_0_403, refute_0_389])).

cnf(refute_0_405, plain,
    (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0 |
     cos(skoSC1) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(split, [], [refute_0_404])).

cnf(refute_0_406, plain,
    (skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     cos(skoSC1) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(arithmetic, [], [refute_0_405])).

cnf(refute_0_407, plain,
    (skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     1 + skoSC1 * (skoSC1 * -1/2) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)))),
    inference(resolve, [], [refute_0_406, refute_0_398])).

cnf(refute_0_408, plain,
    ((1 + skoSC1 * (skoSC1 * -1/2)) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_407, refute_0_397])).

cnf(refute_0_409, plain,
    (skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2 |
     skoSC1 *
     (12 +
      skoSC1 * (9 + skoSC1 * (-5 + skoSC1 * (-5/2 + skoSC1 * 1/2)))) <=
     -2), inference(arithmetic, [], [refute_0_408])).

cnf(refute_0_410, plain,
    (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24))) <
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) |
     0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_411, plain,
    (cos(skoSC1) <
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) |
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_412, plain,
    (0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <= cos(skoSC1)),
    inference(split, [], [refute_0_404])).

cnf(refute_0_413, plain,
    (skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <= cos(skoSC1)),
    inference(arithmetic, [], [refute_0_412])).

cnf(refute_0_414, plain,
    (skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) /
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_413, refute_0_411])).

cnf(refute_0_415, plain,
    (0 <= 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))) <=
     skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3))) |
     skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2),
    inference(resolve, [], [refute_0_414, refute_0_410])).

cnf(refute_0_416, plain,
    (skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) <= 2 |
     skoSC1 *
     (12 +
      skoSC1 *
      (9 +
       skoSC1 *
       (-5 +
        skoSC1 *
        (-29/12 +
         skoSC1 * (7/12 + skoSC1 * (-1/24 + skoSC1 * -1/24)))))) <= -2),
    inference(arithmetic, [], [refute_0_415])).

cnf(refute_0_417, plain,
    (2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) < 0 |
     0 < 2 + skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     0 <= skoSC1 * (-10 + skoSC1 * (-11 + skoSC1 * (3 + skoSC1 * 3)))),
    inference(split, [], [refute_0_404])).

cnf(refute_0_418, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     skoSC1 * (10 + skoSC1 * (11 + skoSC1 * (-3 + skoSC1 * -3))) <= 0),
    inference(arithmetic, [], [refute_0_417])).

cnf(refute_0_419, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) |
     0 < skoSC1 * (10 + skoSC1 * (11 + skoSC1 * (-3 + skoSC1 * -3)))),
    inference(decision, [], [refute_0_67, refute_0_35])).

cnf(refute_0_420, plain,
    (2 < skoSC1 * (-2 + skoSC1 * (1 + skoSC1)) |
     -2 < skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1))),
    inference(resolve, [], [refute_0_418, refute_0_419])).

cnf(refute_0_421, plain,
    (-2 <
     skoSC1 *
     (12 + skoSC1 * (9 + skoSC1 * (-5 + skoSC1 * (-5/2 + skoSC1 * 1/2)))) |
     skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(decision, [],
              [refute_0_416, refute_0_420, refute_0_67, refute_0_35])).

cnf(refute_0_422, plain,
    (skoSC1 * (2 + skoSC1 * (-1 + skoSC1 * -1)) <= -2),
    inference(resolve, [], [refute_0_409, refute_0_421])).

cnf(refute_0_423, plain,
    (2 <
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400))))))))))),
    inference(decision, [],
              [refute_0_422, refute_0_416, refute_0_420, refute_0_67,
               refute_0_35])).

cnf(refute_0_424, plain,
    (skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     cos(skoSC1)), inference(resolve, [], [refute_0_396, refute_0_423])).

cnf(refute_0_425, plain,
    (skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800))))))))))) /
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))),
    inference(resolve, [], [refute_0_424, refute_0_5])).

cnf(refute_0_426, plain,
    (0 <=
     2 +
     skoSC1 *
     (2 +
      skoSC1 *
      (-1 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (1/12 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (-1/360 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (1/20160 +
             skoSC1 *
             (1/20160 +
              skoSC1 * (-1/1814400 + skoSC1 * -1/1814400)))))))))) |
     (1 + skoSC1 * (skoSC1 * (-1/2 + skoSC1 * (skoSC1 * 1/24)))) *
     (2 +
      skoSC1 *
      (2 +
       skoSC1 *
       (-1 +
        skoSC1 *
        (-1 +
         skoSC1 *
         (1/12 +
          skoSC1 *
          (1/12 +
           skoSC1 *
           (-1/360 +
            skoSC1 *
            (-1/360 +
             skoSC1 *
             (1/20160 +
              skoSC1 *
              (1/20160 +
               skoSC1 * (-1/1814400 + skoSC1 * -1/1814400))))))))))) <=
     skoSC1 *
     (-10 +
      skoSC1 *
      (-11 +
       skoSC1 *
       (3 +
        skoSC1 *
        (35/12 +
         skoSC1 *
         (-5/12 +
          skoSC1 *
          (-89/360 +
           skoSC1 *
           (1/72 +
            skoSC1 *
            (167/20160 +
             skoSC1 *
             (-1/4032 +
              skoSC1 *
              (-269/1814400 +
               skoSC1 * (1/362880 + skoSC1 * 1/604800)))))))))))),
    inference(resolve, [], [refute_0_425, refute_0_0])).

cnf(refute_0_427, plain,
    (skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2 |
     skoSC1 *
     (12 +
      skoSC1 *
      (9 +
       skoSC1 *
       (-5 +
        skoSC1 *
        (-9/4 +
         skoSC1 *
         (13/12 +
          skoSC1 *
          (29/180 +
           skoSC1 *
           (-1/10 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (13/2520 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-29/201600 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (17/7257600 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) <= -2),
    inference(arithmetic, [], [refute_0_426])).

cnf(refute_0_428, plain,
    (-2 <
     skoSC1 *
     (12 +
      skoSC1 *
      (9 +
       skoSC1 *
       (-5 +
        skoSC1 *
        (-9/4 +
         skoSC1 *
         (13/12 +
          skoSC1 *
          (29/180 +
           skoSC1 *
           (-1/10 +
            skoSC1 *
            (-17/5040 +
             skoSC1 *
             (13/2520 +
              skoSC1 *
              (13/1814400 +
               skoSC1 *
               (-29/201600 +
                skoSC1 *
                (1/1451520 +
                 skoSC1 *
                 (17/7257600 +
                  skoSC1 *
                  (-1/43545600 + skoSC1 * -1/43545600)))))))))))))) |
     skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(decision, [],
              [refute_0_422, refute_0_416, refute_0_420, refute_0_67,
               refute_0_35])).

cnf(refute_0_429, plain,
    (skoSC1 *
     (-2 +
      skoSC1 *
      (1 +
       skoSC1 *
       (1 +
        skoSC1 *
        (-1/12 +
         skoSC1 *
         (-1/12 +
          skoSC1 *
          (1/360 +
           skoSC1 *
           (1/360 +
            skoSC1 *
            (-1/20160 +
             skoSC1 *
             (-1/20160 +
              skoSC1 * (1/1814400 + skoSC1 * 1/1814400)))))))))) <= 2),
    inference(resolve, [], [refute_0_427, refute_0_428])).

cnf(refute_0_430, plain, ($false),
    inference(resolve, [], [refute_0_429, refute_0_423])).
% SZS output end CNFRefutation for Arthan1.tptp
