% SZS output start CNFRefutation for asin-5.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(arcsin_lower_bound_pos_2, axiom,
    (X < 0 | 1 < X | ~ lgen(R, Y, 3/40 * X ^ 5 + 1/6 * X ^ 3 + X) |
     lgen(R, Y, arcsin(X)))).

cnf(arcsin_lower_bound_pos_5, axiom,
    (X < 0 | 1 < X |
     ~ lgen(R, Y,
            63/2816 * X ^ 11 + 35/1152 * X ^ 9 + 5/112 * X ^ 7 +
            3/40 * X ^ 5 + 1/6 * X ^ 3 + X) | lgen(R, Y, arcsin(X)))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_2, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            8 * X * (X + 1) * (X ^ 2 + 6 * X + 1) /
            (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1)) |
     lgen(R, Y, sqrt(X)))).

fof(asin_5, conjecture,
    (! [X] :
       ((0 < X & X < 59/100) =>
        6 * (sqrt(1 + X) - sqrt(1 - X)) / (4 + sqrt(1 + X) + sqrt(1 - X)) <
        arcsin(X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 59/100) =>
        6 * (sqrt(1 + X) - sqrt(1 - X)) / (4 + sqrt(1 + X) + sqrt(1 - X)) <
        arcsin(X))), inference(strip, [], [asin_5])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 59/100) =>
          6 * (sqrt(1 + X) - sqrt(1 - X)) /
          (4 + sqrt(1 + X) + sqrt(1 - X)) < arcsin(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 59/100 &
        arcsin(X) <=
        6 * (sqrt(1 + X) - sqrt(1 - X)) /
        (4 + sqrt(1 + X) + sqrt(1 - X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 59/100 & 0 < skoXC1 &
     arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 59/100),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704)))))))))))))) <
     (-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-3112/105 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (-1669/2520 +
               skoXC1 *
               (skoXC1 *
                (-7621/6160 +
                 skoXC1 *
                 (skoXC1 *
                  (-8132881/221760 +
                   skoXC1 *
                   (skoXC1 *
                    (2112983/88704 +
                     skoXC1 *
                     (skoXC1 *
                      (-18109/9216 +
                       skoXC1 * (skoXC1 * 63/22528))))))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -12288 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (18432 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-7104 +
         skoXC1 *
         (4128/5 +
          skoXC1 *
          (528 +
           skoXC1 *
           (-3112/105 +
            skoXC1 *
            (-3/4 +
             skoXC1 *
             (-1669/2520 +
              skoXC1 *
              (skoXC1 *
               (-7621/6160 +
                skoXC1 *
                (skoXC1 *
                 (-8132881/221760 +
                  skoXC1 *
                  (skoXC1 *
                   (2112983/88704 +
                    skoXC1 *
                    (skoXC1 *
                     (-18109/9216 +
                      skoXC1 * (skoXC1 * 63/22528)))))))))))))))))) <=
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) |
     0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 *
           (751/42 +
            skoXC1 *
            (-484/35 +
             skoXC1 *
             (27989/2520 +
              skoXC1 *
              (-580/63 +
               skoXC1 *
               (86147/11088 +
                skoXC1 *
                (-2653/396 +
                 skoXC1 *
                 (45745/12672 +
                  skoXC1 * (-63/88 + skoXC1 * 63/2816)))))))))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    (sqrt(X_000048) < X_000049 | ~ lgen(1, sqrt(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (X_000048 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000048 ^ 4 + 28 * X_000048 ^ 3 + 70 * X_000048 ^ 2 +
             28 * X_000048 + 1) /
            ((X_000048 + 1) * (X_000048 ^ 2 + 6 * X_000048 + 1)),
            X_000049) | lgen(1, sqrt(X_000048), X_000049)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_4, plain,
    (sqrt(X_000048) < X_000049 | X_000048 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000048 ^ 4 + 28 * X_000048 ^ 3 + 70 * X_000048 ^ 2 +
             28 * X_000048 + 1) /
            ((X_000048 + 1) * (X_000048 ^ 2 + 6 * X_000048 + 1)),
            X_000049)), inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (sqrt(X_000048) < X_000049 | X_000048 < 0 |
     X_000049 <=
     (1/8 +
      X_000048 *
      (7/2 + X_000048 * (35/4 + X_000048 * (7/2 + X_000048 * 1/8)))) /
     (1 + X_000048 * (7 + X_000048 * (7 + X_000048)))),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) <
     (768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 *
            (2882/105 +
             skoXC1 *
             (-731/35 +
              skoXC1 *
              (5239/315 +
               skoXC1 *
               (-860/63 +
                skoXC1 *
                (63671/5544 +
                 skoXC1 *
                 (-15575/1584 +
                  skoXC1 *
                  (1253/198 +
                   skoXC1 * (-63/32 + skoXC1 * 63/352))))))))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     768 +
     skoXC1 *
     (-1792 +
      skoXC1 *
      (1312 +
       skoXC1 *
       (-704/3 +
        skoXC1 *
        (-280/3 +
         skoXC1 *
         (944/15 +
          skoXC1 *
          (-116/3 +
           skoXC1 *
           (2882/105 +
            skoXC1 *
            (-731/35 +
             skoXC1 *
             (5239/315 +
              skoXC1 *
              (-860/63 +
               skoXC1 *
               (63671/5544 +
                skoXC1 *
                (-15575/1584 +
                 skoXC1 *
                 (1253/198 +
                  skoXC1 * (-63/32 + skoXC1 * 63/352)))))))))))))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    ((skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     6 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_9, plain,
    (X_000063 < sqrt(X_000062) | ~ lgen(1, X_000063, sqrt(X_000062))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_10, plain,
    (X_000062 < 0 |
     ~ lgen(1, X_000063,
            8 * X_000062 * (X_000062 + 1) *
            (X_000062 ^ 2 + 6 * X_000062 + 1) /
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1)) | lgen(1, X_000063, sqrt(X_000062))),
    inference(subst, [], [sqrt_lower_bound_2])).

cnf(refute_0_11, plain,
    (X_000062 < 0 | X_000063 < sqrt(X_000062) |
     ~ lgen(1, X_000063,
            8 * X_000062 * (X_000062 + 1) *
            (X_000062 ^ 2 + 6 * X_000062 + 1) /
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1))),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (X_000062 < 0 | X_000063 < sqrt(X_000062) |
     X_000062 * (8 + X_000062 * (56 + X_000062 * (56 + X_000062 * 8))) /
     (1 +
      X_000062 * (28 + X_000062 * (70 + X_000062 * (28 + X_000062)))) <=
     X_000063), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (1 + skoXC1 * -1 < 0 |
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <
     sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_15, plain,
    (arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_16, plain,
    (arcsin(skoXC1) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (X_000119 < arcsin(X_000118) | ~ lgen(1, X_000119, arcsin(X_000118))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_18, plain,
    (1 < X_000118 | X_000118 < 0 |
     ~ lgen(1, X_000119,
            63/2816 * X_000118 ^ 11 + 35/1152 * X_000118 ^ 9 +
            5/112 * X_000118 ^ 7 + 3/40 * X_000118 ^ 5 +
            1/6 * X_000118 ^ 3 + X_000118) |
     lgen(1, X_000119, arcsin(X_000118))),
    inference(subst, [], [arcsin_lower_bound_pos_5])).

cnf(refute_0_19, plain,
    (1 < X_000118 | X_000118 < 0 | X_000119 < arcsin(X_000118) |
     ~ lgen(1, X_000119,
            63/2816 * X_000118 ^ 11 + 35/1152 * X_000118 ^ 9 +
            5/112 * X_000118 ^ 7 + 3/40 * X_000118 ^ 5 +
            1/6 * X_000118 ^ 3 + X_000118)),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (1 < X_000118 | X_000118 < 0 | X_000119 < arcsin(X_000118) |
     X_000118 *
     (1 +
      X_000118 *
      (X_000118 *
       (1/6 +
        X_000118 *
        (X_000118 *
         (3/40 +
          X_000118 *
          (X_000118 *
           (5/112 +
            X_000118 *
            (X_000118 *
             (35/1152 + X_000118 * (X_000118 * 63/2816)))))))))) <=
     X_000119), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (skoXC1 < 0 |
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) < arcsin(skoXC1) |
     1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 < 0 | 1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_16, refute_0_21])).

cnf(refute_0_23, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_24, plain, (skoXC1 < 59/100),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_25, plain, (1 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_25, refute_0_22])).

cnf(refute_0_27, plain, (skoXC1 <= 1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_28, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_28, refute_0_14])).

cnf(refute_0_30, plain,
    (sqrt(1 + skoXC1 * -1) *
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (-10 + skoXC1 * 6 < (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) |
     2 + skoXC1 <= 0 |
     4 + skoXC1 * (2 + skoXC1 * -2) <= (-10 + skoXC1 * 6) * (2 + skoXC1)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_32, plain,
    ((-10 + skoXC1 * 6) / (2 + skoXC1 * -1) <
     (2 + skoXC1 * 2) / (2 + skoXC1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) * (2 + skoXC1 * -1) <=
     -10 + skoXC1 * 6 | 2 + skoXC1 * -1 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_33, plain,
    (X_000042 < sqrt(X_000041) | ~ lgen(1, X_000042, sqrt(X_000041))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_34, plain,
    (X_000041 < 0 | ~ lgen(1, X_000042, 2 * X_000041 / (X_000041 + 1)) |
     lgen(1, X_000042, sqrt(X_000041))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_35, plain,
    (X_000041 < 0 | X_000042 < sqrt(X_000041) |
     ~ lgen(1, X_000042, 2 * X_000041 / (X_000041 + 1))),
    inference(resolve, [], [refute_0_34, refute_0_33])).

cnf(refute_0_36, plain,
    (X_000041 < 0 | X_000042 < sqrt(X_000041) |
     X_000041 * 2 / (1 + X_000041) <= X_000042),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (1 + skoXC1 < 0 |
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1) < sqrt(1 + skoXC1) |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < (2 + skoXC1 * -2) / (2 + skoXC1 * -1) |
     2 + skoXC1 * -2 <= (-4 + sqrt(1 + skoXC1) * -1) * (2 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_39, plain,
    (1 + skoXC1 * -1 < 0 |
     -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     -4 + sqrt(1 + skoXC1) * -1), inference(subst, [], [refute_0_36])).

cnf(refute_0_40, plain,
    (1 < skoXC1 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(resolve, [], [refute_0_27, refute_0_40])).

cnf(refute_0_42, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     2 + skoXC1 * -2 <= (-4 + sqrt(1 + skoXC1) * -1) * (2 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_41, refute_0_38])).

cnf(refute_0_43, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     sqrt(1 + skoXC1) * (2 + skoXC1 * -1) <= -10 + skoXC1 * 6),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain, (skoXC1 < 2),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_45, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     sqrt(1 + skoXC1) * (2 + skoXC1 * -1) <= -10 + skoXC1 * 6),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0 |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(split, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_47, refute_0_44])).

cnf(refute_0_49, plain,
    (1 + skoXC1 < 0 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_48, refute_0_37])).

cnf(refute_0_50, plain,
    (skoXC1 < -1 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain, (-1 <= skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_52, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_51, refute_0_50])).

cnf(refute_0_53, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) * (2 + skoXC1 * -1) <=
     -10 + skoXC1 * 6 | 2 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_52, refute_0_32])).

cnf(refute_0_54, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) <= -10 + skoXC1 * 6),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) <= -10 + skoXC1 * 6),
    inference(resolve, [], [refute_0_54, refute_0_44])).

cnf(refute_0_56, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 + skoXC1 <= 0 |
     4 + skoXC1 * (2 + skoXC1 * -2) <= (-10 + skoXC1 * 6) * (2 + skoXC1)),
    inference(resolve, [], [refute_0_55, refute_0_31])).

cnf(refute_0_57, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     skoXC1 * (skoXC1 * -8) <= -24 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain, (-24 < skoXC1 * (skoXC1 * -8) | skoXC1 <= -2),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_59, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain, (-2 < skoXC1),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_61, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (sqrt(1 + skoXC1 * -1) *
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(resolve, [], [refute_0_30, refute_0_61])).

cnf(refute_0_63, plain,
    (6 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0 |
     sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(split, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     -6 |
     sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (-768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10)))))))) <
     (-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-1816/15 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (8911/120 +
               skoXC1 *
               (skoXC1 *
                (-1579/240 + skoXC1 * (skoXC1 * 3/320))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -12288 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (18432 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-7104 +
         skoXC1 *
         (4128/5 +
          skoXC1 *
          (528 +
           skoXC1 *
           (-1816/15 +
            skoXC1 *
            (-3/4 +
             skoXC1 *
             (8911/120 +
              skoXC1 *
              (skoXC1 *
               (-1579/240 + skoXC1 * (skoXC1 * 3/320)))))))))))) <=
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_66, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) |
     0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40)))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_67, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_68, plain,
    (skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40))))) <
     (768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 * (326/15 + skoXC1 * (-33/5 + skoXC1 * 3/5))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     768 +
     skoXC1 *
     (-1792 +
      skoXC1 *
      (1312 +
       skoXC1 *
       (-704/3 +
        skoXC1 *
        (-280/3 +
         skoXC1 *
         (944/15 +
          skoXC1 *
          (-116/3 +
           skoXC1 * (326/15 + skoXC1 * (-33/5 + skoXC1 * 3/5)))))))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_69, plain,
    ((skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40))))) |
     6 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_70, plain,
    (1 + skoXC1 * -1 < 0 |
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <
     sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_71, plain,
    ((sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_72, plain,
    (X_000091 < arcsin(X_000090) | ~ lgen(1, X_000091, arcsin(X_000090))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_73, plain,
    (1 < X_000090 | X_000090 < 0 |
     ~ lgen(1, X_000091,
            3/40 * X_000090 ^ 5 + 1/6 * X_000090 ^ 3 + X_000090) |
     lgen(1, X_000091, arcsin(X_000090))),
    inference(subst, [], [arcsin_lower_bound_pos_2])).

cnf(refute_0_74, plain,
    (1 < X_000090 | X_000090 < 0 | X_000091 < arcsin(X_000090) |
     ~ lgen(1, X_000091,
            3/40 * X_000090 ^ 5 + 1/6 * X_000090 ^ 3 + X_000090)),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    (1 < X_000090 | X_000090 < 0 | X_000091 < arcsin(X_000090) |
     X_000090 *
     (1 + X_000090 * (X_000090 * (1/6 + X_000090 * (X_000090 * 3/40)))) <=
     X_000091), inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 < 0 |
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) < arcsin(skoXC1) |
     1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 < 0 | 1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_16, refute_0_76])).

cnf(refute_0_78, plain,
    (1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_25, refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_27, refute_0_78])).

cnf(refute_0_80, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_79, refute_0_71])).

cnf(refute_0_81, plain,
    (sqrt(1 + skoXC1 * -1) *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40))))) |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (sqrt(1 + skoXC1 * -1) *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))),
    inference(resolve, [], [refute_0_81, refute_0_61])).

cnf(refute_0_83, plain,
    (6 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     0 |
     sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(split, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     -6 |
     sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (-6 <
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_86, plain,
    (sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (1 + skoXC1 * -1 < 0 |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(resolve, [], [refute_0_86, refute_0_70])).

cnf(refute_0_88, plain,
    (1 < skoXC1 |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) /
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))))),
    inference(resolve, [], [refute_0_27, refute_0_88])).

cnf(refute_0_90, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (6 +
      skoXC1 *
      (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40))))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40))))) |
     6 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     0), inference(resolve, [], [refute_0_89, refute_0_69])).

cnf(refute_0_91, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * (1/6 + skoXC1 * (skoXC1 * 3/40)))) <=
     -6 |
     (768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 * (326/15 + skoXC1 * (-33/5 + skoXC1 * 3/5))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    ((768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 * (326/15 + skoXC1 * (-33/5 + skoXC1 * 3/5))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     skoXC1 *
     (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))),
    inference(resolve, [], [refute_0_91, refute_0_85])).

cnf(refute_0_93, plain,
    (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     768 +
     skoXC1 *
     (-1792 +
      skoXC1 *
      (1312 +
       skoXC1 *
       (-704/3 +
        skoXC1 *
        (-280/3 +
         skoXC1 *
         (944/15 +
          skoXC1 *
          (-116/3 +
           skoXC1 * (326/15 + skoXC1 * (-33/5 + skoXC1 * 3/5)))))))) <=
     (skoXC1 *
      (-4 + skoXC1 * (skoXC1 * (-2/3 + skoXC1 * (skoXC1 * -3/10)))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (skoXC1 * -3/40)))))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(resolve, [], [refute_0_92, refute_0_68])).

cnf(refute_0_94, plain,
    (skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <= -128 |
     sqrt(1 + skoXC1) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (-128 < skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_96, plain,
    (sqrt(1 + skoXC1) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(resolve, [], [refute_0_94, refute_0_95])).

cnf(refute_0_97, plain,
    (0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40)))))))) |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 * (-73/6 + skoXC1 * (12/5 + skoXC1 * -3/40)))))))) <=
     -768 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (-768 <
     skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 * (-73/6 + skoXC1 * (12/5 + skoXC1 * -3/40))))))))),
    inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_100, plain,
    ((-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (1 + skoXC1 < 0 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_100, refute_0_67])).

cnf(refute_0_102, plain,
    (skoXC1 < -1 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    ((-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_51, refute_0_102])).

cnf(refute_0_104, plain,
    (0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40)))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 * (73/6 + skoXC1 * (-12/5 + skoXC1 * 3/40))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(resolve, [], [refute_0_103, refute_0_66])).

cnf(refute_0_105, plain,
    (skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 * (-73/6 + skoXC1 * (12/5 + skoXC1 * -3/40)))))))) <=
     -768 |
     (-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-1816/15 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (8911/120 +
               skoXC1 *
               (skoXC1 *
                (-1579/240 + skoXC1 * (skoXC1 * 3/320))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    ((-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-1816/15 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (8911/120 +
               skoXC1 *
               (skoXC1 *
                (-1579/240 + skoXC1 * (skoXC1 * 3/320))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))),
    inference(resolve, [], [refute_0_105, refute_0_99])).

cnf(refute_0_107, plain,
    (-12288 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (18432 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-7104 +
         skoXC1 *
         (4128/5 +
          skoXC1 *
          (528 +
           skoXC1 *
           (-1816/15 +
            skoXC1 *
            (-3/4 +
             skoXC1 *
             (8911/120 +
              skoXC1 *
              (skoXC1 *
               (-1579/240 + skoXC1 * (skoXC1 * 3/320)))))))))))) <=
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 * (-352/5 + skoXC1 * (81/5 + skoXC1 * -9/10))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_106, refute_0_65])).

cnf(refute_0_108, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (64/15 +
          skoXC1 *
          (-32/15 +
           skoXC1 *
           (-8248/15 +
            skoXC1 *
            (5473/20 +
             skoXC1 *
             (32047/120 +
              skoXC1 *
              (-70 +
               skoXC1 *
               (-3307/240 + skoXC1 * (9/10 + skoXC1 * 3/320)))))))))))) <=
     0 | skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (64/15 +
          skoXC1 *
          (-32/15 +
           skoXC1 *
           (-8248/15 +
            skoXC1 *
            (5473/20 +
             skoXC1 *
             (32047/120 +
              skoXC1 *
              (-70 +
               skoXC1 *
               (-3307/240 + skoXC1 * (9/10 + skoXC1 * 3/320)))))))))))) <=
     0), inference(decision, [], [refute_0_24, refute_0_23])).

cnf(refute_0_110, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (64/15 +
          skoXC1 *
          (-32/15 +
           skoXC1 *
           (-8248/15 +
            skoXC1 *
            (5473/20 +
             skoXC1 *
             (32047/120 +
              skoXC1 *
              (-70 +
               skoXC1 *
               (-3307/240 + skoXC1 * (9/10 + skoXC1 * 3/320)))))))))))) <=
     0), inference(resolve, [], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain,
    (-6 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))),
    inference(decision, [], [refute_0_110, refute_0_24, refute_0_23])).

cnf(refute_0_112, plain,
    (sqrt(1 + skoXC1 * -1) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_64, refute_0_111])).

cnf(refute_0_113, plain,
    (1 + skoXC1 * -1 < 0 |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_112, refute_0_13])).

cnf(refute_0_114, plain,
    (1 < skoXC1 |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_27, refute_0_114])).

cnf(refute_0_116, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (6 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     6 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0), inference(resolve, [], [refute_0_115, refute_0_8])).

cnf(refute_0_117, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     -6 |
     (768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 *
            (2882/105 +
             skoXC1 *
             (-731/35 +
              skoXC1 *
              (5239/315 +
               skoXC1 *
               (-860/63 +
                skoXC1 *
                (63671/5544 +
                 skoXC1 *
                 (-15575/1584 +
                  skoXC1 *
                  (1253/198 +
                   skoXC1 * (-63/32 + skoXC1 * 63/352))))))))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    ((768 +
      skoXC1 *
      (-1792 +
       skoXC1 *
       (1312 +
        skoXC1 *
        (-704/3 +
         skoXC1 *
         (-280/3 +
          skoXC1 *
          (944/15 +
           skoXC1 *
           (-116/3 +
            skoXC1 *
            (2882/105 +
             skoXC1 *
             (-731/35 +
              skoXC1 *
              (5239/315 +
               skoXC1 *
               (-860/63 +
                skoXC1 *
                (63671/5544 +
                 skoXC1 *
                 (-15575/1584 +
                  skoXC1 *
                  (1253/198 +
                   skoXC1 * (-63/32 + skoXC1 * 63/352))))))))))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (6 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(resolve, [], [refute_0_117, refute_0_111])).

cnf(refute_0_119, plain,
    (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     768 +
     skoXC1 *
     (-1792 +
      skoXC1 *
      (1312 +
       skoXC1 *
       (-704/3 +
        skoXC1 *
        (-280/3 +
         skoXC1 *
         (944/15 +
          skoXC1 *
          (-116/3 +
           skoXC1 *
           (2882/105 +
            skoXC1 *
            (-731/35 +
             skoXC1 *
             (5239/315 +
              skoXC1 *
              (-860/63 +
               skoXC1 *
               (63671/5544 +
                skoXC1 *
                (-15575/1584 +
                 skoXC1 *
                 (1253/198 +
                  skoXC1 * (-63/32 + skoXC1 * 63/352)))))))))))))) <=
     (skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (6 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(resolve, [], [refute_0_118, refute_0_7])).

cnf(refute_0_120, plain,
    (skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <= -128 |
     sqrt(1 + skoXC1) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (sqrt(1 + skoXC1) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(resolve, [], [refute_0_120, refute_0_95])).

cnf(refute_0_122, plain,
    (0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 *
           (751/42 +
            skoXC1 *
            (-484/35 +
             skoXC1 *
             (27989/2520 +
              skoXC1 *
              (-580/63 +
               skoXC1 *
               (86147/11088 +
                skoXC1 *
                (-2653/396 +
                 skoXC1 *
                 (45745/12672 +
                  skoXC1 * (-63/88 + skoXC1 * 63/2816)))))))))))))) |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 *
           (-751/42 +
            skoXC1 *
            (484/35 +
             skoXC1 *
             (-27989/2520 +
              skoXC1 *
              (580/63 +
               skoXC1 *
               (-86147/11088 +
                skoXC1 *
                (2653/396 +
                 skoXC1 *
                 (-45745/12672 +
                  skoXC1 * (63/88 + skoXC1 * -63/2816)))))))))))))) <=
     -768 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (-768 <
     skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 *
           (-751/42 +
            skoXC1 *
            (484/35 +
             skoXC1 *
             (-27989/2520 +
              skoXC1 *
              (580/63 +
               skoXC1 *
               (-86147/11088 +
                skoXC1 *
                (2653/396 +
                 skoXC1 *
                 (-45745/12672 +
                  skoXC1 * (63/88 + skoXC1 * -63/2816))))))))))))))),
    inference(decision, [], [refute_0_110, refute_0_24, refute_0_23])).

cnf(refute_0_125, plain,
    ((-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (1 + skoXC1 < 0 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_125, refute_0_6])).

cnf(refute_0_127, plain,
    (skoXC1 < -1 |
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    ((-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) /
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_51, refute_0_127])).

cnf(refute_0_129, plain,
    (0 <=
     -768 +
     skoXC1 *
     (1664 +
      skoXC1 *
      (-1216 +
       skoXC1 *
       (1120/3 +
        skoXC1 *
        (-242/3 +
         skoXC1 *
         (559/15 +
          skoXC1 *
          (-368/15 +
           skoXC1 *
           (751/42 +
            skoXC1 *
            (-484/35 +
             skoXC1 *
             (27989/2520 +
              skoXC1 *
              (-580/63 +
               skoXC1 *
               (86147/11088 +
                skoXC1 *
                (-2653/396 +
                 skoXC1 *
                 (45745/12672 +
                  skoXC1 * (-63/88 + skoXC1 * 63/2816)))))))))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-768 +
      skoXC1 *
      (1664 +
       skoXC1 *
       (-1216 +
        skoXC1 *
        (1120/3 +
         skoXC1 *
         (-242/3 +
          skoXC1 *
          (559/15 +
           skoXC1 *
           (-368/15 +
            skoXC1 *
            (751/42 +
             skoXC1 *
             (-484/35 +
              skoXC1 *
              (27989/2520 +
               skoXC1 *
               (-580/63 +
                skoXC1 *
                (86147/11088 +
                 skoXC1 *
                 (-2653/396 +
                  skoXC1 *
                  (45745/12672 +
                   skoXC1 * (-63/88 + skoXC1 * 63/2816))))))))))))))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(resolve, [], [refute_0_128, refute_0_1])).

cnf(refute_0_130, plain,
    (skoXC1 *
     (-1664 +
      skoXC1 *
      (1216 +
       skoXC1 *
       (-1120/3 +
        skoXC1 *
        (242/3 +
         skoXC1 *
         (-559/15 +
          skoXC1 *
          (368/15 +
           skoXC1 *
           (-751/42 +
            skoXC1 *
            (484/35 +
             skoXC1 *
             (-27989/2520 +
              skoXC1 *
              (580/63 +
               skoXC1 *
               (-86147/11088 +
                skoXC1 *
                (2653/396 +
                 skoXC1 *
                 (-45745/12672 +
                  skoXC1 * (63/88 + skoXC1 * -63/2816)))))))))))))) <=
     -768 |
     (-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-3112/105 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (-1669/2520 +
               skoXC1 *
               (skoXC1 *
                (-7621/6160 +
                 skoXC1 *
                 (skoXC1 *
                  (-8132881/221760 +
                   skoXC1 *
                   (skoXC1 *
                    (2112983/88704 +
                     skoXC1 *
                     (skoXC1 *
                      (-18109/9216 +
                       skoXC1 * (skoXC1 * 63/22528))))))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(arithmetic, [], [refute_0_129])).

cnf(refute_0_131, plain,
    ((-12288 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (18432 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-7104 +
          skoXC1 *
          (4128/5 +
           skoXC1 *
           (528 +
            skoXC1 *
            (-3112/105 +
             skoXC1 *
             (-3/4 +
              skoXC1 *
              (-1669/2520 +
               skoXC1 *
               (skoXC1 *
                (-7621/6160 +
                 skoXC1 *
                 (skoXC1 *
                  (-8132881/221760 +
                   skoXC1 *
                   (skoXC1 *
                    (2112983/88704 +
                     skoXC1 *
                     (skoXC1 *
                      (-18109/9216 +
                       skoXC1 * (skoXC1 * 63/22528))))))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -768 +
     skoXC1 *
     (1280 +
      skoXC1 *
      (-288 +
       skoXC1 *
       (-1472/3 +
        skoXC1 *
        (392 +
         skoXC1 *
         (-212 +
          skoXC1 *
          (684/5 +
           skoXC1 *
           (-3464/35 +
            skoXC1 *
            (381/5 +
             skoXC1 *
             (-38467/630 +
              skoXC1 *
              (1060/21 +
               skoXC1 *
               (-78655/1848 +
                skoXC1 *
                (6447/176 +
                 skoXC1 *
                 (-21931/1056 +
                  skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))),
    inference(resolve, [], [refute_0_130, refute_0_124])).

cnf(refute_0_132, plain,
    (-12288 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (18432 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-7104 +
         skoXC1 *
         (4128/5 +
          skoXC1 *
          (528 +
           skoXC1 *
           (-3112/105 +
            skoXC1 *
            (-3/4 +
             skoXC1 *
             (-1669/2520 +
              skoXC1 *
              (skoXC1 *
               (-7621/6160 +
                skoXC1 *
                (skoXC1 *
                 (-8132881/221760 +
                  skoXC1 *
                  (skoXC1 *
                   (2112983/88704 +
                    skoXC1 *
                    (skoXC1 *
                     (-18109/9216 +
                      skoXC1 * (skoXC1 * 63/22528)))))))))))))))))) <=
     (-768 +
      skoXC1 *
      (1280 +
       skoXC1 *
       (-288 +
        skoXC1 *
        (-1472/3 +
         skoXC1 *
         (392 +
          skoXC1 *
          (-212 +
           skoXC1 *
           (684/5 +
            skoXC1 *
            (-3464/35 +
             skoXC1 *
             (381/5 +
              skoXC1 *
              (-38467/630 +
               skoXC1 *
               (1060/21 +
                skoXC1 *
                (-78655/1848 +
                 skoXC1 *
                 (6447/176 +
                  skoXC1 *
                  (-21931/1056 +
                   skoXC1 * (1701/352 + skoXC1 * -189/704))))))))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_131, refute_0_0])).

cnf(refute_0_133, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (64/15 +
          skoXC1 *
          (-32/15 +
           skoXC1 *
           (-136/105 +
            skoXC1 *
            (-89/140 +
             skoXC1 *
             (47/120 +
              skoXC1 *
              (-110/21 +
               skoXC1 *
               (150643/55440 +
                skoXC1 *
                (-57523/6930 +
                 skoXC1 *
                 (-46209601/221760 +
                  skoXC1 *
                  (8179/84 +
                   skoXC1 *
                   (7379111/88704 +
                    skoXC1 *
                    (-22295/1056 +
                     skoXC1 *
                     (-416927/101376 +
                      skoXC1 *
                      (189/704 + skoXC1 * 63/22528)))))))))))))))))) <= 0 |
     skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (64/15 +
          skoXC1 *
          (-32/15 +
           skoXC1 *
           (-136/105 +
            skoXC1 *
            (-89/140 +
             skoXC1 *
             (47/120 +
              skoXC1 *
              (-110/21 +
               skoXC1 *
               (150643/55440 +
                skoXC1 *
                (-57523/6930 +
                 skoXC1 *
                 (-46209601/221760 +
                  skoXC1 *
                  (8179/84 +
                   skoXC1 *
                   (7379111/88704 +
                    skoXC1 *
                    (-22295/1056 +
                     skoXC1 *
                     (-416927/101376 +
                      skoXC1 *
                      (189/704 + skoXC1 * 63/22528)))))))))))))))))) |
     skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(decision, [], [refute_0_110, refute_0_24, refute_0_23])).

cnf(refute_0_135, plain, (skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain, (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1))),
    inference(decision, [], [refute_0_110, refute_0_24, refute_0_23])).

cnf(refute_0_137, plain, ($false),
    inference(resolve, [], [refute_0_135, refute_0_136])).
% SZS output end CNFRefutation for asin-5.tptp
