% SZS output start CNFRefutation for asin-5-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(arcsin_lower_bound_pos_1, axiom,
    (X < 0 | 1 < X | ~ lgen(R, Y, 1/6 * X ^ 3 + X) |
     lgen(R, Y, arcsin(X)))).

cnf(arcsin_lower_bound_pos_5, axiom,
    (X < 0 | 1 < X |
     ~ lgen(R, Y,
            63/2816 * X ^ 11 + 35/1152 * X ^ 9 + 5/112 * X ^ 7 +
            3/40 * X ^ 5 + 1/6 * X ^ 3 + X) | lgen(R, Y, arcsin(X)))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_id, axiom,
    (X < 0 | 1 < X | ~ lgen(R, Y, X) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_2, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            8 * X * (X + 1) * (X ^ 2 + 6 * X + 1) /
            (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1)) |
     lgen(R, Y, sqrt(X)))).

fof(asin_5, conjecture,
    (! [X] :
       ((0 < X & X < 1) =>
        6 * (sqrt(1 + X) - sqrt(1 - X)) / (4 + sqrt(1 + X) + sqrt(1 - X)) <
        23/100 + arcsin(X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 1) =>
        6 * (sqrt(1 + X) - sqrt(1 - X)) / (4 + sqrt(1 + X) + sqrt(1 - X)) <
        23/100 + arcsin(X))), inference(strip, [], [asin_5])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 1) =>
          6 * (sqrt(1 + X) - sqrt(1 - X)) /
          (4 + sqrt(1 + X) + sqrt(1 - X)) < 23/100 + arcsin(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 1 &
        23/100 + arcsin(X) <=
        6 * (sqrt(1 + X) - sqrt(1 - X)) /
        (4 + sqrt(1 + X) + sqrt(1 - X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 1 & 0 < skoXC1 &
     23/100 + arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (23/100 + arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408))))))))))) <
     (-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (1691/2400 +
           skoXC1 *
           (59/40 +
            skoXC1 *
            (393/280 +
             skoXC1 *
             (2053/1680 +
              skoXC1 *
              (2113/2016 +
               skoXC1 *
               (6137/6720 +
                skoXC1 *
                (35107/44352 +
                 skoXC1 *
                 (20807/29568 +
                  skoXC1 *
                  (1099/16896 +
                   skoXC1 *
                   (-27601/101376 +
                    skoXC1 *
                    (-945/11264 + skoXC1 * -63/22528)))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -4616/25 +
     skoXC1 *
     (-6124/25 +
      skoXC1 *
      (46/25 +
       skoXC1 *
       (6193/75 +
        skoXC1 *
        (1411/80 +
         skoXC1 *
         (1691/2400 +
          skoXC1 *
          (59/40 +
           skoXC1 *
           (393/280 +
            skoXC1 *
            (2053/1680 +
             skoXC1 *
             (2113/2016 +
              skoXC1 *
              (6137/6720 +
               skoXC1 *
               (35107/44352 +
                skoXC1 *
                (20807/29568 +
                 skoXC1 *
                 (1099/16896 +
                  skoXC1 *
                  (-27601/101376 +
                   skoXC1 *
                   (-945/11264 + skoXC1 * -63/22528))))))))))))))) <=
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) |
     0 <=
     -577/50 +
     skoXC1 *
     (777/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (1/3 +
        skoXC1 *
        (-1/6 +
         skoXC1 *
         (3/20 +
          skoXC1 *
          (-3/40 +
           skoXC1 *
           (5/56 +
            skoXC1 *
            (-5/112 +
             skoXC1 *
             (35/576 +
              skoXC1 *
              (-35/1152 +
               skoXC1 * (63/1408 + skoXC1 * -63/2816))))))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    (sqrt(X_000048) < X_000049 | ~ lgen(1, sqrt(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (X_000048 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000048 ^ 4 + 28 * X_000048 ^ 3 + 70 * X_000048 ^ 2 +
             28 * X_000048 + 1) /
            ((X_000048 + 1) * (X_000048 ^ 2 + 6 * X_000048 + 1)),
            X_000049) | lgen(1, sqrt(X_000048), X_000049)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_4, plain,
    (sqrt(X_000048) < X_000049 | X_000048 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000048 ^ 4 + 28 * X_000048 ^ 3 + 70 * X_000048 ^ 2 +
             28 * X_000048 + 1) /
            ((X_000048 + 1) * (X_000048 ^ 2 + 6 * X_000048 + 1)),
            X_000049)), inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (sqrt(X_000048) < X_000049 | X_000048 < 0 |
     X_000049 <=
     (1/8 +
      X_000048 *
      (7/2 + X_000048 * (35/4 + X_000048 * (7/2 + X_000048 * 1/8)))) /
     (1 + X_000048 * (7 + X_000048 * (7 + X_000048)))),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) |
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) <
     (623/50 +
      skoXC1 *
      (-523/50 +
       skoXC1 *
       (-2 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/3 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/20 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/56 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/576 +
                skoXC1 * (63/1408 + skoXC1 * -63/1408)))))))))))) /
     (2 + skoXC1 * -1) | 2 + skoXC1 * -1 <= 0 |
     623/50 +
     skoXC1 *
     (-523/50 +
      skoXC1 *
      (-2 +
       skoXC1 *
       (1/3 +
        skoXC1 *
        (-1/3 +
         skoXC1 *
         (3/20 +
          skoXC1 *
          (-3/20 +
           skoXC1 *
           (5/56 +
            skoXC1 *
            (-5/56 +
             skoXC1 *
             (35/576 +
              skoXC1 *
              (-35/576 +
               skoXC1 * (63/1408 + skoXC1 * -63/1408))))))))))) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) *
     (2 + skoXC1 * -1)), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    ((-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) *
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     623/100 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_9, plain,
    (X_000042 < sqrt(X_000041) | ~ lgen(1, X_000042, sqrt(X_000041))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_10, plain,
    (X_000041 < 0 | ~ lgen(1, X_000042, 2 * X_000041 / (X_000041 + 1)) |
     lgen(1, X_000042, sqrt(X_000041))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_11, plain,
    (X_000041 < 0 | X_000042 < sqrt(X_000041) |
     ~ lgen(1, X_000042, 2 * X_000041 / (X_000041 + 1))),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (X_000041 < 0 | X_000042 < sqrt(X_000041) |
     X_000041 * 2 / (1 + X_000041) <= X_000042),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (1 + skoXC1 * -1 < 0 |
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <
     sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_15, plain,
    (23/100 + arcsin(skoXC1) <=
     6 * (sqrt(1 + skoXC1) - sqrt(1 - skoXC1)) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 - skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_16, plain,
    (23/100 + arcsin(skoXC1) <=
     (sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <
     23/100 + arcsin(skoXC1) |
     (23/100 + arcsin(skoXC1)) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_18, plain,
    ((23/100 + arcsin(skoXC1)) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     sqrt(1 + skoXC1) * 6 + sqrt(1 + skoXC1 * -1) * -6 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (arcsin(skoXC1) * (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100 |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (-10 + skoXC1 * 6 < (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) |
     2 + skoXC1 <= 0 |
     4 + skoXC1 * (2 + skoXC1 * -2) <= (-10 + skoXC1 * 6) * (2 + skoXC1)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_21, plain,
    ((-10 + skoXC1 * 6) / (2 + skoXC1 * -1) <
     (2 + skoXC1 * 2) / (2 + skoXC1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) * (2 + skoXC1 * -1) <=
     -10 + skoXC1 * 6 | 2 + skoXC1 * -1 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_22, plain,
    (1 + skoXC1 < 0 |
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1) < sqrt(1 + skoXC1) |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_23, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < (2 + skoXC1 * -2) / (2 + skoXC1 * -1) |
     2 + skoXC1 * -2 <= (-4 + sqrt(1 + skoXC1) * -1) * (2 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_24, plain,
    (1 + skoXC1 * -1 < 0 |
     -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     -4 + sqrt(1 + skoXC1) * -1), inference(subst, [], [refute_0_12])).

cnf(refute_0_25, plain,
    (1 < skoXC1 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (skoXC1 < 1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_28, plain, (skoXC1 <= 1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(resolve, [], [refute_0_28, refute_0_25])).

cnf(refute_0_30, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     2 + skoXC1 * -2 <= (-4 + sqrt(1 + skoXC1) * -1) * (2 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_29, refute_0_23])).

cnf(refute_0_31, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     sqrt(1 + skoXC1) * (2 + skoXC1 * -1) <= -10 + skoXC1 * 6),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain, (skoXC1 < 2),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_33, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     sqrt(1 + skoXC1) * (2 + skoXC1 * -1) <= -10 + skoXC1 * 6),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     2 + skoXC1 * -1 <= 0 |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(split, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     sqrt(1 + skoXC1) <= (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_35, refute_0_32])).

cnf(refute_0_37, plain,
    (1 + skoXC1 < 0 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1) * 2 / (1 + (1 + skoXC1)) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_36, refute_0_22])).

cnf(refute_0_38, plain,
    (skoXC1 < -1 | -4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain, (-1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_40, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) <=
     (-10 + skoXC1 * 6) / (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_39, refute_0_38])).

cnf(refute_0_41, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (2 + skoXC1 * 2) / (2 + skoXC1) * (2 + skoXC1 * -1) <=
     -10 + skoXC1 * 6 | 2 + skoXC1 * -1 <= 0),
    inference(resolve, [], [refute_0_40, refute_0_21])).

cnf(refute_0_42, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 <= skoXC1 |
     (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) <= -10 + skoXC1 * 6),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     (4 + skoXC1 * (2 + skoXC1 * -2)) / (2 + skoXC1) <= -10 + skoXC1 * 6),
    inference(resolve, [], [refute_0_42, refute_0_32])).

cnf(refute_0_44, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | 2 + skoXC1 <= 0 |
     4 + skoXC1 * (2 + skoXC1 * -2) <= (-10 + skoXC1 * 6) * (2 + skoXC1)),
    inference(resolve, [], [refute_0_43, refute_0_20])).

cnf(refute_0_45, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) |
     skoXC1 * (skoXC1 * -8) <= -24 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain, (-24 < skoXC1 * (skoXC1 * -8) | skoXC1 <= -2),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_47, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain, (-2 < skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_49, plain,
    (-4 + sqrt(1 + skoXC1) * -1 < sqrt(1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (arcsin(skoXC1) * (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100),
    inference(resolve, [], [refute_0_19, refute_0_49])).

cnf(refute_0_51, plain,
    (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0 |
     arcsin(skoXC1) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(split, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (arcsin(skoXC1) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (arcsin(skoXC1) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_52, refute_0_49])).

cnf(refute_0_54, plain,
    (X_000119 < arcsin(X_000118) | ~ lgen(1, X_000119, arcsin(X_000118))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_55, plain,
    (1 < X_000118 | X_000118 < 0 |
     ~ lgen(1, X_000119,
            63/2816 * X_000118 ^ 11 + 35/1152 * X_000118 ^ 9 +
            5/112 * X_000118 ^ 7 + 3/40 * X_000118 ^ 5 +
            1/6 * X_000118 ^ 3 + X_000118) |
     lgen(1, X_000119, arcsin(X_000118))),
    inference(subst, [], [arcsin_lower_bound_pos_5])).

cnf(refute_0_56, plain,
    (1 < X_000118 | X_000118 < 0 | X_000119 < arcsin(X_000118) |
     ~ lgen(1, X_000119,
            63/2816 * X_000118 ^ 11 + 35/1152 * X_000118 ^ 9 +
            5/112 * X_000118 ^ 7 + 3/40 * X_000118 ^ 5 +
            1/6 * X_000118 ^ 3 + X_000118)),
    inference(resolve, [], [refute_0_55, refute_0_54])).

cnf(refute_0_57, plain,
    (1 < X_000118 | X_000118 < 0 | X_000119 < arcsin(X_000118) |
     X_000118 *
     (1 +
      X_000118 *
      (X_000118 *
       (1/6 +
        X_000118 *
        (X_000118 *
         (3/40 +
          X_000118 *
          (X_000118 *
           (5/112 +
            X_000118 *
            (X_000118 *
             (35/1152 + X_000118 * (X_000118 * 63/2816)))))))))) <=
     X_000119), inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 < 0 |
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) < arcsin(skoXC1) |
     1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(subst, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (skoXC1 < 0 | 1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_53, refute_0_58])).

cnf(refute_0_60, plain, (1 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_27, refute_0_26])).

cnf(refute_0_61, plain,
    (1 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_28, refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_62, refute_0_14])).

cnf(refute_0_64, plain,
    (sqrt(1 + skoXC1 * -1) *
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (sqrt(1 + skoXC1 * -1) *
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(resolve, [], [refute_0_64, refute_0_49])).

cnf(refute_0_66, plain,
    (623/100 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0 |
     sqrt(1 + skoXC1 * -1) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(split, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     -623/100 |
     sqrt(1 + skoXC1 * -1) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (1 + skoXC1 * 1/2 <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (1 + skoXC1 * 1/2) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_69, plain,
    (sqrt(X_000034) < X_000035 | ~ lgen(1, sqrt(X_000034), X_000035)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_70, plain,
    (X_000034 < 0 | ~ lgen(1, (X_000034 + 1) / 2, X_000035) |
     lgen(1, sqrt(X_000034), X_000035)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_71, plain,
    (sqrt(X_000034) < X_000035 | X_000034 < 0 |
     ~ lgen(1, (X_000034 + 1) / 2, X_000035)),
    inference(resolve, [], [refute_0_70, refute_0_69])).

cnf(refute_0_72, plain,
    (sqrt(X_000034) < X_000035 | X_000034 < 0 |
     X_000035 <= 1/2 + X_000034 * 1/2),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     1/2 + (1 + skoXC1) * 1/2), inference(subst, [], [refute_0_72])).

cnf(refute_0_74, plain,
    ((-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) < 1 + skoXC1 * -1 |
     (1 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_75, plain,
    (X_000077 < sqrt(X_000076) | ~ lgen(1, X_000077, sqrt(X_000076))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_76, plain,
    (1 < X_000076 | X_000076 < 0 | ~ lgen(1, X_000077, X_000076) |
     lgen(1, X_000077, sqrt(X_000076))),
    inference(subst, [], [sqrt_lower_bound_id])).

cnf(refute_0_77, plain,
    (1 < X_000076 | X_000076 < 0 | X_000077 < sqrt(X_000076) |
     ~ lgen(1, X_000077, X_000076)),
    inference(resolve, [], [refute_0_76, refute_0_75])).

cnf(refute_0_78, plain,
    (1 < X_000076 | X_000076 < 0 | X_000077 < sqrt(X_000076) |
     X_000076 <= X_000077), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (1 + skoXC1 * -1 < 0 | 1 < 1 + skoXC1 * -1 |
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <
     sqrt(1 + skoXC1 * -1) |
     1 + skoXC1 * -1 <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(subst, [], [refute_0_78])).

cnf(refute_0_80, plain,
    ((-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_81, plain,
    (X_000084 < arcsin(X_000083) | ~ lgen(1, X_000084, arcsin(X_000083))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_82, plain,
    (1 < X_000083 | X_000083 < 0 |
     ~ lgen(1, X_000084, 1/6 * X_000083 ^ 3 + X_000083) |
     lgen(1, X_000084, arcsin(X_000083))),
    inference(subst, [], [arcsin_lower_bound_pos_1])).

cnf(refute_0_83, plain,
    (1 < X_000083 | X_000083 < 0 | X_000084 < arcsin(X_000083) |
     ~ lgen(1, X_000084, 1/6 * X_000083 ^ 3 + X_000083)),
    inference(resolve, [], [refute_0_82, refute_0_81])).

cnf(refute_0_84, plain,
    (1 < X_000083 | X_000083 < 0 | X_000084 < arcsin(X_000083) |
     X_000083 * (1 + X_000083 * (X_000083 * 1/6)) <= X_000084),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (skoXC1 < 0 |
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) < arcsin(skoXC1) |
     1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(subst, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (skoXC1 < 0 | 1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_53, refute_0_85])).

cnf(refute_0_87, plain,
    (1 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_60, refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <=
     (-23/25 + sqrt(1 + skoXC1) * 577/100 +
      sqrt(1 + skoXC1 * -1) * -623/100) /
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_28, refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) *
     (4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1)) <=
     -23/25 + sqrt(1 + skoXC1) * 577/100 +
     sqrt(1 + skoXC1 * -1) * -623/100 |
     4 + sqrt(1 + skoXC1) + sqrt(1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_80])).

cnf(refute_0_90, plain,
    (sqrt(1 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     sqrt(1 + skoXC1 * -1) <= -4 + sqrt(1 + skoXC1) * -1),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (sqrt(1 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))),
    inference(resolve, [], [refute_0_90, refute_0_49])).

cnf(refute_0_92, plain,
    (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0 |
     sqrt(1 + skoXC1 * -1) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(split, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= -623/100 |
     sqrt(1 + skoXC1 * -1) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (-623/100 < skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_95, plain,
    (sqrt(1 + skoXC1 * -1) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (1 + skoXC1 * -1 < 0 | 1 < 1 + skoXC1 * -1 |
     1 + skoXC1 * -1 <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_95, refute_0_79])).

cnf(refute_0_97, plain,
    (skoXC1 < 0 | 1 < skoXC1 |
     1 + skoXC1 * -1 <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (1 < skoXC1 |
     1 + skoXC1 * -1 <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_60, refute_0_97])).

cnf(refute_0_99, plain,
    (1 + skoXC1 * -1 <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_28, refute_0_98])).

cnf(refute_0_100, plain,
    ((1 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(resolve, [], [refute_0_99, refute_0_74])).

cnf(refute_0_101, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= -623/100 |
     sqrt(1 + skoXC1) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (sqrt(1 + skoXC1) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_101, refute_0_94])).

cnf(refute_0_103, plain,
    (0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) <= -577/100 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (-577/100 < skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_106, plain,
    ((-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain,
    (1 + skoXC1 < 0 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     1/2 + (1 + skoXC1) * 1/2),
    inference(resolve, [], [refute_0_106, refute_0_73])).

cnf(refute_0_108, plain,
    (skoXC1 < -1 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     1 + skoXC1 * 1/2), inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    ((-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     1 + skoXC1 * 1/2),
    inference(resolve, [], [refute_0_39, refute_0_108])).

cnf(refute_0_110, plain,
    (0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (1 + skoXC1 * 1/2) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_109, refute_0_68])).

cnf(refute_0_111, plain,
    (skoXC1 *
     (-623/200 + skoXC1 * (-1/2 + skoXC1 * (1 + skoXC1 * -1/12))) <=
     -69/50 | skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) <= -577/100),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-577/100 < skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) |
     skoXC1 *
     (-623/200 + skoXC1 * (-1/2 + skoXC1 * (1 + skoXC1 * -1/12))) <=
     -69/50), inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_113, plain,
    (skoXC1 *
     (-623/200 + skoXC1 * (-1/2 + skoXC1 * (1 + skoXC1 * -1/12))) <=
     -69/50), inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (-143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6))) <
     (-114588161/3200 +
      skoXC1 *
      (-2117911/100 +
       skoXC1 *
       (406881/320 +
        skoXC1 *
        (7464119/4800 +
         skoXC1 *
         (7727669/9600 +
          skoXC1 * (4819/48 + skoXC1 * (113/48 + skoXC1 * 1/192))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     -114588161/3200 +
     skoXC1 *
     (-2117911/100 +
      skoXC1 *
      (406881/320 +
       skoXC1 *
       (7464119/4800 +
        skoXC1 *
        (7727669/9600 +
         skoXC1 * (4819/48 + skoXC1 * (113/48 + skoXC1 * 1/192)))))) <=
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) *
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_115, plain,
    ((198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_116, plain,
    (sqrt(X_000055) < X_000056 | ~ lgen(1, sqrt(X_000055), X_000056)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_117, plain,
    (X_000055 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000055 + 17920 * X_000055 ^ 2 +
             448 * X_000055 ^ 3 + X_000055 ^ 4) /
            ((X_000055 + 16) * (256 + 96 * X_000055 + X_000055 ^ 2)),
            X_000056) | lgen(1, sqrt(X_000055), X_000056)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_118, plain,
    (sqrt(X_000055) < X_000056 | X_000055 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000055 + 17920 * X_000055 ^ 2 +
             448 * X_000055 ^ 3 + X_000055 ^ 4) /
            ((X_000055 + 16) * (256 + 96 * X_000055 + X_000055 ^ 2)),
            X_000056)),
    inference(resolve, [], [refute_0_117, refute_0_116])).

cnf(refute_0_119, plain,
    (sqrt(X_000055) < X_000056 | X_000055 < 1 |
     X_000056 <=
     (2048 +
      X_000055 *
      (3584 + X_000055 * (560 + X_000055 * (14 + X_000055 * 1/32)))) /
     (4096 + X_000055 * (1792 + X_000055 * (112 + X_000055)))),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (1 + skoXC1 < 1 |
     sqrt(1 + skoXC1) <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (2048 +
      (1 + skoXC1) *
      (3584 +
       (1 + skoXC1) * (560 + (1 + skoXC1) * (14 + (1 + skoXC1) * 1/32)))) /
     (4096 + (1 + skoXC1) * (1792 + (1 + skoXC1) * (112 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (1 + skoXC1 < 1 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (2048 +
      (1 + skoXC1) *
      (3584 +
       (1 + skoXC1) * (560 + (1 + skoXC1) * (14 + (1 + skoXC1) * 1/32)))) /
     (4096 + (1 + skoXC1) * (1792 + (1 + skoXC1) * (112 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_106, refute_0_120])).

cnf(refute_0_122, plain,
    (skoXC1 < 0 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_124, plain,
    ((-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)))),
    inference(resolve, [], [refute_0_123, refute_0_122])).

cnf(refute_0_125, plain,
    (0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_124, refute_0_115])).

cnf(refute_0_126, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) <= -577/100 |
     (-114588161/3200 +
      skoXC1 *
      (-2117911/100 +
       skoXC1 *
       (406881/320 +
        skoXC1 *
        (7464119/4800 +
         skoXC1 *
         (7727669/9600 +
          skoXC1 * (4819/48 + skoXC1 * (113/48 + skoXC1 * 1/192))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    ((-114588161/3200 +
      skoXC1 *
      (-2117911/100 +
       skoXC1 *
       (406881/320 +
        skoXC1 *
        (7464119/4800 +
         skoXC1 *
         (7727669/9600 +
          skoXC1 * (4819/48 + skoXC1 * (113/48 + skoXC1 * 1/192))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_126, refute_0_105])).

cnf(refute_0_128, plain,
    (-114588161/3200 +
     skoXC1 *
     (-2117911/100 +
      skoXC1 *
      (406881/320 +
       skoXC1 *
       (7464119/4800 +
        skoXC1 *
        (7727669/9600 +
         skoXC1 * (4819/48 + skoXC1 * (113/48 + skoXC1 * 1/192)))))) <=
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) *
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_127, refute_0_114])).

cnf(refute_0_129, plain,
    (skoXC1 *
     (-1412449/100 +
      skoXC1 *
      (-10224987/1600 +
       skoXC1 *
       (7044093/1600 +
        skoXC1 *
        (13162261/9600 +
         skoXC1 * (-6781/48 + skoXC1 * (-767/48 + skoXC1 * -31/192)))))) <=
     -22714719/3200 | skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= -6001),
    inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (-6001 < skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) |
     skoXC1 *
     (-1412449/100 +
      skoXC1 *
      (-10224987/1600 +
       skoXC1 *
       (7044093/1600 +
        skoXC1 *
        (13162261/9600 +
         skoXC1 * (-6781/48 + skoXC1 * (-767/48 + skoXC1 * -31/192)))))) <=
     -22714719/3200),
    inference(decision, [], [refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_131, plain,
    (skoXC1 *
     (-1412449/100 +
      skoXC1 *
      (-10224987/1600 +
       skoXC1 *
       (7044093/1600 +
        skoXC1 *
        (13162261/9600 +
         skoXC1 * (-6781/48 + skoXC1 * (-767/48 + skoXC1 * -31/192)))))) <=
     -22714719/3200),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (-143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6))) <
     (-2308/25 +
      skoXC1 *
      (-4216/25 +
       skoXC1 *
       (-417/5 +
        skoXC1 *
        (-31/75 +
         skoXC1 *
         (20669/2400 +
          skoXC1 * (83/24 + skoXC1 * (2/3 + skoXC1 * 1/48))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -2308/25 +
     skoXC1 *
     (-4216/25 +
      skoXC1 *
      (-417/5 +
       skoXC1 *
       (-31/75 +
        skoXC1 *
        (20669/2400 +
         skoXC1 * (83/24 + skoXC1 * (2/3 + skoXC1 * 1/48)))))) <=
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_133, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_134, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_135, plain,
    (1 + skoXC1 < 0 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_106, refute_0_134])).

cnf(refute_0_136, plain,
    (skoXC1 < -1 |
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain,
    ((-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) /
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_39, refute_0_136])).

cnf(refute_0_138, plain,
    (0 <= -577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_137, refute_0_133])).

cnf(refute_0_139, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) <= -577/100 |
     (-2308/25 +
      skoXC1 *
      (-4216/25 +
       skoXC1 *
       (-417/5 +
        skoXC1 *
        (-31/75 +
         skoXC1 *
         (20669/2400 +
          skoXC1 * (83/24 + skoXC1 * (2/3 + skoXC1 * 1/48))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    ((-2308/25 +
      skoXC1 *
      (-4216/25 +
       skoXC1 *
       (-417/5 +
        skoXC1 *
        (-31/75 +
         skoXC1 *
         (20669/2400 +
          skoXC1 * (83/24 + skoXC1 * (2/3 + skoXC1 * 1/48))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/20 +
     skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_139, refute_0_105])).

cnf(refute_0_141, plain,
    (-2308/25 +
     skoXC1 *
     (-4216/25 +
      skoXC1 *
      (-417/5 +
       skoXC1 *
       (-31/75 +
        skoXC1 *
        (20669/2400 +
         skoXC1 * (83/24 + skoXC1 * (2/3 + skoXC1 * 1/48)))))) <=
     (-143/20 +
      skoXC1 * (123/100 + skoXC1 * (1 + skoXC1 * (-5/6 + skoXC1 * 1/6)))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_140, refute_0_132])).

cnf(refute_0_142, plain,
    (skoXC1 *
     (-418/25 +
      skoXC1 *
      (-2871/50 +
       skoXC1 *
       (-1623/100 +
        skoXC1 *
        (35317/2400 +
         skoXC1 * (163/24 + skoXC1 * (-1/6 + skoXC1 * -7/48)))))) <=
     -552/25 | skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1)) |
     skoXC1 *
     (-418/25 +
      skoXC1 *
      (-2871/50 +
       skoXC1 *
       (-1623/100 +
        skoXC1 *
        (35317/2400 +
         skoXC1 * (163/24 + skoXC1 * (-1/6 + skoXC1 * -7/48)))))) <=
     -552/25),
    inference(decision, [],
              [refute_0_131, refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_144, plain,
    (skoXC1 *
     (-418/25 +
      skoXC1 *
      (-2871/50 +
       skoXC1 *
       (-1623/100 +
        skoXC1 *
        (35317/2400 +
         skoXC1 * (163/24 + skoXC1 * (-1/6 + skoXC1 * -7/48)))))) <=
     -552/25), inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (1 + skoXC1 * 1/2 <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (1 + skoXC1 * 1/2) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_146, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     1/2 + (1 + skoXC1) * 1/2), inference(subst, [], [refute_0_72])).

cnf(refute_0_147, plain,
    (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) <
     (623/50 +
      skoXC1 *
      (-523/50 + skoXC1 * (-2 + skoXC1 * (1/3 + skoXC1 * -1/3)))) /
     (2 + skoXC1 * -1) | 2 + skoXC1 * -1 <= 0 |
     623/50 +
     skoXC1 * (-523/50 + skoXC1 * (-2 + skoXC1 * (1/3 + skoXC1 * -1/3))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) *
     (2 + skoXC1 * -1)), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_148, plain,
    ((-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_149, plain,
    (1 + skoXC1 * -1 < 0 |
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <
     sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_150, plain,
    (1 + skoXC1 * -1 < 0 |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_95, refute_0_149])).

cnf(refute_0_151, plain,
    (1 < skoXC1 |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    ((2 + skoXC1 * -2) / (2 + skoXC1 * -1) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_28, refute_0_151])).

cnf(refute_0_153, plain,
    ((2 + skoXC1 * -2) / (2 + skoXC1 * -1) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(resolve, [], [refute_0_152, refute_0_148])).

cnf(refute_0_154, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= -623/100 |
     (623/50 +
      skoXC1 *
      (-523/50 + skoXC1 * (-2 + skoXC1 * (1/3 + skoXC1 * -1/3)))) /
     (2 + skoXC1 * -1) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))),
    inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain,
    ((623/50 +
      skoXC1 *
      (-523/50 + skoXC1 * (-2 + skoXC1 * (1/3 + skoXC1 * -1/3)))) /
     (2 + skoXC1 * -1) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))),
    inference(resolve, [], [refute_0_154, refute_0_94])).

cnf(refute_0_156, plain,
    (2 + skoXC1 * -1 <= 0 |
     623/50 +
     skoXC1 * (-523/50 + skoXC1 * (-2 + skoXC1 * (1/3 + skoXC1 * -1/3))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) *
     (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_155, refute_0_147])).

cnf(refute_0_157, plain,
    (2 <= skoXC1 |
     sqrt(1 + skoXC1) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(arithmetic, [], [refute_0_156])).

cnf(refute_0_158, plain,
    (sqrt(1 + skoXC1) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_157, refute_0_32])).

cnf(refute_0_159, plain,
    (0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6))) <=
     -577/50 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (-577/50 <
     skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6)))),
    inference(decision, [],
              [refute_0_144, refute_0_131, refute_0_113, refute_0_26,
               refute_0_27])).

cnf(refute_0_162, plain,
    ((-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (1 + skoXC1 < 0 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     1/2 + (1 + skoXC1) * 1/2),
    inference(resolve, [], [refute_0_162, refute_0_146])).

cnf(refute_0_164, plain,
    (skoXC1 < -1 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     1 + skoXC1 * 1/2), inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    ((-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     1 + skoXC1 * 1/2),
    inference(resolve, [], [refute_0_39, refute_0_164])).

cnf(refute_0_166, plain,
    (0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (1 + skoXC1 * 1/2) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_165, refute_0_145])).

cnf(refute_0_167, plain,
    (skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6))) <=
     -577/50 |
     skoXC1 *
     (-69/50 +
      skoXC1 *
      (-623/200 + skoXC1 * (3/2 + skoXC1 * (-1 + skoXC1 * -1/12)))) <=
     -69/25), inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain,
    (-577/50 <
     skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6))) |
     skoXC1 *
     (-69/50 +
      skoXC1 *
      (-623/200 + skoXC1 * (3/2 + skoXC1 * (-1 + skoXC1 * -1/12)))) <=
     -69/25),
    inference(decision, [],
              [refute_0_144, refute_0_131, refute_0_113, refute_0_26,
               refute_0_27])).

cnf(refute_0_169, plain,
    (skoXC1 *
     (-69/50 +
      skoXC1 *
      (-623/200 + skoXC1 * (3/2 + skoXC1 * (-1 + skoXC1 * -1/12)))) <=
     -69/25), inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (-143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1))) <
     (-114588161/1600 +
      skoXC1 *
      (-20958143/3200 +
       skoXC1 *
       (18977693/800 +
        skoXC1 *
        (8825023/4800 +
         skoXC1 *
         (1757/32 +
          skoXC1 *
          (-5800069/9600 +
           skoXC1 *
           (-1531/16 + skoXC1 * (-75/32 + skoXC1 * -1/192)))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     -114588161/1600 +
     skoXC1 *
     (-20958143/3200 +
      skoXC1 *
      (18977693/800 +
       skoXC1 *
       (8825023/4800 +
        skoXC1 *
        (1757/32 +
         skoXC1 *
         (-5800069/9600 +
          skoXC1 * (-1531/16 + skoXC1 * (-75/32 + skoXC1 * -1/192))))))) <=
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) *
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_171, plain,
    ((198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_172, plain,
    (1 + skoXC1 < 1 |
     sqrt(1 + skoXC1) <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (2048 +
      (1 + skoXC1) *
      (3584 +
       (1 + skoXC1) * (560 + (1 + skoXC1) * (14 + (1 + skoXC1) * 1/32)))) /
     (4096 + (1 + skoXC1) * (1792 + (1 + skoXC1) * (112 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_119])).

cnf(refute_0_173, plain,
    (1 + skoXC1 < 1 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (2048 +
      (1 + skoXC1) *
      (3584 +
       (1 + skoXC1) * (560 + (1 + skoXC1) * (14 + (1 + skoXC1) * 1/32)))) /
     (4096 + (1 + skoXC1) * (1792 + (1 + skoXC1) * (112 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_162, refute_0_172])).

cnf(refute_0_174, plain,
    (skoXC1 < 0 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)))),
    inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    ((-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)))),
    inference(resolve, [], [refute_0_123, refute_0_174])).

cnf(refute_0_176, plain,
    (0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (198593/32 +
      skoXC1 *
      (37969/8 + skoXC1 * (9635/16 + skoXC1 * (113/8 + skoXC1 * 1/32)))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_175, refute_0_171])).

cnf(refute_0_177, plain,
    (skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6))) <=
     -577/50 |
     (-114588161/1600 +
      skoXC1 *
      (-20958143/3200 +
       skoXC1 *
       (18977693/800 +
        skoXC1 *
        (8825023/4800 +
         skoXC1 *
         (1757/32 +
          skoXC1 *
          (-5800069/9600 +
           skoXC1 *
           (-1531/16 + skoXC1 * (-75/32 + skoXC1 * -1/192)))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    ((-114588161/1600 +
      skoXC1 *
      (-20958143/3200 +
       skoXC1 *
       (18977693/800 +
        skoXC1 *
        (8825023/4800 +
         skoXC1 *
         (1757/32 +
          skoXC1 *
          (-5800069/9600 +
           skoXC1 *
           (-1531/16 + skoXC1 * (-75/32 + skoXC1 * -1/192)))))))) /
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_177, refute_0_161])).

cnf(refute_0_179, plain,
    (-114588161/1600 +
     skoXC1 *
     (-20958143/3200 +
      skoXC1 *
      (18977693/800 +
       skoXC1 *
       (8825023/4800 +
        skoXC1 *
        (1757/32 +
         skoXC1 *
         (-5800069/9600 +
          skoXC1 * (-1531/16 + skoXC1 * (-75/32 + skoXC1 * -1/192))))))) <=
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) *
     (6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1))) |
     6001 + skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_178, refute_0_170])).

cnf(refute_0_180, plain,
    (skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) <= -6001 |
     skoXC1 *
     (6524481/3200 +
      skoXC1 *
      (-13970883/800 +
       skoXC1 *
       (-1037099/1600 +
        skoXC1 *
        (-2619579/800 +
         skoXC1 *
         (-7800023/3200 +
          skoXC1 *
          (-10033/48 + skoXC1 * (-107/32 + skoXC1 * -1/192))))))) <=
     -22714719/1600), inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (-6001 < skoXC1 * (2019 + skoXC1 * (115 + skoXC1)) |
     skoXC1 *
     (6524481/3200 +
      skoXC1 *
      (-13970883/800 +
       skoXC1 *
       (-1037099/1600 +
        skoXC1 *
        (-2619579/800 +
         skoXC1 *
         (-7800023/3200 +
          skoXC1 *
          (-10033/48 + skoXC1 * (-107/32 + skoXC1 * -1/192))))))) <=
     -22714719/1600),
    inference(decision, [],
              [refute_0_169, refute_0_144, refute_0_131, refute_0_113,
               refute_0_26, refute_0_27])).

cnf(refute_0_182, plain,
    (skoXC1 *
     (6524481/3200 +
      skoXC1 *
      (-13970883/800 +
       skoXC1 *
       (-1037099/1600 +
        skoXC1 *
        (-2619579/800 +
         skoXC1 *
         (-7800023/3200 +
          skoXC1 *
          (-10033/48 + skoXC1 * (-107/32 + skoXC1 * -1/192))))))) <=
     -22714719/1600),
    inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (-143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1))) <
     (-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (-4069/2400 +
           skoXC1 * (-17/8 + skoXC1 * (-5/8 + skoXC1 * -1/48)))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -4616/25 +
     skoXC1 *
     (-6124/25 +
      skoXC1 *
      (46/25 +
       skoXC1 *
       (6193/75 +
        skoXC1 *
        (1411/80 +
         skoXC1 *
         (-4069/2400 +
          skoXC1 * (-17/8 + skoXC1 * (-5/8 + skoXC1 * -1/48))))))) <=
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_184, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_185, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_186, plain,
    (1 + skoXC1 < 0 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_162, refute_0_185])).

cnf(refute_0_187, plain,
    (skoXC1 < -1 |
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_186])).

cnf(refute_0_188, plain,
    ((-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) /
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_39, refute_0_187])).

cnf(refute_0_189, plain,
    (0 <=
     -577/50 +
     skoXC1 * (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 + skoXC1 * (-1 + skoXC1 * (1/3 + skoXC1 * -1/6)))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_188, refute_0_184])).

cnf(refute_0_190, plain,
    (skoXC1 * (-777/100 + skoXC1 * (1 + skoXC1 * (-1/3 + skoXC1 * 1/6))) <=
     -577/50 |
     (-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (-4069/2400 +
           skoXC1 * (-17/8 + skoXC1 * (-5/8 + skoXC1 * -1/48)))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(arithmetic, [], [refute_0_189])).

cnf(refute_0_191, plain,
    ((-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (-4069/2400 +
           skoXC1 * (-17/8 + skoXC1 * (-5/8 + skoXC1 * -1/48)))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/10 +
     skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))),
    inference(resolve, [], [refute_0_190, refute_0_161])).

cnf(refute_0_192, plain,
    (-4616/25 +
     skoXC1 *
     (-6124/25 +
      skoXC1 *
      (46/25 +
       skoXC1 *
       (6193/75 +
        skoXC1 *
        (1411/80 +
         skoXC1 *
         (-4069/2400 +
          skoXC1 * (-17/8 + skoXC1 * (-5/8 + skoXC1 * -1/48))))))) <=
     (-143/10 +
      skoXC1 * (169/50 + skoXC1 * (6 + skoXC1 * (-5/3 + skoXC1)))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_191, refute_0_183])).

cnf(refute_0_193, plain,
    (skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16 |
     skoXC1 *
     (1104/25 +
      skoXC1 *
      (-807/25 +
       skoXC1 *
       (-2713/50 +
        skoXC1 *
        (-8697/400 +
         skoXC1 *
         (-12023/800 +
          skoXC1 * (-251/24 + skoXC1 * (-13/8 + skoXC1 * -1/48))))))) <=
     -1104/25), inference(arithmetic, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1)) |
     skoXC1 *
     (1104/25 +
      skoXC1 *
      (-807/25 +
       skoXC1 *
       (-2713/50 +
        skoXC1 *
        (-8697/400 +
         skoXC1 *
         (-12023/800 +
          skoXC1 * (-251/24 + skoXC1 * (-13/8 + skoXC1 * -1/48))))))) <=
     -1104/25),
    inference(decision, [],
              [refute_0_182, refute_0_169, refute_0_144, refute_0_131,
               refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_195, plain,
    (skoXC1 *
     (1104/25 +
      skoXC1 *
      (-807/25 +
       skoXC1 *
       (-2713/50 +
        skoXC1 *
        (-8697/400 +
         skoXC1 *
         (-12023/800 +
          skoXC1 * (-251/24 + skoXC1 * (-13/8 + skoXC1 * -1/48))))))) <=
     -1104/25), inference(resolve, [], [refute_0_193, refute_0_194])).

cnf(refute_0_196, plain,
    (1 + skoXC1 * 1/2 <
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) |
     0 <=
     -18464/25 +
     skoXC1 *
     (40128/25 +
      skoXC1 *
      (-5896/5 +
       skoXC1 *
       (27448/75 +
        skoXC1 *
        (-24131/300 +
         skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6)))))) |
     (1 + skoXC1 * 1/2) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_197, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     1/2 + (1 + skoXC1) * 1/2), inference(subst, [], [refute_0_72])).

cnf(refute_0_198, plain,
    (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) <
     (19936/25 +
      skoXC1 *
      (-9328/5 +
       skoXC1 *
       (34364/25 +
        skoXC1 *
        (-19118/75 +
         skoXC1 *
         (-6862/75 +
          skoXC1 * (160/3 + skoXC1 * (-44/3 + skoXC1 * 4/3))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     19936/25 +
     skoXC1 *
     (-9328/5 +
      skoXC1 *
      (34364/25 +
       skoXC1 *
       (-19118/75 +
        skoXC1 *
        (-6862/75 +
         skoXC1 * (160/3 + skoXC1 * (-44/3 + skoXC1 * 4/3)))))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_199, plain,
    ((-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_200, plain,
    (X_000063 < sqrt(X_000062) | ~ lgen(1, X_000063, sqrt(X_000062))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_201, plain,
    (X_000062 < 0 |
     ~ lgen(1, X_000063,
            8 * X_000062 * (X_000062 + 1) *
            (X_000062 ^ 2 + 6 * X_000062 + 1) /
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1)) | lgen(1, X_000063, sqrt(X_000062))),
    inference(subst, [], [sqrt_lower_bound_2])).

cnf(refute_0_202, plain,
    (X_000062 < 0 | X_000063 < sqrt(X_000062) |
     ~ lgen(1, X_000063,
            8 * X_000062 * (X_000062 + 1) *
            (X_000062 ^ 2 + 6 * X_000062 + 1) /
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1))),
    inference(resolve, [], [refute_0_201, refute_0_200])).

cnf(refute_0_203, plain,
    (X_000062 < 0 | X_000063 < sqrt(X_000062) |
     X_000062 * (8 + X_000062 * (56 + X_000062 * (56 + X_000062 * 8))) /
     (1 +
      X_000062 * (28 + X_000062 * (70 + X_000062 * (28 + X_000062)))) <=
     X_000063), inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (1 + skoXC1 * -1 < 0 |
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <
     sqrt(1 + skoXC1 * -1) |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(subst, [], [refute_0_203])).

cnf(refute_0_205, plain,
    (1 + skoXC1 * -1 < 0 |
     (1 + skoXC1 * -1) *
     (8 +
      (1 + skoXC1 * -1) *
      (56 + (1 + skoXC1 * -1) * (56 + (1 + skoXC1 * -1) * 8))) /
     (1 +
      (1 + skoXC1 * -1) *
      (28 +
       (1 + skoXC1 * -1) *
       (70 + (1 + skoXC1 * -1) * (28 + (1 + skoXC1 * -1))))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_95, refute_0_204])).

cnf(refute_0_206, plain,
    (1 < skoXC1 |
     (128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) /
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)))),
    inference(resolve, [], [refute_0_28, refute_0_206])).

cnf(refute_0_208, plain,
    ((128 +
      skoXC1 * (-320 + skoXC1 * (272 + skoXC1 * (-88 + skoXC1 * 8)))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) *
     (623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6))) |
     623/100 + skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(resolve, [], [refute_0_207, refute_0_199])).

cnf(refute_0_209, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/6)) <= -623/100 |
     (19936/25 +
      skoXC1 *
      (-9328/5 +
       skoXC1 *
       (34364/25 +
        skoXC1 *
        (-19118/75 +
         skoXC1 *
         (-6862/75 +
          skoXC1 * (160/3 + skoXC1 * (-44/3 + skoXC1 * 4/3))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))),
    inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    ((19936/25 +
      skoXC1 *
      (-9328/5 +
       skoXC1 *
       (34364/25 +
        skoXC1 *
        (-19118/75 +
         skoXC1 *
         (-6862/75 +
          skoXC1 * (160/3 + skoXC1 * (-44/3 + skoXC1 * 4/3))))))) /
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))) <=
     -23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
     sqrt(1 + skoXC1) *
     (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))),
    inference(resolve, [], [refute_0_209, refute_0_94])).

cnf(refute_0_211, plain,
    (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <=
     0 |
     19936/25 +
     skoXC1 *
     (-9328/5 +
      skoXC1 *
      (34364/25 +
       skoXC1 *
       (-19118/75 +
        skoXC1 *
        (-6862/75 +
         skoXC1 * (160/3 + skoXC1 * (-44/3 + skoXC1 * 4/3)))))) <=
     (-23/25 + skoXC1 * (-4 + skoXC1 * (skoXC1 * -2/3)) +
      sqrt(1 + skoXC1) *
      (577/100 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)))) *
     (128 + skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))))),
    inference(resolve, [], [refute_0_210, refute_0_198])).

cnf(refute_0_212, plain,
    (skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1))) <= -128 |
     sqrt(1 + skoXC1) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(arithmetic, [], [refute_0_211])).

cnf(refute_0_213, plain,
    (-128 < skoXC1 * (-256 + skoXC1 * (160 + skoXC1 * (-32 + skoXC1)))),
    inference(decision, [],
              [refute_0_195, refute_0_182, refute_0_169, refute_0_144,
               refute_0_131, refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_214, plain,
    (sqrt(1 + skoXC1) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(resolve, [], [refute_0_212, refute_0_213])).

cnf(refute_0_215, plain,
    (0 <=
     -18464/25 +
     skoXC1 *
     (40128/25 +
      skoXC1 *
      (-5896/5 +
       skoXC1 *
       (27448/75 +
        skoXC1 *
        (-24131/300 +
         skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6)))))) |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (skoXC1 *
     (-40128/25 +
      skoXC1 *
      (5896/5 +
       skoXC1 *
       (-27448/75 +
        skoXC1 *
        (24131/300 +
         skoXC1 * (-83/3 + skoXC1 * (16/3 + skoXC1 * -1/6)))))) <=
     -18464/25 |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_215])).

cnf(refute_0_217, plain,
    (-18464/25 <
     skoXC1 *
     (-40128/25 +
      skoXC1 *
      (5896/5 +
       skoXC1 *
       (-27448/75 +
        skoXC1 *
        (24131/300 +
         skoXC1 * (-83/3 + skoXC1 * (16/3 + skoXC1 * -1/6))))))),
    inference(decision, [],
              [refute_0_195, refute_0_182, refute_0_169, refute_0_144,
               refute_0_131, refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_218, plain,
    ((-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_216, refute_0_217])).

cnf(refute_0_219, plain,
    (1 + skoXC1 < 0 |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     1/2 + (1 + skoXC1) * 1/2),
    inference(resolve, [], [refute_0_218, refute_0_197])).

cnf(refute_0_220, plain,
    (skoXC1 < -1 |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     1 + skoXC1 * 1/2), inference(arithmetic, [], [refute_0_219])).

cnf(refute_0_221, plain,
    ((-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     1 + skoXC1 * 1/2),
    inference(resolve, [], [refute_0_39, refute_0_220])).

cnf(refute_0_222, plain,
    (0 <=
     -18464/25 +
     skoXC1 *
     (40128/25 +
      skoXC1 *
      (-5896/5 +
       skoXC1 *
       (27448/75 +
        skoXC1 *
        (-24131/300 +
         skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6)))))) |
     (1 + skoXC1 * 1/2) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(resolve, [], [refute_0_221, refute_0_196])).

cnf(refute_0_223, plain,
    (skoXC1 *
     (-40128/25 +
      skoXC1 *
      (5896/5 +
       skoXC1 *
       (-27448/75 +
        skoXC1 *
        (24131/300 +
         skoXC1 * (-83/3 + skoXC1 * (16/3 + skoXC1 * -1/6)))))) <=
     -18464/25 |
     skoXC1 *
     (-8832/25 +
      skoXC1 *
      (3028/25 +
       skoXC1 *
       (5434/25 +
        skoXC1 *
        (-28669/100 +
         skoXC1 *
         (90869/600 +
          skoXC1 * (-55/2 + skoXC1 * (-1/2 + skoXC1 * 1/12))))))) <=
     -4416/25), inference(arithmetic, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (-18464/25 <
     skoXC1 *
     (-40128/25 +
      skoXC1 *
      (5896/5 +
       skoXC1 *
       (-27448/75 +
        skoXC1 *
        (24131/300 +
         skoXC1 * (-83/3 + skoXC1 * (16/3 + skoXC1 * -1/6)))))) |
     skoXC1 *
     (-8832/25 +
      skoXC1 *
      (3028/25 +
       skoXC1 *
       (5434/25 +
        skoXC1 *
        (-28669/100 +
         skoXC1 *
         (90869/600 +
          skoXC1 * (-55/2 + skoXC1 * (-1/2 + skoXC1 * 1/12))))))) <=
     -4416/25),
    inference(decision, [],
              [refute_0_195, refute_0_182, refute_0_169, refute_0_144,
               refute_0_131, refute_0_113, refute_0_26, refute_0_27])).

cnf(refute_0_225, plain,
    (skoXC1 *
     (-8832/25 +
      skoXC1 *
      (3028/25 +
       skoXC1 *
       (5434/25 +
        skoXC1 *
        (-28669/100 +
         skoXC1 *
         (90869/600 +
          skoXC1 * (-55/2 + skoXC1 * (-1/2 + skoXC1 * 1/12))))))) <=
     -4416/25), inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (-4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2)))))) <
     (-295424/25 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (443136/25 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-170792/25 +
          skoXC1 *
          (672 +
           skoXC1 *
           (12694/25 +
            skoXC1 *
            (328/3 +
             skoXC1 *
             (-577/800 +
              skoXC1 * (-349/24 + skoXC1 * (skoXC1 * 1/48))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     -295424/25 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (443136/25 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-170792/25 +
         skoXC1 *
         (672 +
          skoXC1 *
          (12694/25 +
           skoXC1 *
           (328/3 +
            skoXC1 *
            (-577/800 +
             skoXC1 * (-349/24 + skoXC1 * (skoXC1 * 1/48)))))))))) <=
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_227, plain,
    ((16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) |
     0 <=
     -18464/25 +
     skoXC1 *
     (40128/25 +
      skoXC1 *
      (-5896/5 +
       skoXC1 *
       (27448/75 +
        skoXC1 *
        (-24131/300 +
         skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6)))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_228, plain,
    (1 + skoXC1 < 0 |
     sqrt(1 + skoXC1) <
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_229, plain,
    (1 + skoXC1 < 0 |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_218, refute_0_228])).

cnf(refute_0_230, plain,
    (skoXC1 < -1 |
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    ((-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) /
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_39, refute_0_230])).

cnf(refute_0_232, plain,
    (0 <=
     -18464/25 +
     skoXC1 *
     (40128/25 +
      skoXC1 *
      (-5896/5 +
       skoXC1 *
       (27448/75 +
        skoXC1 *
        (-24131/300 +
         skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6)))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-18464/25 +
      skoXC1 *
      (40128/25 +
       skoXC1 *
       (-5896/5 +
        skoXC1 *
        (27448/75 +
         skoXC1 *
         (-24131/300 +
          skoXC1 * (83/3 + skoXC1 * (-16/3 + skoXC1 * 1/6))))))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(resolve, [], [refute_0_231, refute_0_227])).

cnf(refute_0_233, plain,
    (skoXC1 *
     (-40128/25 +
      skoXC1 *
      (5896/5 +
       skoXC1 *
       (-27448/75 +
        skoXC1 *
        (24131/300 +
         skoXC1 * (-83/3 + skoXC1 * (16/3 + skoXC1 * -1/6)))))) <=
     -18464/25 |
     (-295424/25 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (443136/25 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-170792/25 +
          skoXC1 *
          (672 +
           skoXC1 *
           (12694/25 +
            skoXC1 *
            (328/3 +
             skoXC1 *
             (-577/800 +
              skoXC1 * (-349/24 + skoXC1 * (skoXC1 * 1/48))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(arithmetic, [], [refute_0_232])).

cnf(refute_0_234, plain,
    ((-295424/25 +
      skoXC1 *
      (2048 +
       skoXC1 *
       (443136/25 +
        skoXC1 *
        (-8192/3 +
         skoXC1 *
         (-170792/25 +
          skoXC1 *
          (672 +
           skoXC1 *
           (12694/25 +
            skoXC1 *
            (328/3 +
             skoXC1 *
             (-577/800 +
              skoXC1 * (-349/24 + skoXC1 * (skoXC1 * 1/48))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -4576/5 +
     skoXC1 *
     (39728/25 +
      skoXC1 *
      (-12444/25 +
       skoXC1 *
       (-33074/75 +
        skoXC1 *
        (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))),
    inference(resolve, [], [refute_0_233, refute_0_217])).

cnf(refute_0_235, plain,
    (-295424/25 +
     skoXC1 *
     (2048 +
      skoXC1 *
      (443136/25 +
       skoXC1 *
       (-8192/3 +
        skoXC1 *
        (-170792/25 +
         skoXC1 *
         (672 +
          skoXC1 *
          (12694/25 +
           skoXC1 *
           (328/3 +
            skoXC1 *
            (-577/800 +
             skoXC1 * (-349/24 + skoXC1 * (skoXC1 * 1/48)))))))))) <=
     (-4576/5 +
      skoXC1 *
      (39728/25 +
       skoXC1 *
       (-12444/25 +
        skoXC1 *
        (-33074/75 +
         skoXC1 *
         (9731/25 + skoXC1 * (-164 + skoXC1 * (36 + skoXC1 * -2))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_234, refute_0_226])).

cnf(refute_0_236, plain,
    (skoXC1 *
     (-35328/25 +
      skoXC1 *
      (-82432/25 +
       skoXC1 *
       (32384/25 +
        skoXC1 *
        (22816/25 +
         skoXC1 *
         (-17072/15 +
          skoXC1 *
          (31226/75 +
           skoXC1 *
           (39607/75 +
            skoXC1 *
            (-118977/800 +
             skoXC1 * (-733/24 + skoXC1 * (2 + skoXC1 * 1/48)))))))))) <=
     -70656/25 | skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(arithmetic, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1)) |
     skoXC1 *
     (-35328/25 +
      skoXC1 *
      (-82432/25 +
       skoXC1 *
       (32384/25 +
        skoXC1 *
        (22816/25 +
         skoXC1 *
         (-17072/15 +
          skoXC1 *
          (31226/75 +
           skoXC1 *
           (39607/75 +
            skoXC1 *
            (-118977/800 +
             skoXC1 * (-733/24 + skoXC1 * (2 + skoXC1 * 1/48)))))))))) <=
     -70656/25),
    inference(decision, [],
              [refute_0_225, refute_0_195, refute_0_182, refute_0_169,
               refute_0_144, refute_0_131, refute_0_113, refute_0_26,
               refute_0_27])).

cnf(refute_0_238, plain,
    (skoXC1 *
     (-35328/25 +
      skoXC1 *
      (-82432/25 +
       skoXC1 *
       (32384/25 +
        skoXC1 *
        (22816/25 +
         skoXC1 *
         (-17072/15 +
          skoXC1 *
          (31226/75 +
           skoXC1 *
           (39607/75 +
            skoXC1 *
            (-118977/800 +
             skoXC1 * (-733/24 + skoXC1 * (2 + skoXC1 * 1/48)))))))))) <=
     -70656/25), inference(resolve, [], [refute_0_236, refute_0_237])).

cnf(refute_0_239, plain,
    (-623/100 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))),
    inference(decision, [],
              [refute_0_27, refute_0_26, refute_0_113, refute_0_131,
               refute_0_144, refute_0_169, refute_0_182, refute_0_195,
               refute_0_225, refute_0_238])).

cnf(refute_0_240, plain,
    (sqrt(1 + skoXC1 * -1) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_67, refute_0_239])).

cnf(refute_0_241, plain,
    (1 + skoXC1 * -1 < 0 |
     (1 + skoXC1 * -1) * 2 / (1 + (1 + skoXC1 * -1)) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_240, refute_0_13])).

cnf(refute_0_242, plain,
    (1 < skoXC1 |
     (2 + skoXC1 * -2) / (2 + skoXC1 * -1) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    ((2 + skoXC1 * -2) / (2 + skoXC1 * -1) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) /
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))))),
    inference(resolve, [], [refute_0_28, refute_0_242])).

cnf(refute_0_244, plain,
    ((2 + skoXC1 * -2) / (2 + skoXC1 * -1) *
     (623/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (1/6 +
         skoXC1 *
         (skoXC1 *
          (3/40 +
           skoXC1 *
           (skoXC1 *
            (5/112 +
             skoXC1 *
             (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816))))))))))) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816))))))))))) |
     623/100 +
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     0), inference(resolve, [], [refute_0_243, refute_0_8])).

cnf(refute_0_245, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (3/40 +
          skoXC1 *
          (skoXC1 *
           (5/112 +
            skoXC1 *
            (skoXC1 * (35/1152 + skoXC1 * (skoXC1 * 63/2816)))))))))) <=
     -623/100 |
     (623/50 +
      skoXC1 *
      (-523/50 +
       skoXC1 *
       (-2 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/3 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/20 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/56 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/576 +
                skoXC1 * (63/1408 + skoXC1 * -63/1408)))))))))))) /
     (2 + skoXC1 * -1) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(arithmetic, [], [refute_0_244])).

cnf(refute_0_246, plain,
    ((623/50 +
      skoXC1 *
      (-523/50 +
       skoXC1 *
       (-2 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/3 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/20 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/56 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/576 +
                skoXC1 * (63/1408 + skoXC1 * -63/1408)))))))))))) /
     (2 + skoXC1 * -1) <=
     -23/25 +
     skoXC1 *
     (-4 +
      skoXC1 *
      (skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (-3/10 +
          skoXC1 *
          (skoXC1 *
           (-5/28 +
            skoXC1 *
            (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
     sqrt(1 + skoXC1) *
     (577/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (-1/6 +
         skoXC1 *
         (skoXC1 *
          (-3/40 +
           skoXC1 *
           (skoXC1 *
            (-5/112 +
             skoXC1 *
             (skoXC1 * (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))),
    inference(resolve, [], [refute_0_245, refute_0_239])).

cnf(refute_0_247, plain,
    (2 + skoXC1 * -1 <= 0 |
     623/50 +
     skoXC1 *
     (-523/50 +
      skoXC1 *
      (-2 +
       skoXC1 *
       (1/3 +
        skoXC1 *
        (-1/3 +
         skoXC1 *
         (3/20 +
          skoXC1 *
          (-3/20 +
           skoXC1 *
           (5/56 +
            skoXC1 *
            (-5/56 +
             skoXC1 *
             (35/576 +
              skoXC1 *
              (-35/576 +
               skoXC1 * (63/1408 + skoXC1 * -63/1408))))))))))) <=
     (-23/25 +
      skoXC1 *
      (-4 +
       skoXC1 *
       (skoXC1 *
        (-2/3 +
         skoXC1 *
         (skoXC1 *
          (-3/10 +
           skoXC1 *
           (skoXC1 *
            (-5/28 +
             skoXC1 *
             (skoXC1 * (-35/288 + skoXC1 * (skoXC1 * -63/704)))))))))) +
      sqrt(1 + skoXC1) *
      (577/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (skoXC1 *
         (-1/6 +
          skoXC1 *
          (skoXC1 *
           (-3/40 +
            skoXC1 *
            (skoXC1 *
             (-5/112 +
              skoXC1 *
              (skoXC1 *
               (-35/1152 + skoXC1 * (skoXC1 * -63/2816)))))))))))) *
     (2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_246, refute_0_7])).

cnf(refute_0_248, plain,
    (2 <= skoXC1 |
     sqrt(1 + skoXC1) *
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (sqrt(1 + skoXC1) *
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(resolve, [], [refute_0_248, refute_0_32])).

cnf(refute_0_250, plain,
    (0 <=
     -577/50 +
     skoXC1 *
     (777/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (1/3 +
        skoXC1 *
        (-1/6 +
         skoXC1 *
         (3/20 +
          skoXC1 *
          (-3/40 +
           skoXC1 *
           (5/56 +
            skoXC1 *
            (-5/112 +
             skoXC1 *
             (35/576 +
              skoXC1 *
              (-35/1152 +
               skoXC1 * (63/1408 + skoXC1 * -63/2816))))))))))) |
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     sqrt(1 + skoXC1)), inference(split, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (skoXC1 *
     (-777/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-1/3 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (-3/20 +
          skoXC1 *
          (3/40 +
           skoXC1 *
           (-5/56 +
            skoXC1 *
            (5/112 +
             skoXC1 *
             (-35/576 +
              skoXC1 *
              (35/1152 +
               skoXC1 * (-63/1408 + skoXC1 * 63/2816))))))))))) <=
     -577/50 |
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     sqrt(1 + skoXC1)), inference(arithmetic, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (-577/50 <
     skoXC1 *
     (-777/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-1/3 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (-3/20 +
          skoXC1 *
          (3/40 +
           skoXC1 *
           (-5/56 +
            skoXC1 *
            (5/112 +
             skoXC1 *
             (-35/576 +
              skoXC1 *
              (35/1152 + skoXC1 * (-63/1408 + skoXC1 * 63/2816)))))))))))),
    inference(decision, [],
              [refute_0_27, refute_0_26, refute_0_113, refute_0_131,
               refute_0_144, refute_0_169, refute_0_182, refute_0_195,
               refute_0_225, refute_0_238])).

cnf(refute_0_253, plain,
    ((-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_251, refute_0_252])).

cnf(refute_0_254, plain,
    (1 + skoXC1 < 0 |
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     (1/8 +
      (1 + skoXC1) *
      (7/2 +
       (1 + skoXC1) *
       (35/4 + (1 + skoXC1) * (7/2 + (1 + skoXC1) * 1/8)))) /
     (1 + (1 + skoXC1) * (7 + (1 + skoXC1) * (7 + (1 + skoXC1))))),
    inference(resolve, [], [refute_0_253, refute_0_6])).

cnf(refute_0_255, plain,
    (skoXC1 < -1 |
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(arithmetic, [], [refute_0_254])).

cnf(refute_0_256, plain,
    ((-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) /
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)))),
    inference(resolve, [], [refute_0_39, refute_0_255])).

cnf(refute_0_257, plain,
    (0 <=
     -577/50 +
     skoXC1 *
     (777/100 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (1/3 +
        skoXC1 *
        (-1/6 +
         skoXC1 *
         (3/20 +
          skoXC1 *
          (-3/40 +
           skoXC1 *
           (5/56 +
            skoXC1 *
            (-5/112 +
             skoXC1 *
             (35/576 +
              skoXC1 *
              (-35/1152 +
               skoXC1 * (63/1408 + skoXC1 * -63/2816))))))))))) |
     (16 + skoXC1 * (32 + skoXC1 * (20 + skoXC1 * (4 + skoXC1 * 1/8)))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) *
     (-577/50 +
      skoXC1 *
      (777/100 +
       skoXC1 *
       (-1 +
        skoXC1 *
        (1/3 +
         skoXC1 *
         (-1/6 +
          skoXC1 *
          (3/20 +
           skoXC1 *
           (-3/40 +
            skoXC1 *
            (5/56 +
             skoXC1 *
             (-5/112 +
              skoXC1 *
              (35/576 +
               skoXC1 *
               (-35/1152 +
                skoXC1 * (63/1408 + skoXC1 * -63/2816)))))))))))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(resolve, [], [refute_0_256, refute_0_1])).

cnf(refute_0_258, plain,
    (skoXC1 *
     (-777/100 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-1/3 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (-3/20 +
          skoXC1 *
          (3/40 +
           skoXC1 *
           (-5/56 +
            skoXC1 *
            (5/112 +
             skoXC1 *
             (-35/576 +
              skoXC1 *
              (35/1152 +
               skoXC1 * (-63/1408 + skoXC1 * 63/2816))))))))))) <=
     -577/50 |
     (-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (1691/2400 +
           skoXC1 *
           (59/40 +
            skoXC1 *
            (393/280 +
             skoXC1 *
             (2053/1680 +
              skoXC1 *
              (2113/2016 +
               skoXC1 *
               (6137/6720 +
                skoXC1 *
                (35107/44352 +
                 skoXC1 *
                 (20807/29568 +
                  skoXC1 *
                  (1099/16896 +
                   skoXC1 *
                   (-27601/101376 +
                    skoXC1 *
                    (-945/11264 + skoXC1 * -63/22528)))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain,
    ((-4616/25 +
      skoXC1 *
      (-6124/25 +
       skoXC1 *
       (46/25 +
        skoXC1 *
        (6193/75 +
         skoXC1 *
         (1411/80 +
          skoXC1 *
          (1691/2400 +
           skoXC1 *
           (59/40 +
            skoXC1 *
            (393/280 +
             skoXC1 *
             (2053/1680 +
              skoXC1 *
              (2113/2016 +
               skoXC1 *
               (6137/6720 +
                skoXC1 *
                (35107/44352 +
                 skoXC1 *
                 (20807/29568 +
                  skoXC1 *
                  (1099/16896 +
                   skoXC1 *
                   (-27601/101376 +
                    skoXC1 *
                    (-945/11264 + skoXC1 * -63/22528)))))))))))))))) /
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) <=
     -143/10 +
     skoXC1 *
     (169/50 +
      skoXC1 *
      (6 +
       skoXC1 *
       (-5/3 +
        skoXC1 *
        (1 +
         skoXC1 *
         (-3/4 +
          skoXC1 *
          (9/20 +
           skoXC1 *
           (-25/56 +
            skoXC1 *
            (15/56 +
             skoXC1 *
             (-175/576 +
              skoXC1 *
              (35/192 +
               skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))),
    inference(resolve, [], [refute_0_258, refute_0_252])).

cnf(refute_0_260, plain,
    (-4616/25 +
     skoXC1 *
     (-6124/25 +
      skoXC1 *
      (46/25 +
       skoXC1 *
       (6193/75 +
        skoXC1 *
        (1411/80 +
         skoXC1 *
         (1691/2400 +
          skoXC1 *
          (59/40 +
           skoXC1 *
           (393/280 +
            skoXC1 *
            (2053/1680 +
             skoXC1 *
             (2113/2016 +
              skoXC1 *
              (6137/6720 +
               skoXC1 *
               (35107/44352 +
                skoXC1 *
                (20807/29568 +
                 skoXC1 *
                 (1099/16896 +
                  skoXC1 *
                  (-27601/101376 +
                   skoXC1 *
                   (-945/11264 + skoXC1 * -63/22528))))))))))))))) <=
     (-143/10 +
      skoXC1 *
      (169/50 +
       skoXC1 *
       (6 +
        skoXC1 *
        (-5/3 +
         skoXC1 *
         (1 +
          skoXC1 *
          (-3/4 +
           skoXC1 *
           (9/20 +
            skoXC1 *
            (-25/56 +
             skoXC1 *
             (15/56 +
              skoXC1 *
              (-175/576 +
               skoXC1 *
               (35/192 +
                skoXC1 * (-315/1408 + skoXC1 * 189/1408)))))))))))) *
     (16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1))) |
     16 + skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_259, refute_0_0])).

cnf(refute_0_261, plain,
    (skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16 |
     skoXC1 *
     (1104/25 +
      skoXC1 *
      (-807/25 +
       skoXC1 *
       (-2713/50 +
        skoXC1 *
        (-8697/400 +
         skoXC1 *
         (-503/800 +
          skoXC1 *
          (473/120 +
           skoXC1 *
           (1189/280 +
            skoXC1 *
            (6553/1680 +
             skoXC1 *
             (11743/3360 +
              skoXC1 *
              (20537/6720 +
               skoXC1 *
               (13633/4928 +
                skoXC1 *
                (213439/88704 +
                 skoXC1 *
                 (-18613/16896 +
                  skoXC1 *
                  (-141001/101376 +
                   skoXC1 *
                   (-2457/11264 + skoXC1 * -63/22528))))))))))))))) <=
     -1104/25), inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (-1104/25 <
     skoXC1 *
     (1104/25 +
      skoXC1 *
      (-807/25 +
       skoXC1 *
       (-2713/50 +
        skoXC1 *
        (-8697/400 +
         skoXC1 *
         (-503/800 +
          skoXC1 *
          (473/120 +
           skoXC1 *
           (1189/280 +
            skoXC1 *
            (6553/1680 +
             skoXC1 *
             (11743/3360 +
              skoXC1 *
              (20537/6720 +
               skoXC1 *
               (13633/4928 +
                skoXC1 *
                (213439/88704 +
                 skoXC1 *
                 (-18613/16896 +
                  skoXC1 *
                  (-141001/101376 +
                   skoXC1 *
                   (-2457/11264 + skoXC1 * -63/22528))))))))))))))) |
     skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(decision, [],
              [refute_0_27, refute_0_26, refute_0_113, refute_0_131,
               refute_0_144, refute_0_169, refute_0_182, refute_0_195,
               refute_0_225, refute_0_238])).

cnf(refute_0_263, plain, (skoXC1 * (24 + skoXC1 * (10 + skoXC1)) <= -16),
    inference(resolve, [], [refute_0_261, refute_0_262])).

cnf(refute_0_264, plain, (-16 < skoXC1 * (24 + skoXC1 * (10 + skoXC1))),
    inference(decision, [],
              [refute_0_27, refute_0_26, refute_0_113, refute_0_131,
               refute_0_144, refute_0_169, refute_0_182, refute_0_195,
               refute_0_225, refute_0_238])).

cnf(refute_0_265, plain, ($false),
    inference(resolve, [], [refute_0_263, refute_0_264])).
% SZS output end CNFRefutation for asin-5-weak.tptp
