% SZS output start CNFRefutation for atan-error-analysis-2-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(atan_upper_bound_case_42, axiom,
    (0 <= X |
     ~ lgen(R,
            1/35 *
            (15159 * X ^ 6 + 147455 * X ^ 4 + 345345 * X ^ 2 + 225225) *
            X /
            (35 * X ^ 8 + 1260 * X ^ 6 + 6930 * X ^ 4 + 12012 * X ^ 2 +
             6435), Y) | lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_43, axiom,
    (X < 0 |
     ~ lgen(R,
            1/35 * (256 * X ^ 6 + 5943 * X ^ 4 + 19250 * X ^ 2 + 15015) *
            X / (35 * X ^ 6 + 315 * X ^ 4 + 693 * X ^ 2 + 429), Y) |
     lgen(R, arctan(X), Y))).

cnf(atan_lower_bound_case_42, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            1/35 * (256 * X ^ 6 + 5943 * X ^ 4 + 19250 * X ^ 2 + 15015) *
            X / (35 * X ^ 6 + 315 * X ^ 4 + 693 * X ^ 2 + 429)) |
     lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_43, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            1/35 *
            (15159 * X ^ 6 + 147455 * X ^ 4 + 345345 * X ^ 2 + 225225) *
            X /
            (35 * X ^ 8 + 1260 * X ^ 6 + 6930 * X ^ 4 + 12012 * X ^ 2 +
             6435)) | lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, 3 * X / (X ^ 2 + 3), Y) | lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_32, axiom,
    (0 <= X |
     ~ lgen(R,
            7/5 * (33 * X ^ 4 + 170 * X ^ 2 + 165) * X /
            (5 * X ^ 6 + 105 * X ^ 4 + 315 * X ^ 2 + 231), Y) |
     lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_33, axiom,
    (X < 0 |
     ~ lgen(R,
            1/15 * (64 * X ^ 4 + 735 * X ^ 2 + 945) * X /
            (15 * X ^ 4 + 70 * X ^ 2 + 63), Y) | lgen(R, arctan(X), Y))).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_32, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            1/15 * (64 * X ^ 4 + 735 * X ^ 2 + 945) * X /
            (15 * X ^ 4 + 70 * X ^ 2 + 63)) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_33, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            7/5 * (33 * X ^ 4 + 170 * X ^ 2 + 165) * X /
            (5 * X ^ 6 + 105 * X ^ 4 + 315 * X ^ 2 + 231)) |
     lgen(R, Y, arctan(X)))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(atan_error_analysis_2, conjecture,
    (! [X] :
       (abs(X) <= 1 =>
        abs(arctan(X) -
            (499933/500000 * X - 660599/2000000 * X ^ 3 +
             180141/1000000 * X ^ 5 - 85133/1000000 * X ^ 7 +
             208351/10000000 * X ^ 9)) <= 3/200000))).

fof(subgoal_0, plain,
    (! [X] :
       (abs(X) <= 1 =>
        abs(arctan(X) -
            (499933/500000 * X - 660599/2000000 * X ^ 3 +
             180141/1000000 * X ^ 5 - 85133/1000000 * X ^ 7 +
             208351/10000000 * X ^ 9)) <= 3/200000)),
    inference(strip, [], [atan_error_analysis_2])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (abs(X) <= 1 =>
          abs(arctan(X) -
              (499933/500000 * X - 660599/2000000 * X ^ 3 +
               180141/1000000 * X ^ 5 - 85133/1000000 * X ^ 7 +
               208351/10000000 * X ^ 9)) <= 3/200000)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (3/200000 <
        abs(arctan(X) -
            (499933/500000 * X - 660599/2000000 * X ^ 3 +
             180141/1000000 * X ^ 5 - 85133/1000000 * X ^ 7 +
             208351/10000000 * X ^ 9)) & abs(X) <= 1)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (3/200000 <
     abs(arctan(skoXC1) -
         (499933/500000 * skoXC1 - 660599/2000000 * skoXC1 ^ 3 +
          180141/1000000 * skoXC1 ^ 5 - 85133/1000000 * skoXC1 ^ 7 +
          208351/10000000 * skoXC1 ^ 9)) & abs(skoXC1) <= 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (3/200000 <
     abs(arctan(skoXC1) -
         (499933/500000 * skoXC1 - 660599/2000000 * skoXC1 ^ 3 +
          180141/1000000 * skoXC1 ^ 5 - 85133/1000000 * skoXC1 ^ 7 +
          208351/10000000 * skoXC1 ^ 9))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (abs(skoXC1) <= 1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000183, arctan(X_000182)) | X_000183 <= arctan(X_000182)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (X_000182 < 0 |
     ~ lgen(0, X_000183,
            1/35 *
            (15159 * X_000182 ^ 6 + 147455 * X_000182 ^ 4 +
             345345 * X_000182 ^ 2 + 225225) * X_000182 /
            (35 * X_000182 ^ 8 + 1260 * X_000182 ^ 6 +
             6930 * X_000182 ^ 4 + 12012 * X_000182 ^ 2 + 6435)) |
     lgen(0, X_000183, arctan(X_000182))),
    inference(subst, [], [atan_lower_bound_case_43])).

cnf(refute_0_3, plain,
    (X_000182 < 0 |
     ~ lgen(0, X_000183,
            1/35 *
            (15159 * X_000182 ^ 6 + 147455 * X_000182 ^ 4 +
             345345 * X_000182 ^ 2 + 225225) * X_000182 /
            (35 * X_000182 ^ 8 + 1260 * X_000182 ^ 6 +
             6930 * X_000182 ^ 4 + 12012 * X_000182 ^ 2 + 6435)) |
     X_000183 <= arctan(X_000182)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000182 *
     (6435 +
      X_000182 *
      (X_000182 *
       (9867 +
        X_000182 *
        (X_000182 * (4213 + X_000182 * (X_000182 * 15159/35)))))) /
     (6435 +
      X_000182 *
      (X_000182 *
       (12012 +
        X_000182 *
        (X_000182 *
         (6930 +
          X_000182 *
          (X_000182 * (1260 + X_000182 * (X_000182 * 35)))))))) <
     X_000183 | X_000182 < 0 | X_000183 <= arctan(X_000182)),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (3/200000 <
     abs(arctan(skoXC1) -
         (499933/500000 * skoXC1 - 660599/2000000 * skoXC1 ^ 3 +
          180141/1000000 * skoXC1 ^ 5 - 85133/1000000 * skoXC1 ^ 7 +
          208351/10000000 * skoXC1 ^ 9))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (3/200000 <
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1))), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (0 <=
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) =
     -(skoXC1 *
      (-499933/500000 +
       skoXC1 *
       (skoXC1 *
        (660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (-180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
      arctan(skoXC1))), inference(subst, [], [abs_negative])).

cnf(refute_0_9, plain,
    (3/200000 <
     -(skoXC1 *
      (-499933/500000 +
       skoXC1 *
       (skoXC1 *
        (660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (-180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
      arctan(skoXC1)) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) !=
     -(skoXC1 *
      (-499933/500000 +
       skoXC1 *
       (skoXC1 *
        (660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (-180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
      arctan(skoXC1)) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) <= 3/200000), introduced(tautology, [equality])).

cnf(refute_0_10, plain,
    (3/200000 <
     -(skoXC1 *
      (-499933/500000 +
       skoXC1 *
       (skoXC1 *
        (660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (-180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
      arctan(skoXC1)) |
     0 <=
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) <= 3/200000),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (3/200000 <
     -(skoXC1 *
      (-499933/500000 +
       skoXC1 *
       (skoXC1 *
        (660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (-180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
      arctan(skoXC1)) |
     0 <=
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1)), inference(resolve, [], [refute_0_10, refute_0_7])).

cnf(refute_0_12, plain,
    (arctan(skoXC1) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_14, plain,
    (~ lgen(0, arctan(X_000175), X_000176) | arctan(X_000175) <= X_000176),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_15, plain,
    (~ lgen(0,
            1/35 *
            (15159 * X_000175 ^ 6 + 147455 * X_000175 ^ 4 +
             345345 * X_000175 ^ 2 + 225225) * X_000175 /
            (35 * X_000175 ^ 8 + 1260 * X_000175 ^ 6 +
             6930 * X_000175 ^ 4 + 12012 * X_000175 ^ 2 + 6435),
            X_000176) | 0 <= X_000175 |
     lgen(0, arctan(X_000175), X_000176)),
    inference(subst, [], [atan_upper_bound_case_42])).

cnf(refute_0_16, plain,
    (~ lgen(0,
            1/35 *
            (15159 * X_000175 ^ 6 + 147455 * X_000175 ^ 4 +
             345345 * X_000175 ^ 2 + 225225) * X_000175 /
            (35 * X_000175 ^ 8 + 1260 * X_000175 ^ 6 +
             6930 * X_000175 ^ 4 + 12012 * X_000175 ^ 2 + 6435),
            X_000176) | 0 <= X_000175 | arctan(X_000175) <= X_000176),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    (X_000176 <
     X_000175 *
     (6435 +
      X_000175 *
      (X_000175 *
       (9867 +
        X_000175 *
        (X_000175 * (4213 + X_000175 * (X_000175 * 15159/35)))))) /
     (6435 +
      X_000175 *
      (X_000175 *
       (12012 +
        X_000175 *
        (X_000175 *
         (6930 +
          X_000175 *
          (X_000175 * (1260 + X_000175 * (X_000175 * 35)))))))) |
     0 <= X_000175 | arctan(X_000175) <= X_000176),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     0 <= skoXC1 |
     arctan(skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) < 0 |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) =
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1)), inference(subst, [], [abs_nonnegative])).

cnf(refute_0_20, plain,
    (3/200000 <
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) !=
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) <= 3/200000), introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) < 0 |
     3/200000 <
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) |
     abs(skoXC1 *
         (-499933/500000 +
          skoXC1 *
          (skoXC1 *
           (660599/2000000 +
            skoXC1 *
            (skoXC1 *
             (-180141/1000000 +
              skoXC1 *
              (skoXC1 *
               (85133/1000000 +
                skoXC1 * (skoXC1 * -208351/10000000)))))))) +
         arctan(skoXC1)) <= 3/200000),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1) < 0 |
     3/200000 <
     skoXC1 *
     (-499933/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000)))))))) +
     arctan(skoXC1)), inference(resolve, [], [refute_0_21, refute_0_7])).

cnf(refute_0_23, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_5, refute_0_12])).

cnf(refute_0_25, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_26, plain,
    (~ lgen(0, X_000155, arctan(X_000154)) | X_000155 <= arctan(X_000154)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000155,
            1/35 *
            (256 * X_000154 ^ 6 + 5943 * X_000154 ^ 4 +
             19250 * X_000154 ^ 2 + 15015) * X_000154 /
            (35 * X_000154 ^ 6 + 315 * X_000154 ^ 4 + 693 * X_000154 ^ 2 +
             429)) | 0 <= X_000154 | lgen(0, X_000155, arctan(X_000154))),
    inference(subst, [], [atan_lower_bound_case_42])).

cnf(refute_0_28, plain,
    (~ lgen(0, X_000155,
            1/35 *
            (256 * X_000154 ^ 6 + 5943 * X_000154 ^ 4 +
             19250 * X_000154 ^ 2 + 15015) * X_000154 /
            (35 * X_000154 ^ 6 + 315 * X_000154 ^ 4 + 693 * X_000154 ^ 2 +
             429)) | X_000155 <= arctan(X_000154) | 0 <= X_000154),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (X_000154 *
     (429 +
      X_000154 *
      (X_000154 *
       (550 +
        X_000154 *
        (X_000154 * (849/5 + X_000154 * (X_000154 * 256/35)))))) /
     (429 +
      X_000154 *
      (X_000154 *
       (693 +
        X_000154 * (X_000154 * (315 + X_000154 * (X_000154 * 35)))))) <
     X_000155 | X_000155 <= arctan(X_000154) | 0 <= X_000154),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_30, refute_0_12])).

cnf(refute_0_32, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(resolve, [], [refute_0_25, refute_0_31])).

cnf(refute_0_33, plain,
    (1287/200000 <
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_35, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_29])).

cnf(refute_0_36, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (0 <
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) /
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) |
     231 +
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_39, plain,
    (~ lgen(0, X_000127, arctan(X_000126)) | X_000127 <= arctan(X_000126)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_40, plain,
    (X_000126 < 0 |
     ~ lgen(0, X_000127,
            7/5 * (33 * X_000126 ^ 4 + 170 * X_000126 ^ 2 + 165) *
            X_000126 /
            (5 * X_000126 ^ 6 + 105 * X_000126 ^ 4 + 315 * X_000126 ^ 2 +
             231)) | lgen(0, X_000127, arctan(X_000126))),
    inference(subst, [], [atan_lower_bound_case_33])).

cnf(refute_0_41, plain,
    (X_000126 < 0 |
     ~ lgen(0, X_000127,
            7/5 * (33 * X_000126 ^ 4 + 170 * X_000126 ^ 2 + 165) *
            X_000126 /
            (5 * X_000126 ^ 6 + 105 * X_000126 ^ 4 + 315 * X_000126 ^ 2 +
             231)) | X_000127 <= arctan(X_000126)),
    inference(resolve, [], [refute_0_40, refute_0_39])).

cnf(refute_0_42, plain,
    (X_000126 *
     (231 +
      X_000126 * (X_000126 * (238 + X_000126 * (X_000126 * 231/5)))) /
     (231 +
      X_000126 *
      (X_000126 *
       (315 +
        X_000126 * (X_000126 * (105 + X_000126 * (X_000126 * 5)))))) <
     X_000127 | X_000126 < 0 | X_000127 <= arctan(X_000126)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) /
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     231 +
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= 0),
    inference(resolve, [], [refute_0_38, refute_0_43])).

cnf(refute_0_45, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= -231 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_47, plain,
    (~ lgen(0, X_000099, arctan(X_000098)) | X_000099 <= arctan(X_000098)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_48, plain,
    (~ lgen(0, X_000099,
            1/15 * (64 * X_000098 ^ 4 + 735 * X_000098 ^ 2 + 945) *
            X_000098 / (15 * X_000098 ^ 4 + 70 * X_000098 ^ 2 + 63)) |
     0 <= X_000098 | lgen(0, X_000099, arctan(X_000098))),
    inference(subst, [], [atan_lower_bound_case_32])).

cnf(refute_0_49, plain,
    (~ lgen(0, X_000099,
            1/15 * (64 * X_000098 ^ 4 + 735 * X_000098 ^ 2 + 945) *
            X_000098 / (15 * X_000098 ^ 4 + 70 * X_000098 ^ 2 + 63)) |
     X_000099 <= arctan(X_000098) | 0 <= X_000098),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    (X_000098 *
     (63 + X_000098 * (X_000098 * (49 + X_000098 * (X_000098 * 64/15)))) /
     (63 + X_000098 * (X_000098 * (70 + X_000098 * (X_000098 * 15)))) <
     X_000099 | X_000099 <= arctan(X_000098) | 0 <= X_000098),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_51, refute_0_12])).

cnf(refute_0_53, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(resolve, [], [refute_0_46, refute_0_52])).

cnf(refute_0_54, plain,
    (189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= -63 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_56, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_50])).

cnf(refute_0_57, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (0 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= -63 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (skoXC1 * 3 <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (3 + skoXC1 * skoXC1) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) | 3 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_60, plain,
    (~ lgen(0, X_000071, arctan(X_000070)) | X_000071 <= arctan(X_000070)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_61, plain,
    (X_000070 < 0 |
     ~ lgen(0, X_000071, 3 * X_000070 / (X_000070 ^ 2 + 3)) |
     lgen(0, X_000071, arctan(X_000070))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_62, plain,
    (X_000070 < 0 |
     ~ lgen(0, X_000071, 3 * X_000070 / (X_000070 ^ 2 + 3)) |
     X_000071 <= arctan(X_000070)),
    inference(resolve, [], [refute_0_61, refute_0_60])).

cnf(refute_0_63, plain,
    (X_000070 * 3 / (3 + X_000070 * X_000070) < X_000071 | X_000070 < 0 |
     X_000071 <= arctan(X_000070)),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_64, refute_0_12])).

cnf(refute_0_66, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (3 + skoXC1 * skoXC1) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_59, refute_0_65])).

cnf(refute_0_67, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) | skoXC1 * skoXC1 <= -3),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * 3 <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (3 + skoXC1 * skoXC1) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) | 3 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_69, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_63])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (3 + skoXC1 * skoXC1) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) | skoXC1 * skoXC1 <= -3),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (~ lgen(0, X_000043, arctan(X_000042)) | X_000043 <= arctan(X_000042)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_73, plain,
    (~ lgen(0, X_000043, X_000042) | 0 <= X_000042 |
     lgen(0, X_000043, arctan(X_000042))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_74, plain,
    (~ lgen(0, X_000043, X_000042) | X_000043 <= arctan(X_000042) |
     0 <= X_000042), inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    (X_000042 < X_000043 | X_000043 <= arctan(X_000042) | 0 <= X_000042),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_76, refute_0_12])).

cnf(refute_0_78, plain,
    (3/200000 <
     skoXC1 *
     (-67/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain, (abs(skoXC1) <= 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_80, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_81, plain,
    (1 < abs(skoXC1) | abs(skoXC1) != skoXC1 | skoXC1 <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_82, plain, (skoXC1 < 0 | 1 < abs(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_80, refute_0_81])).

cnf(refute_0_83, plain, (skoXC1 < 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_79, refute_0_82])).

cnf(refute_0_84, plain, (0 <= skoXC1 | skoXC1 <= 1),
    introduced(tautology, [decision])).

cnf(refute_0_85, plain, (skoXC1 <= 1),
    inference(resolve, [], [refute_0_84, refute_0_83])).

cnf(refute_0_86, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_87, plain,
    (1 < abs(skoXC1) | abs(skoXC1) != -skoXC1 | -skoXC1 <= 1),
    introduced(tautology, [equality])).

cnf(refute_0_88, plain, (1 < abs(skoXC1) | 0 <= skoXC1 | -skoXC1 <= 1),
    inference(resolve, [], [refute_0_86, refute_0_87])).

cnf(refute_0_89, plain, (0 <= skoXC1 | -skoXC1 <= 1),
    inference(resolve, [], [refute_0_79, refute_0_88])).

cnf(refute_0_90, plain, (-1 <= skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain, (skoXC1 < 0 | -1 <= skoXC1),
    inference(decision, [], [refute_0_85])).

cnf(refute_0_92, plain, (-1 <= skoXC1),
    inference(resolve, [], [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     -3 < skoXC1 * skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [], [refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_94, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_71, refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     -3 < skoXC1 * skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_94, refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_96, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_67, refute_0_95])).

cnf(refute_0_97, plain,
    (-63 < skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) |
     0 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_96, refute_0_94, refute_0_78, refute_0_85,
               refute_0_92])).

cnf(refute_0_98, plain,
    (0 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_58, refute_0_97])).

cnf(refute_0_99, plain,
    (-63 < skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) |
     189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_98, refute_0_96, refute_0_94, refute_0_78,
               refute_0_85, refute_0_92])).

cnf(refute_0_100, plain,
    (189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_54, refute_0_99])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 |
     -231 <
     skoXC1 *
     (skoXC1 * (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) |
     0 <
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_100, refute_0_98, refute_0_96, refute_0_94,
               refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_102, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_45, refute_0_101])).

cnf(refute_0_103, plain,
    (-429 <
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) |
     0 <
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_102, refute_0_100, refute_0_98, refute_0_96,
               refute_0_94, refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_104, plain,
    (0 <
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_37, refute_0_103])).

cnf(refute_0_105, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) <=
     1287/200000 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_104, refute_0_102, refute_0_100, refute_0_98,
               refute_0_96, refute_0_94, refute_0_78, refute_0_85,
               refute_0_92])).

cnf(refute_0_106, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_105, refute_0_33])).

cnf(refute_0_107, plain,
    (-429 <
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_104, refute_0_102, refute_0_100, refute_0_98,
               refute_0_96, refute_0_94, refute_0_78, refute_0_85,
               refute_0_92])).

cnf(refute_0_108, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(resolve, [], [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(resolve, [], [refute_0_108, refute_0_24])).

cnf(refute_0_110, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_109])).

cnf(refute_0_111, plain,
    (3861/40000 <
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_113, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(subst, [], [refute_0_4])).

cnf(refute_0_114, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(resolve, [], [refute_0_108, refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_112, refute_0_114])).

cnf(refute_0_116, plain,
    (0 <
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)), inference(arithmetic, [], [refute_0_115])).

cnf(refute_0_117, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     0 <
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_108, refute_0_104, refute_0_102, refute_0_100,
               refute_0_98, refute_0_96, refute_0_94, refute_0_78,
               refute_0_85, refute_0_92])).

cnf(refute_0_118, plain,
    (0 <
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(resolve, [], [refute_0_116, refute_0_117])).

cnf(refute_0_119, plain,
    (skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     3861/40000 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_118, refute_0_108, refute_0_104, refute_0_102,
               refute_0_100, refute_0_98, refute_0_96, refute_0_94,
               refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_120, plain,
    (skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(resolve, [], [refute_0_119, refute_0_111])).

cnf(refute_0_121, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(decision, [],
              [refute_0_118, refute_0_108, refute_0_104, refute_0_102,
               refute_0_100, refute_0_98, refute_0_96, refute_0_94,
               refute_0_78, refute_0_85, refute_0_92])).

cnf(refute_0_122, plain,
    (skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     arctan(skoXC1)),
    inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     arctan(skoXC1)), inference(resolve, [], [refute_0_122, refute_0_23])).

cnf(refute_0_124, plain,
    (3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_18, refute_0_123])).

cnf(refute_0_125, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) |
     0 <= skoXC1 |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_13, refute_0_124])).

cnf(refute_0_126, plain,
    (3861/40000 <
     skoXC1 *
     (86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (-35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (-538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (-21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (-23272571/1000000 +
                    skoXC1 * (skoXC1 * -1458457/2000000)))))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435), inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_128, plain,
    (arctan(X_000175) < X_000176 | ~ lgen(1, arctan(X_000175), X_000176)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_129, plain,
    (~ lgen(1,
            1/35 *
            (15159 * X_000175 ^ 6 + 147455 * X_000175 ^ 4 +
             345345 * X_000175 ^ 2 + 225225) * X_000175 /
            (35 * X_000175 ^ 8 + 1260 * X_000175 ^ 6 +
             6930 * X_000175 ^ 4 + 12012 * X_000175 ^ 2 + 6435),
            X_000176) | 0 <= X_000175 |
     lgen(1, arctan(X_000175), X_000176)),
    inference(subst, [], [atan_upper_bound_case_42])).

cnf(refute_0_130, plain,
    (arctan(X_000175) < X_000176 |
     ~ lgen(1,
            1/35 *
            (15159 * X_000175 ^ 6 + 147455 * X_000175 ^ 4 +
             345345 * X_000175 ^ 2 + 225225) * X_000175 /
            (35 * X_000175 ^ 8 + 1260 * X_000175 ^ 6 +
             6930 * X_000175 ^ 4 + 12012 * X_000175 ^ 2 + 6435),
            X_000176) | 0 <= X_000175),
    inference(resolve, [], [refute_0_129, refute_0_128])).

cnf(refute_0_131, plain,
    (arctan(X_000175) < X_000176 |
     X_000176 <=
     X_000175 *
     (6435 +
      X_000175 *
      (X_000175 *
       (9867 +
        X_000175 *
        (X_000175 * (4213 + X_000175 * (X_000175 * 15159/35)))))) /
     (6435 +
      X_000175 *
      (X_000175 *
       (12012 +
        X_000175 *
        (X_000175 *
         (6930 +
          X_000175 *
          (X_000175 * (1260 + X_000175 * (X_000175 * 35)))))))) |
     0 <= X_000175), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))))),
    inference(subst, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_132, refute_0_127])).

cnf(refute_0_134, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     0), inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0 |
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_136, plain,
    (~ lgen(0, arctan(X_000147), X_000148) | arctan(X_000147) <= X_000148),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_137, plain,
    (X_000147 < 0 |
     ~ lgen(0,
            1/35 *
            (256 * X_000147 ^ 6 + 5943 * X_000147 ^ 4 +
             19250 * X_000147 ^ 2 + 15015) * X_000147 /
            (35 * X_000147 ^ 6 + 315 * X_000147 ^ 4 + 693 * X_000147 ^ 2 +
             429), X_000148) | lgen(0, arctan(X_000147), X_000148)),
    inference(subst, [], [atan_upper_bound_case_43])).

cnf(refute_0_138, plain,
    (X_000147 < 0 |
     ~ lgen(0,
            1/35 *
            (256 * X_000147 ^ 6 + 5943 * X_000147 ^ 4 +
             19250 * X_000147 ^ 2 + 15015) * X_000147 /
            (35 * X_000147 ^ 6 + 315 * X_000147 ^ 4 + 693 * X_000147 ^ 2 +
             429), X_000148) | arctan(X_000147) <= X_000148),
    inference(resolve, [], [refute_0_137, refute_0_136])).

cnf(refute_0_139, plain,
    (X_000147 < 0 |
     X_000148 <
     X_000147 *
     (429 +
      X_000147 *
      (X_000147 *
       (550 +
        X_000147 *
        (X_000147 * (849/5 + X_000147 * (X_000147 * 256/35)))))) /
     (429 +
      X_000147 *
      (X_000147 *
       (693 +
        X_000147 * (X_000147 * (315 + X_000147 * (X_000147 * 35)))))) |
     arctan(X_000147) <= X_000148),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     arctan(skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))))),
    inference(resolve, [], [refute_0_140, refute_0_123])).

cnf(refute_0_142, plain,
    (skoXC1 < 0 |
     (3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(resolve, [], [refute_0_135, refute_0_141])).

cnf(refute_0_143, plain,
    (skoXC1 < 0 |
     1287/200000 <
     skoXC1 *
     (28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (-483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (-167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (-7166803/2000000 +
                  skoXC1 * (skoXC1 * -1458457/2000000)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429),
    inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <=
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_145, plain,
    (arctan(X_000147) < X_000148 | ~ lgen(1, arctan(X_000147), X_000148)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_146, plain,
    (X_000147 < 0 |
     ~ lgen(1,
            1/35 *
            (256 * X_000147 ^ 6 + 5943 * X_000147 ^ 4 +
             19250 * X_000147 ^ 2 + 15015) * X_000147 /
            (35 * X_000147 ^ 6 + 315 * X_000147 ^ 4 + 693 * X_000147 ^ 2 +
             429), X_000148) | lgen(1, arctan(X_000147), X_000148)),
    inference(subst, [], [atan_upper_bound_case_43])).

cnf(refute_0_147, plain,
    (arctan(X_000147) < X_000148 | X_000147 < 0 |
     ~ lgen(1,
            1/35 *
            (256 * X_000147 ^ 6 + 5943 * X_000147 ^ 4 +
             19250 * X_000147 ^ 2 + 15015) * X_000147 /
            (35 * X_000147 ^ 6 + 315 * X_000147 ^ 4 + 693 * X_000147 ^ 2 +
             429), X_000148)),
    inference(resolve, [], [refute_0_146, refute_0_145])).

cnf(refute_0_148, plain,
    (arctan(X_000147) < X_000148 | X_000147 < 0 |
     X_000148 <=
     X_000147 *
     (429 +
      X_000147 *
      (X_000147 *
       (550 +
        X_000147 *
        (X_000147 * (849/5 + X_000147 * (X_000147 * 256/35)))))) /
     (429 +
      X_000147 *
      (X_000147 *
       (693 +
        X_000147 * (X_000147 * (315 + X_000147 * (X_000147 * 35))))))),
    inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))))),
    inference(subst, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <=
     skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(resolve, [], [refute_0_149, refute_0_144])).

cnf(refute_0_151, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) /
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) <=
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) |
     231 +
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_153, plain,
    (arctan(X_000119) < X_000120 | ~ lgen(1, arctan(X_000119), X_000120)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_154, plain,
    (~ lgen(1,
            7/5 * (33 * X_000119 ^ 4 + 170 * X_000119 ^ 2 + 165) *
            X_000119 /
            (5 * X_000119 ^ 6 + 105 * X_000119 ^ 4 + 315 * X_000119 ^ 2 +
             231), X_000120) | 0 <= X_000119 |
     lgen(1, arctan(X_000119), X_000120)),
    inference(subst, [], [atan_upper_bound_case_32])).

cnf(refute_0_155, plain,
    (arctan(X_000119) < X_000120 |
     ~ lgen(1,
            7/5 * (33 * X_000119 ^ 4 + 170 * X_000119 ^ 2 + 165) *
            X_000119 /
            (5 * X_000119 ^ 6 + 105 * X_000119 ^ 4 + 315 * X_000119 ^ 2 +
             231), X_000120) | 0 <= X_000119),
    inference(resolve, [], [refute_0_154, refute_0_153])).

cnf(refute_0_156, plain,
    (arctan(X_000119) < X_000120 |
     X_000120 <=
     X_000119 *
     (231 +
      X_000119 * (X_000119 * (238 + X_000119 * (X_000119 * 231/5)))) /
     (231 +
      X_000119 *
      (X_000119 *
       (315 +
        X_000119 * (X_000119 * (105 + X_000119 * (X_000119 * 5)))))) |
     0 <= X_000119), inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) /
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))))),
    inference(subst, [], [refute_0_156])).

cnf(refute_0_158, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (231 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5)))))) <=
     skoXC1 *
     (231 + skoXC1 * (skoXC1 * (238 + skoXC1 * (skoXC1 * 231/5)))) |
     231 +
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= 0),
    inference(resolve, [], [refute_0_157, refute_0_152])).

cnf(refute_0_159, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) <= -231 |
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0 |
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_161, plain,
    (~ lgen(0, arctan(X_000091), X_000092) | arctan(X_000091) <= X_000092),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_162, plain,
    (X_000091 < 0 |
     ~ lgen(0,
            1/15 * (64 * X_000091 ^ 4 + 735 * X_000091 ^ 2 + 945) *
            X_000091 / (15 * X_000091 ^ 4 + 70 * X_000091 ^ 2 + 63),
            X_000092) | lgen(0, arctan(X_000091), X_000092)),
    inference(subst, [], [atan_upper_bound_case_33])).

cnf(refute_0_163, plain,
    (X_000091 < 0 |
     ~ lgen(0,
            1/15 * (64 * X_000091 ^ 4 + 735 * X_000091 ^ 2 + 945) *
            X_000091 / (15 * X_000091 ^ 4 + 70 * X_000091 ^ 2 + 63),
            X_000092) | arctan(X_000091) <= X_000092),
    inference(resolve, [], [refute_0_162, refute_0_161])).

cnf(refute_0_164, plain,
    (X_000091 < 0 |
     X_000092 <
     X_000091 *
     (63 + X_000091 * (X_000091 * (49 + X_000091 * (X_000091 * 64/15)))) /
     (63 + X_000091 * (X_000091 * (70 + X_000091 * (X_000091 * 15)))) |
     arctan(X_000091) <= X_000092),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     arctan(skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [refute_0_164])).

cnf(refute_0_166, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))))),
    inference(resolve, [], [refute_0_165, refute_0_123])).

cnf(refute_0_167, plain,
    (skoXC1 < 0 |
     (3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(resolve, [], [refute_0_160, refute_0_166])).

cnf(refute_0_168, plain,
    (skoXC1 < 0 |
     189/200000 <
     skoXC1 *
     (4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (-363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (780569/750000 +
          skoXC1 *
          (skoXC1 *
           (-4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (-90731/500000 +
                skoXC1 * (skoXC1 * -625053/2000000)))))))))))) |
     skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= -63),
    inference(arithmetic, [], [refute_0_167])).

cnf(refute_0_169, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <=
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_170, plain,
    (arctan(X_000091) < X_000092 | ~ lgen(1, arctan(X_000091), X_000092)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_171, plain,
    (X_000091 < 0 |
     ~ lgen(1,
            1/15 * (64 * X_000091 ^ 4 + 735 * X_000091 ^ 2 + 945) *
            X_000091 / (15 * X_000091 ^ 4 + 70 * X_000091 ^ 2 + 63),
            X_000092) | lgen(1, arctan(X_000091), X_000092)),
    inference(subst, [], [atan_upper_bound_case_33])).

cnf(refute_0_172, plain,
    (arctan(X_000091) < X_000092 | X_000091 < 0 |
     ~ lgen(1,
            1/15 * (64 * X_000091 ^ 4 + 735 * X_000091 ^ 2 + 945) *
            X_000091 / (15 * X_000091 ^ 4 + 70 * X_000091 ^ 2 + 63),
            X_000092)),
    inference(resolve, [], [refute_0_171, refute_0_170])).

cnf(refute_0_173, plain,
    (arctan(X_000091) < X_000092 | X_000091 < 0 |
     X_000092 <=
     X_000091 *
     (63 + X_000091 * (X_000091 * (49 + X_000091 * (X_000091 * 64/15)))) /
     (63 + X_000091 * (X_000091 * (70 + X_000091 * (X_000091 * 15))))),
    inference(arithmetic, [], [refute_0_172])).

cnf(refute_0_174, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <=
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))))),
    inference(subst, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <=
     skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(resolve, [], [refute_0_174, refute_0_169])).

cnf(refute_0_176, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= -63 |
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_175])).

cnf(refute_0_177, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (3 + skoXC1 * skoXC1) < skoXC1 * 3 | 3 + skoXC1 * skoXC1 <= 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_178, plain,
    (~ lgen(0, arctan(X_000063), X_000064) | arctan(X_000063) <= X_000064),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_179, plain,
    (~ lgen(0, 3 * X_000063 / (X_000063 ^ 2 + 3), X_000064) |
     0 <= X_000063 | lgen(0, arctan(X_000063), X_000064)),
    inference(subst, [], [atan_upper_bound_case_12])).

cnf(refute_0_180, plain,
    (~ lgen(0, 3 * X_000063 / (X_000063 ^ 2 + 3), X_000064) |
     0 <= X_000063 | arctan(X_000063) <= X_000064),
    inference(resolve, [], [refute_0_179, refute_0_178])).

cnf(refute_0_181, plain,
    (X_000064 < X_000063 * 3 / (3 + X_000063 * X_000063) | 0 <= X_000063 |
     arctan(X_000063) <= X_000064),
    inference(arithmetic, [], [refute_0_180])).

cnf(refute_0_182, plain,
    (3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) | 0 <= skoXC1 |
     arctan(skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_182, refute_0_123])).

cnf(refute_0_184, plain,
    ((3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (3 + skoXC1 * skoXC1) < skoXC1 * 3 | 0 <= skoXC1 |
     3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_177, refute_0_183])).

cnf(refute_0_185, plain,
    (9/200000 <
     skoXC1 *
     (201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (-3587/400000 +
        skoXC1 *
        (skoXC1 *
         (-420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (37629/500000 +
            skoXC1 *
            (skoXC1 *
             (226277/10000000 +
              skoXC1 * (skoXC1 * -208351/10000000)))))))))) | 0 <= skoXC1 |
     skoXC1 * skoXC1 <= -3), inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (~ lgen(0, arctan(X_000035), X_000036) | arctan(X_000035) <= X_000036),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_187, plain,
    (X_000035 < 0 | ~ lgen(0, X_000035, X_000036) |
     lgen(0, arctan(X_000035), X_000036)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_188, plain,
    (X_000035 < 0 | ~ lgen(0, X_000035, X_000036) |
     arctan(X_000035) <= X_000036),
    inference(resolve, [], [refute_0_187, refute_0_186])).

cnf(refute_0_189, plain,
    (X_000036 < X_000035 | X_000035 < 0 | arctan(X_000035) <= X_000036),
    inference(arithmetic, [], [refute_0_188])).

cnf(refute_0_190, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1 |
     arctan(skoXC1) <=
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(subst, [], [refute_0_189])).

cnf(refute_0_191, plain,
    (skoXC1 < 0 |
     3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) <
     skoXC1), inference(resolve, [], [refute_0_190, refute_0_123])).

cnf(refute_0_192, plain,
    (skoXC1 < 0 |
     3/200000 <
     skoXC1 *
     (67/500000 +
      skoXC1 *
      (skoXC1 *
       (660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (-180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (85133/1000000 + skoXC1 * (skoXC1 * -208351/10000000))))))))),
    inference(arithmetic, [], [refute_0_191])).

cnf(refute_0_193, plain,
    (9/200000 <
     skoXC1 *
     (201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (-3587/400000 +
        skoXC1 *
        (skoXC1 *
         (-420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (37629/500000 +
            skoXC1 *
            (skoXC1 *
             (226277/10000000 +
              skoXC1 * (skoXC1 * -208351/10000000)))))))))) |
     -3 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_192, refute_0_92, refute_0_85])).

cnf(refute_0_194, plain,
    (9/200000 <
     skoXC1 *
     (201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (-3587/400000 +
        skoXC1 *
        (skoXC1 *
         (-420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (37629/500000 +
            skoXC1 *
            (skoXC1 *
             (226277/10000000 +
              skoXC1 * (skoXC1 * -208351/10000000)))))))))) | 0 <= skoXC1),
    inference(resolve, [], [refute_0_185, refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 < 0 | -63 < skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) |
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) <= 0),
    inference(decision, [],
              [refute_0_194, refute_0_192, refute_0_92, refute_0_85])).

cnf(refute_0_196, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) <= 0),
    inference(resolve, [], [refute_0_176, refute_0_195])).

cnf(refute_0_197, plain,
    (skoXC1 < 0 | -63 < skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) |
     189/200000 <
     skoXC1 *
     (4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (-363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (780569/750000 +
          skoXC1 *
          (skoXC1 *
           (-4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (-90731/500000 +
                skoXC1 * (skoXC1 * -625053/2000000)))))))))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(decision, [],
              [refute_0_196, refute_0_194, refute_0_192, refute_0_92,
               refute_0_85])).

cnf(refute_0_198, plain,
    (skoXC1 < 0 |
     189/200000 <
     skoXC1 *
     (4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (-363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (780569/750000 +
          skoXC1 *
          (skoXC1 *
           (-4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (-90731/500000 +
                skoXC1 * (skoXC1 * -625053/2000000)))))))))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_168, refute_0_197])).

cnf(refute_0_199, plain,
    (-231 <
     skoXC1 *
     (skoXC1 * (315 + skoXC1 * (skoXC1 * (105 + skoXC1 * (skoXC1 * 5))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) <= 0),
    inference(decision, [],
              [refute_0_198, refute_0_196, refute_0_194, refute_0_192,
               refute_0_92, refute_0_85])).

cnf(refute_0_200, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (-15477/500000 +
      skoXC1 *
      (skoXC1 *
       (1317211/2000000 +
        skoXC1 *
        (skoXC1 *
         (-7291683/2000000 +
          skoXC1 *
          (skoXC1 *
           (14793149/2000000 +
            skoXC1 *
            (skoXC1 *
             (-23703397/5000000 +
              skoXC1 *
              (skoXC1 *
               (-2950407/2000000 +
                skoXC1 *
                (skoXC1 *
                 (3524041/2000000 +
                  skoXC1 * (skoXC1 * 208351/2000000)))))))))))))) <= 0),
    inference(resolve, [], [refute_0_159, refute_0_199])).

cnf(refute_0_201, plain,
    (skoXC1 < 0 |
     -429 <
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) <= 0),
    inference(decision, [],
              [refute_0_200, refute_0_198, refute_0_196, refute_0_194,
               refute_0_192, refute_0_92, refute_0_85])).

cnf(refute_0_202, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (skoXC1 *
       (483461/400000 +
        skoXC1 *
        (skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) <= 0),
    inference(resolve, [], [refute_0_151, refute_0_201])).

cnf(refute_0_203, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (-483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (-167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (-7166803/2000000 +
                  skoXC1 * (skoXC1 * -1458457/2000000)))))))))))))) <=
     1287/200000),
    inference(decision, [],
              [refute_0_202, refute_0_200, refute_0_198, refute_0_196,
               refute_0_194, refute_0_192, refute_0_92, refute_0_85])).

cnf(refute_0_204, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429),
    inference(resolve, [], [refute_0_203, refute_0_143])).

cnf(refute_0_205, plain,
    (skoXC1 < 0 |
     -429 <
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(decision, [],
              [refute_0_202, refute_0_200, refute_0_198, refute_0_196,
               refute_0_194, refute_0_192, refute_0_92, refute_0_85])).

cnf(refute_0_206, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_204, refute_0_205])).

cnf(refute_0_207, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     0),
    inference(decision, [],
              [refute_0_206, refute_0_202, refute_0_200, refute_0_198,
               refute_0_196, refute_0_194, refute_0_192, refute_0_92,
               refute_0_85])).

cnf(refute_0_208, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     0), inference(resolve, [], [refute_0_134, refute_0_207])).

cnf(refute_0_209, plain,
    (skoXC1 < 0 |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     0),
    inference(decision, [],
              [refute_0_206, refute_0_202, refute_0_200, refute_0_198,
               refute_0_196, refute_0_194, refute_0_192, refute_0_92,
               refute_0_85])).

cnf(refute_0_210, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     0), inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (-35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (-538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (-21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (-23272571/1000000 +
                    skoXC1 * (skoXC1 * -1458457/2000000)))))))))))))))) <=
     3861/40000),
    inference(decision, [],
              [refute_0_210, refute_0_206, refute_0_202, refute_0_200,
               refute_0_198, refute_0_196, refute_0_194, refute_0_192,
               refute_0_92, refute_0_85])).

cnf(refute_0_212, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435), inference(resolve, [], [refute_0_211, refute_0_126])).

cnf(refute_0_213, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1),
    inference(decision, [],
              [refute_0_210, refute_0_206, refute_0_202, refute_0_200,
               refute_0_198, refute_0_196, refute_0_194, refute_0_192,
               refute_0_92, refute_0_85])).

cnf(refute_0_214, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_212, refute_0_213])).

cnf(refute_0_215, plain,
    (arctan(skoXC1) <
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_214, refute_0_206])).

cnf(refute_0_216, plain,
    (arctan(skoXC1) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_12, refute_0_215])).

cnf(refute_0_217, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_5, refute_0_216])).

cnf(refute_0_218, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_217])).

cnf(refute_0_219, plain,
    (skoXC1 < 0 |
     3861/40000 <
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435), inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) /
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_30, refute_0_216])).

cnf(refute_0_221, plain,
    (skoXC1 *
     (429 +
      skoXC1 *
      (skoXC1 *
       (550 + skoXC1 * (skoXC1 * (849/5 + skoXC1 * (skoXC1 * 256/35)))))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (429 +
      skoXC1 *
      (skoXC1 *
       (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35)))))) |
     0 <= skoXC1 |
     429 +
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= 0),
    inference(resolve, [], [refute_0_25, refute_0_220])).

cnf(refute_0_222, plain,
    (1287/200000 <
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) |
     0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429),
    inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) /
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000)))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_51, refute_0_216])).

cnf(refute_0_224, plain,
    (skoXC1 * (63 + skoXC1 * (skoXC1 * (49 + skoXC1 * (skoXC1 * 64/15)))) <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15)))) |
     0 <= skoXC1 |
     63 + skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= 0),
    inference(resolve, [], [refute_0_46, refute_0_223])).

cnf(refute_0_225, plain,
    (189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1 |
     skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) <= -63),
    inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     -3/200000 +
     skoXC1 *
     (499933/500000 +
      skoXC1 *
      (skoXC1 *
       (-660599/2000000 +
        skoXC1 *
        (skoXC1 *
         (180141/1000000 +
          skoXC1 *
          (skoXC1 *
           (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))),
    inference(resolve, [], [refute_0_64, refute_0_216])).

cnf(refute_0_227, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 <
     (-3/200000 +
      skoXC1 *
      (499933/500000 +
       skoXC1 *
       (skoXC1 *
        (-660599/2000000 +
         skoXC1 *
         (skoXC1 *
          (180141/1000000 +
           skoXC1 *
           (skoXC1 *
            (-85133/1000000 + skoXC1 * (skoXC1 * 208351/10000000))))))))) *
     (3 + skoXC1 * skoXC1) | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_59, refute_0_226])).

cnf(refute_0_228, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     skoXC1 * skoXC1 <= -3), inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000)))))))))) |
     -3 < skoXC1 * skoXC1),
    inference(decision, [],
              [refute_0_85, refute_0_92, refute_0_192, refute_0_194])).

cnf(refute_0_230, plain,
    (skoXC1 < 0 |
     9/200000 <
     skoXC1 *
     (-201/500000 +
      skoXC1 *
      (-3/200000 +
       skoXC1 *
       (3587/400000 +
        skoXC1 *
        (skoXC1 *
         (420247/2000000 +
          skoXC1 *
          (skoXC1 *
           (-37629/500000 +
            skoXC1 *
            (skoXC1 *
             (-226277/10000000 +
              skoXC1 * (skoXC1 * 208351/10000000))))))))))),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    (-63 < skoXC1 * (skoXC1 * (70 + skoXC1 * (skoXC1 * 15))) |
     189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1),
    inference(decision, [],
              [refute_0_230, refute_0_85, refute_0_92, refute_0_192,
               refute_0_194])).

cnf(refute_0_232, plain,
    (189/200000 <
     skoXC1 *
     (-4221/500000 +
      skoXC1 *
      (-21/20000 +
       skoXC1 *
       (363503/2000000 +
        skoXC1 *
        (-9/40000 +
         skoXC1 *
         (-780569/750000 +
          skoXC1 *
          (skoXC1 *
           (4583997/2000000 +
            skoXC1 *
            (skoXC1 *
             (-19445837/10000000 +
              skoXC1 *
              (skoXC1 *
               (90731/500000 +
                skoXC1 * (skoXC1 * 625053/2000000)))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_225, refute_0_231])).

cnf(refute_0_233, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429 |
     skoXC1 *
     (-28743/500000 +
      skoXC1 *
      (-2079/200000 +
       skoXC1 *
       (483461/400000 +
        skoXC1 *
        (-189/40000 +
         skoXC1 *
         (-12918549/2000000 +
          skoXC1 *
          (-21/40000 +
           skoXC1 *
           (167332729/14000000 +
            skoXC1 *
            (skoXC1 *
             (-24374893/5000000 +
              skoXC1 *
              (skoXC1 *
               (-60732357/10000000 +
                skoXC1 *
                (skoXC1 *
                 (7166803/2000000 +
                  skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))) <=
     1287/200000),
    inference(decision, [],
              [refute_0_232, refute_0_230, refute_0_85, refute_0_92,
               refute_0_192, refute_0_194])).

cnf(refute_0_234, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) <= -429),
    inference(resolve, [], [refute_0_233, refute_0_222])).

cnf(refute_0_235, plain,
    (-429 <
     skoXC1 *
     (skoXC1 *
      (693 + skoXC1 * (skoXC1 * (315 + skoXC1 * (skoXC1 * 35))))) |
     0 <= skoXC1),
    inference(decision, [],
              [refute_0_232, refute_0_230, refute_0_85, refute_0_92,
               refute_0_192, refute_0_194])).

cnf(refute_0_236, plain, (0 <= skoXC1),
    inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     skoXC1 *
     (-86229/100000 +
      skoXC1 *
      (-9009/50000 +
       skoXC1 *
       (35826219/2000000 +
        skoXC1 *
        (-2079/20000 +
         skoXC1 *
         (-92278879/1000000 +
          skoXC1 *
          (-189/10000 +
           skoXC1 *
           (538174617/3500000 +
            skoXC1 *
            (-21/40000 +
             skoXC1 *
             (-8539463/400000 +
              skoXC1 *
              (skoXC1 *
               (-1242832913/10000000 +
                skoXC1 *
                (skoXC1 *
                 (21712299/500000 +
                  skoXC1 *
                  (skoXC1 *
                   (23272571/1000000 +
                    skoXC1 * (skoXC1 * 1458457/2000000)))))))))))))))) <=
     3861/40000),
    inference(decision, [],
              [refute_0_236, refute_0_232, refute_0_230, refute_0_85,
               refute_0_92, refute_0_192, refute_0_194])).

cnf(refute_0_238, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435), inference(resolve, [], [refute_0_237, refute_0_219])).

cnf(refute_0_239, plain,
    (skoXC1 < 0 |
     -6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))),
    inference(decision, [],
              [refute_0_236, refute_0_232, refute_0_230, refute_0_85,
               refute_0_92, refute_0_192, refute_0_194])).

cnf(refute_0_240, plain, (skoXC1 < 0),
    inference(resolve, [], [refute_0_238, refute_0_239])).

cnf(refute_0_241, plain, ($false),
    inference(resolve, [], [refute_0_236, refute_0_240])).
% SZS output end CNFRefutation for atan-error-analysis-2-weak.tptp
