% SZS output start CNFRefutation for atan-problem-1-sqrt-max.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_43, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            1/35 *
            (15159 * X ^ 6 + 147455 * X ^ 4 + 345345 * X ^ 2 + 225225) *
            X /
            (35 * X ^ 8 + 1260 * X ^ 6 + 6930 * X ^ 4 + 12012 * X ^ 2 +
             6435)) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_44, axiom,
    (X <= 1 |
     ~ lgen(R, Y,
            9817477/6250000 -
            1/35 * (256 + 5943 * X ^ 2 + 19250 * X ^ 4 + 15015 * X ^ 6) /
            (X * (35 + 315 * X ^ 2 + 693 * X ^ 4 + 429 * X ^ 6))) |
     lgen(R, Y, arctan(X)))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R, Y,
            32 * X * (X + 16) * (X ^ 2 + 96 * X + 256) /
            (X ^ 4 + 448 * X ^ 3 + 17920 * X ^ 2 + 114688 * X + 65536)) |
     lgen(R, Y, sqrt(X)))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(max_1, axiom, (X < Y | max(X, Y) = X)).

cnf(max_2, axiom, (Y <= X | max(X, Y) = Y)).

fof(atan_problem_1, conjecture,
    (! [X] :
       (0 < X =>
        8 * sqrt(3) * X / (3 * sqrt(3) + max(43/5, 89/10 * X)) <
        203/1000 + arctan(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 < X =>
        8 * sqrt(3) * X / (3 * sqrt(3) + max(43/5, 89/10 * X)) <
        203/1000 + arctan(X))), inference(strip, [], [atan_problem_1])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 < X =>
          8 * sqrt(3) * X / (3 * sqrt(3) + max(43/5, 89/10 * X)) <
          203/1000 + arctan(X))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X &
        203/1000 + arctan(X) <=
        8 * sqrt(3) * X / (3 * sqrt(3) + max(43/5, 89/10 * X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 < skoXC1 &
     203/1000 + arctan(skoXC1) <=
     8 * sqrt(3) * skoXC1 / (3 * sqrt(3) + max(43/5, 89/10 * skoXC1))),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (203/1000 + arctan(skoXC1) <=
     8 * sqrt(3) * skoXC1 / (3 * sqrt(3) + max(43/5, 89/10 * skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (583057/336224 <
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) |
     0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     583057/336224 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (sqrt(X_000063) < X_000064 | ~ lgen(1, sqrt(X_000063), X_000064)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain,
    (X_000063 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000063 + 17920 * X_000063 ^ 2 +
             448 * X_000063 ^ 3 + X_000063 ^ 4) /
            ((X_000063 + 16) * (256 + 96 * X_000063 + X_000063 ^ 2)),
            X_000064) | lgen(1, sqrt(X_000063), X_000064)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_3, plain,
    (sqrt(X_000063) < X_000064 | X_000063 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000063 + 17920 * X_000063 ^ 2 +
             448 * X_000063 ^ 3 + X_000063 ^ 4) /
            ((X_000063 + 16) * (256 + 96 * X_000063 + X_000063 ^ 2)),
            X_000064)), inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (sqrt(X_000063) < X_000064 | X_000063 < 1 |
     X_000064 <=
     (2048 +
      X_000063 *
      (3584 + X_000063 * (560 + X_000063 * (14 + X_000063 * 1/32)))) /
     (4096 + X_000063 * (1792 + X_000063 * (112 + X_000063)))),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (3 < 1 |
     sqrt(3) <
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) |
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) <
     (skoXC1 *
      (23252229/2000 +
       skoXC1 *
       (skoXC1 *
        (54255201/2500 +
         skoXC1 *
         (skoXC1 *
          (12520431/1000 +
           skoXC1 *
           (skoXC1 *
            (1138221/500 + skoXC1 * (skoXC1 * 126469/2000)))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     skoXC1 *
     (23252229/2000 +
      skoXC1 *
      (skoXC1 *
       (54255201/2500 +
        skoXC1 *
        (skoXC1 *
         (12520431/1000 +
          skoXC1 *
          (skoXC1 * (1138221/500 + skoXC1 * (skoXC1 * 126469/2000)))))))) +
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-51480 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-96096 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-55440 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-10080 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) *
     (skoXC1 * 89/10 + sqrt(3) * 3)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_7, plain,
    ((skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    (203/1000 + arctan(skoXC1) <=
     8 * sqrt(3) * skoXC1 / (3 * sqrt(3) + max(43/5, 89/10 * skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    (203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (max(43/5, skoXC1 * 89/10) + sqrt(3) * 3)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (skoXC1 * 89/10 <= 43/5 | max(43/5, skoXC1 * 89/10) = skoXC1 * 89/10),
    inference(subst, [], [max_2])).

cnf(refute_0_11, plain,
    (sqrt(3) * (skoXC1 * 8) / (max(43/5, skoXC1 * 89/10) + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) |
     max(43/5, skoXC1 * 89/10) != skoXC1 * 89/10 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (skoXC1 * 89/10 + sqrt(3) * 3)),
    introduced(tautology, [equality])).

cnf(refute_0_12, plain,
    (sqrt(3) * (skoXC1 * 8) / (max(43/5, skoXC1 * 89/10) + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) | skoXC1 * 89/10 <= 43/5 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (skoXC1 * 89/10 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (skoXC1 * 89/10 <= 43/5 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (skoXC1 * 89/10 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_9, refute_0_12])).

cnf(refute_0_14, plain,
    (203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (skoXC1 * 89/10 + sqrt(3) * 3) |
     skoXC1 <= 86/89), inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (sqrt(3) * (skoXC1 * 8) / (skoXC1 * 89/10 + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) |
     (203/1000 + arctan(skoXC1)) * (skoXC1 * 89/10 + sqrt(3) * 3) <=
     sqrt(3) * (skoXC1 * 8) | skoXC1 * 89/10 + sqrt(3) * 3 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_16, plain,
    ((203/1000 + arctan(skoXC1)) * (skoXC1 * 89/10 + sqrt(3) * 3) <=
     sqrt(3) * (skoXC1 * 8) | skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     skoXC1 <= 86/89), inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain,
    (arctan(skoXC1) * (skoXC1 * 89/10 + sqrt(3) * 3) <=
     skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8) |
     skoXC1 <= 86/89 | sqrt(3) <= skoXC1 * -89/30),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (X_000050 < sqrt(X_000049) | ~ lgen(1, X_000050, sqrt(X_000049))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_19, plain,
    (X_000049 < 0 | ~ lgen(1, X_000050, 2 * X_000049 / (X_000049 + 1)) |
     lgen(1, X_000050, sqrt(X_000049))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_20, plain,
    (X_000049 < 0 | X_000050 < sqrt(X_000049) |
     ~ lgen(1, X_000050, 2 * X_000049 / (X_000049 + 1))),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (X_000049 < 0 | X_000050 < sqrt(X_000049) |
     X_000049 * 2 / (1 + X_000049) <= X_000050),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (3 < 0 | skoXC1 * -89/30 < sqrt(3) |
     3 * 2 / (1 + 3) <= skoXC1 * -89/30),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 * -89/30 < sqrt(3) | skoXC1 <= -45/89),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_25, plain, (-45/89 < skoXC1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 * -89/30 < sqrt(3)),
    inference(resolve, [], [refute_0_23, refute_0_25])).

cnf(refute_0_27, plain,
    (arctan(skoXC1) * (skoXC1 * 89/10 + sqrt(3) * 3) <=
     skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8) |
     skoXC1 <= 86/89), inference(resolve, [], [refute_0_17, refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     arctan(skoXC1) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89),
    inference(split, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (arctan(skoXC1) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89 |
     sqrt(3) <= skoXC1 * -89/30),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (arctan(skoXC1) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89),
    inference(resolve, [], [refute_0_29, refute_0_26])).

cnf(refute_0_31, plain,
    (X_000120 < arctan(X_000119) | ~ lgen(1, X_000120, arctan(X_000119))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_32, plain,
    (X_000119 < 0 |
     ~ lgen(1, X_000120,
            1/35 *
            (15159 * X_000119 ^ 6 + 147455 * X_000119 ^ 4 +
             345345 * X_000119 ^ 2 + 225225) * X_000119 /
            (35 * X_000119 ^ 8 + 1260 * X_000119 ^ 6 +
             6930 * X_000119 ^ 4 + 12012 * X_000119 ^ 2 + 6435)) |
     lgen(1, X_000120, arctan(X_000119))),
    inference(subst, [], [atan_lower_bound_case_43])).

cnf(refute_0_33, plain,
    (X_000119 < 0 | X_000120 < arctan(X_000119) |
     ~ lgen(1, X_000120,
            1/35 *
            (15159 * X_000119 ^ 6 + 147455 * X_000119 ^ 4 +
             345345 * X_000119 ^ 2 + 225225) * X_000119 /
            (35 * X_000119 ^ 8 + 1260 * X_000119 ^ 6 +
             6930 * X_000119 ^ 4 + 12012 * X_000119 ^ 2 + 6435))),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (X_000119 < 0 | X_000120 < arctan(X_000119) |
     X_000119 *
     (6435 +
      X_000119 *
      (X_000119 *
       (9867 +
        X_000119 *
        (X_000119 * (4213 + X_000119 * (X_000119 * 15159/35)))))) /
     (6435 +
      X_000119 *
      (X_000119 *
       (12012 +
        X_000119 *
        (X_000119 *
         (6930 +
          X_000119 *
          (X_000119 * (1260 + X_000119 * (X_000119 * 35)))))))) <=
     X_000120), inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 < 0 |
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) < arctan(skoXC1) |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3)), inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89),
    inference(resolve, [], [refute_0_30, refute_0_35])).

cnf(refute_0_37, plain,
    (583057/336224 <
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) |
     0 <=
     -768/35 +
     skoXC1 *
     (232810767/1250000 +
      skoXC1 *
      (-3947/5 +
       skoXC1 *
       (2095296903/1250000 +
        skoXC1 *
        (-4170 +
         skoXC1 *
         (23048265933/6250000 +
          skoXC1 *
          (-6831 + skoXC1 * (14267974149/6250000 + skoXC1 * -3432))))))) |
     583057/336224 *
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000)))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_38, plain,
    (3 < 1 |
     sqrt(3) <
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) |
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_39, plain,
    (256/35 +
     skoXC1 *
     (-68722339/1250000 +
      skoXC1 *
      (849/5 +
       skoXC1 *
       (-618501051/1250000 +
        skoXC1 *
        (550 +
         skoXC1 *
         (-6803511561/6250000 +
          skoXC1 * (429 + skoXC1 * -4211697633/6250000)))))) <
     (skoXC1 *
      (skoXC1 *
       (126469/2000 +
        skoXC1 *
        (skoXC1 *
         (1138221/2000 +
          skoXC1 *
          (skoXC1 *
           (12520431/10000 + skoXC1 * (skoXC1 * 7750743/10000))))))) +
      sqrt(3) *
      (skoXC1 *
       (4263/200 +
        skoXC1 *
        (-280 +
         skoXC1 *
         (38367/200 +
          skoXC1 *
          (-2520 +
           skoXC1 *
           (422037/1000 +
            skoXC1 *
            (-5544 + skoXC1 * (261261/1000 + skoXC1 * -3432))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     skoXC1 *
     (skoXC1 *
      (126469/2000 +
       skoXC1 *
       (skoXC1 *
        (1138221/2000 +
         skoXC1 *
         (skoXC1 *
          (12520431/10000 + skoXC1 * (skoXC1 * 7750743/10000))))))) +
     sqrt(3) *
     (skoXC1 *
      (4263/200 +
       skoXC1 *
       (-280 +
        skoXC1 *
        (38367/200 +
         skoXC1 *
         (-2520 +
          skoXC1 *
          (422037/1000 +
           skoXC1 *
           (-5544 + skoXC1 * (261261/1000 + skoXC1 * -3432)))))))) <=
     (256/35 +
      skoXC1 *
      (-68722339/1250000 +
       skoXC1 *
       (849/5 +
        skoXC1 *
        (-618501051/1250000 +
         skoXC1 *
         (550 +
          skoXC1 *
          (-6803511561/6250000 +
           skoXC1 * (429 + skoXC1 * -4211697633/6250000))))))) *
     (skoXC1 * 89/10 + sqrt(3) * 3)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_40, plain,
    ((skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <
     (-256/35 +
      skoXC1 *
      (68722339/1250000 +
       skoXC1 *
       (-849/5 +
        skoXC1 *
        (618501051/1250000 +
         skoXC1 *
         (-550 +
          skoXC1 *
          (6803511561/6250000 +
           skoXC1 * (-429 + skoXC1 * 4211697633/6250000))))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) |
     skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) <= 0 |
     -256/35 +
     skoXC1 *
     (68722339/1250000 +
      skoXC1 *
      (-849/5 +
       skoXC1 *
       (618501051/1250000 +
        skoXC1 *
        (-550 +
         skoXC1 *
         (6803511561/6250000 +
          skoXC1 * (-429 + skoXC1 * 4211697633/6250000)))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) *
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_41, plain,
    (X_000106 < arctan(X_000105) | ~ lgen(1, X_000106, arctan(X_000105))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_42, plain,
    (~ lgen(1, X_000106,
            9817477/6250000 -
            1/35 *
            (256 + 5943 * X_000105 ^ 2 + 19250 * X_000105 ^ 4 +
             15015 * X_000105 ^ 6) /
            (X_000105 *
             (35 + 315 * X_000105 ^ 2 + 693 * X_000105 ^ 4 +
              429 * X_000105 ^ 6))) | X_000105 <= 1 |
     lgen(1, X_000106, arctan(X_000105))),
    inference(subst, [], [atan_lower_bound_case_44])).

cnf(refute_0_43, plain,
    (X_000106 < arctan(X_000105) |
     ~ lgen(1, X_000106,
            9817477/6250000 -
            1/35 *
            (256 + 5943 * X_000105 ^ 2 + 19250 * X_000105 ^ 4 +
             15015 * X_000105 ^ 6) /
            (X_000105 *
             (35 + 315 * X_000105 ^ 2 + 693 * X_000105 ^ 4 +
              429 * X_000105 ^ 6))) | X_000105 <= 1),
    inference(resolve, [], [refute_0_42, refute_0_41])).

cnf(refute_0_44, plain,
    (X_000106 < arctan(X_000105) | X_000105 <= 1 |
     9817477/6250000 +
     (256/35 +
      X_000105 *
      (X_000105 *
       (849/5 +
        X_000105 * (X_000105 * (550 + X_000105 * (X_000105 * 429)))))) /
     (X_000105 *
      (35 +
       X_000105 *
       (X_000105 *
        (315 +
         X_000105 * (X_000105 * (693 + X_000105 * (X_000105 * 429))))))) *
     -1 <= X_000106), inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    ((skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) < arctan(skoXC1) |
     9817477/6250000 +
     (256/35 +
      skoXC1 *
      (skoXC1 *
       (849/5 + skoXC1 * (skoXC1 * (550 + skoXC1 * (skoXC1 * 429)))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) *
     -1 <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 1),
    inference(subst, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (9817477/6250000 +
     (256/35 +
      skoXC1 *
      (skoXC1 *
       (849/5 + skoXC1 * (skoXC1 * (550 + skoXC1 * (skoXC1 * 429)))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) *
     -1 <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_30, refute_0_45])).

cnf(refute_0_47, plain,
    ((-256/35 +
      skoXC1 *
      (68722339/1250000 +
       skoXC1 *
       (-849/5 +
        skoXC1 *
        (618501051/1250000 +
         skoXC1 *
         (-550 +
          skoXC1 *
          (6803511561/6250000 +
           skoXC1 * (-429 + skoXC1 * 4211697633/6250000))))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89 | skoXC1 <= 1 |
     skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) = 0),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) != 0 |
     skoXC1 <= 86/89 | skoXC1 <= 1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_49, plain,
    ((-256/35 +
      skoXC1 *
      (68722339/1250000 +
       skoXC1 *
       (-849/5 +
        skoXC1 *
        (618501051/1250000 +
         skoXC1 *
         (-550 +
          skoXC1 *
          (6803511561/6250000 +
           skoXC1 * (-429 + skoXC1 * 4211697633/6250000))))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (86/89 < skoXC1 | skoXC1 <= 1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_51, plain,
    ((-256/35 +
      skoXC1 *
      (68722339/1250000 +
       skoXC1 *
       (-849/5 +
        skoXC1 *
        (618501051/1250000 +
         skoXC1 *
         (-550 +
          skoXC1 *
          (6803511561/6250000 +
           skoXC1 * (-429 + skoXC1 * 4211697633/6250000))))))) /
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain,
    (skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) <= 0 |
     -256/35 +
     skoXC1 *
     (68722339/1250000 +
      skoXC1 *
      (-849/5 +
       skoXC1 *
       (618501051/1250000 +
        skoXC1 *
        (-550 +
         skoXC1 *
         (6803511561/6250000 +
          skoXC1 * (-429 + skoXC1 * 4211697633/6250000)))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) *
     (skoXC1 *
      (35 +
       skoXC1 *
       (skoXC1 *
        (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_51, refute_0_40])).

cnf(refute_0_53, plain,
    (skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) <= 0 |
     (skoXC1 *
      (skoXC1 *
       (126469/2000 +
        skoXC1 *
        (skoXC1 *
         (1138221/2000 +
          skoXC1 *
          (skoXC1 *
           (12520431/10000 + skoXC1 * (skoXC1 * 7750743/10000))))))) +
      sqrt(3) *
      (skoXC1 *
       (4263/200 +
        skoXC1 *
        (-280 +
         skoXC1 *
         (38367/200 +
          skoXC1 *
          (-2520 +
           skoXC1 *
           (422037/1000 +
            skoXC1 *
            (-5544 + skoXC1 * (261261/1000 + skoXC1 * -3432))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <=
     256/35 +
     skoXC1 *
     (-68722339/1250000 +
      skoXC1 *
      (849/5 +
       skoXC1 *
       (-618501051/1250000 +
        skoXC1 *
        (550 +
         skoXC1 *
         (-6803511561/6250000 +
          skoXC1 * (429 + skoXC1 * -4211697633/6250000)))))) |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (0 <
     skoXC1 *
     (35 +
      skoXC1 *
      (skoXC1 *
       (315 + skoXC1 * (skoXC1 * (693 + skoXC1 * (skoXC1 * 429)))))) |
     skoXC1 <= 1), inference(decision, [], [refute_0_24])).

cnf(refute_0_55, plain,
    ((skoXC1 *
      (skoXC1 *
       (126469/2000 +
        skoXC1 *
        (skoXC1 *
         (1138221/2000 +
          skoXC1 *
          (skoXC1 *
           (12520431/10000 + skoXC1 * (skoXC1 * 7750743/10000))))))) +
      sqrt(3) *
      (skoXC1 *
       (4263/200 +
        skoXC1 *
        (-280 +
         skoXC1 *
         (38367/200 +
          skoXC1 *
          (-2520 +
           skoXC1 *
           (422037/1000 +
            skoXC1 *
            (-5544 + skoXC1 * (261261/1000 + skoXC1 * -3432))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <=
     256/35 +
     skoXC1 *
     (-68722339/1250000 +
      skoXC1 *
      (849/5 +
       skoXC1 *
       (-618501051/1250000 +
        skoXC1 *
        (550 +
         skoXC1 *
         (-6803511561/6250000 +
          skoXC1 * (429 + skoXC1 * -4211697633/6250000)))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     skoXC1 *
     (skoXC1 *
      (126469/2000 +
       skoXC1 *
       (skoXC1 *
        (1138221/2000 +
         skoXC1 *
         (skoXC1 *
          (12520431/10000 + skoXC1 * (skoXC1 * 7750743/10000))))))) +
     sqrt(3) *
     (skoXC1 *
      (4263/200 +
       skoXC1 *
       (-280 +
        skoXC1 *
        (38367/200 +
         skoXC1 *
         (-2520 +
          skoXC1 *
          (422037/1000 +
           skoXC1 *
           (-5544 + skoXC1 * (261261/1000 + skoXC1 * -3432)))))))) <=
     (256/35 +
      skoXC1 *
      (-68722339/1250000 +
       skoXC1 *
       (849/5 +
        skoXC1 *
        (-618501051/1250000 +
         skoXC1 *
         (550 +
          skoXC1 *
          (-6803511561/6250000 +
           skoXC1 * (429 + skoXC1 * -4211697633/6250000))))))) *
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_55, refute_0_39])).

cnf(refute_0_57, plain,
    (sqrt(3) *
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) |
     skoXC1 <= 1 | sqrt(3) <= skoXC1 * -89/30),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (sqrt(3) *
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_57, refute_0_26])).

cnf(refute_0_59, plain,
    (0 <=
     -768/35 +
     skoXC1 *
     (232810767/1250000 +
      skoXC1 *
      (-3947/5 +
       skoXC1 *
       (2095296903/1250000 +
        skoXC1 *
        (-4170 +
         skoXC1 *
         (23048265933/6250000 +
          skoXC1 *
          (-6831 + skoXC1 * (14267974149/6250000 + skoXC1 * -3432))))))) |
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     sqrt(3) | skoXC1 <= 1), inference(split, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 *
     (-232810767/1250000 +
      skoXC1 *
      (3947/5 +
       skoXC1 *
       (-2095296903/1250000 +
        skoXC1 *
        (4170 +
         skoXC1 *
         (-23048265933/6250000 +
          skoXC1 *
          (6831 + skoXC1 * (-14267974149/6250000 + skoXC1 * 3432))))))) <=
     -768/35 |
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     sqrt(3) | skoXC1 <= 1), inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (-768/35 <
     skoXC1 *
     (-232810767/1250000 +
      skoXC1 *
      (3947/5 +
       skoXC1 *
       (-2095296903/1250000 +
        skoXC1 *
        (4170 +
         skoXC1 *
         (-23048265933/6250000 +
          skoXC1 *
          (6831 + skoXC1 * (-14267974149/6250000 + skoXC1 * 3432))))))) |
     skoXC1 <= 1), inference(decision, [], [refute_0_24])).

cnf(refute_0_62, plain,
    (skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     sqrt(3) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain,
    (3 < 1 |
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_62, refute_0_38])).

cnf(refute_0_64, plain,
    (skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) /
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     583057/336224 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (0 <=
     -768/35 +
     skoXC1 *
     (232810767/1250000 +
      skoXC1 *
      (-3947/5 +
       skoXC1 *
       (2095296903/1250000 +
        skoXC1 *
        (-4170 +
         skoXC1 *
         (23048265933/6250000 +
          skoXC1 *
          (-6831 + skoXC1 * (14267974149/6250000 + skoXC1 * -3432))))))) |
     583057/336224 *
     (-768/35 +
      skoXC1 *
      (232810767/1250000 +
       skoXC1 *
       (-3947/5 +
        skoXC1 *
        (2095296903/1250000 +
         skoXC1 *
         (-4170 +
          skoXC1 *
          (23048265933/6250000 +
           skoXC1 *
           (-6831 +
            skoXC1 * (14267974149/6250000 + skoXC1 * -3432)))))))) <=
     skoXC1 *
     (11392/175 +
      skoXC1 *
      (-6906719421/12500000 +
       skoXC1 *
       (75561/50 +
        skoXC1 *
        (-62160474789/12500000 +
         skoXC1 *
         (4895 +
          skoXC1 *
          (-683765222679/62500000 +
           skoXC1 * (38181/10 + skoXC1 * -423283233087/62500000))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_64, refute_0_37])).

cnf(refute_0_66, plain,
    (skoXC1 *
     (-232810767/1250000 +
      skoXC1 *
      (3947/5 +
       skoXC1 *
       (-2095296903/1250000 +
        skoXC1 *
        (4170 +
         skoXC1 *
         (-23048265933/6250000 +
          skoXC1 *
          (6831 + skoXC1 * (-14267974149/6250000 + skoXC1 * 3432))))))) <=
     -768/35 |
     skoXC1 *
     (108382920174719/420280000000 +
      skoXC1 *
      (-3350693473529/4104296875 +
       skoXC1 *
       (83791712110353/60040000000 +
        skoXC1 *
        (-296625082798227/131337500000 +
         skoXC1 *
         (450299970013883/300200000000 +
          skoXC1 *
          (-297353432929311/328343750000 +
           skoXC1 *
           (295686863393493/2101400000000 +
            skoXC1 * 539132976920109/656687500000))))))) <=
     13993368/367745 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (-768/35 <
     skoXC1 *
     (-232810767/1250000 +
      skoXC1 *
      (3947/5 +
       skoXC1 *
       (-2095296903/1250000 +
        skoXC1 *
        (4170 +
         skoXC1 *
         (-23048265933/6250000 +
          skoXC1 *
          (6831 + skoXC1 * (-14267974149/6250000 + skoXC1 * 3432))))))) |
     skoXC1 *
     (108382920174719/420280000000 +
      skoXC1 *
      (-3350693473529/4104296875 +
       skoXC1 *
       (83791712110353/60040000000 +
        skoXC1 *
        (-296625082798227/131337500000 +
         skoXC1 *
         (450299970013883/300200000000 +
          skoXC1 *
          (-297353432929311/328343750000 +
           skoXC1 *
           (295686863393493/2101400000000 +
            skoXC1 * 539132976920109/656687500000))))))) <=
     13993368/367745 | skoXC1 <= 1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_68, plain,
    (skoXC1 *
     (108382920174719/420280000000 +
      skoXC1 *
      (-3350693473529/4104296875 +
       skoXC1 *
       (83791712110353/60040000000 +
        skoXC1 *
        (-296625082798227/131337500000 +
         skoXC1 *
         (450299970013883/300200000000 +
          skoXC1 *
          (-297353432929311/328343750000 +
           skoXC1 *
           (295686863393493/2101400000000 +
            skoXC1 * 539132976920109/656687500000))))))) <=
     13993368/367745 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (13993368/367745 <
     skoXC1 *
     (108382920174719/420280000000 +
      skoXC1 *
      (-3350693473529/4104296875 +
       skoXC1 *
       (83791712110353/60040000000 +
        skoXC1 *
        (-296625082798227/131337500000 +
         skoXC1 *
         (450299970013883/300200000000 +
          skoXC1 *
          (-297353432929311/328343750000 +
           skoXC1 *
           (295686863393493/2101400000000 +
            skoXC1 * 539132976920109/656687500000))))))) | skoXC1 <= 1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_70, plain, (skoXC1 <= 1),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain, (0 <= skoXC1 | skoXC1 <= 86/89),
    inference(decision, [], [refute_0_70, refute_0_24])).

cnf(refute_0_72, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89),
    inference(resolve, [], [refute_0_71, refute_0_36])).

cnf(refute_0_73, plain,
    (skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <=
     (skoXC1 * -18067/10000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (skoXC1 * 89/10 + sqrt(3) * 3) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0 | skoXC1 <= 86/89),
    inference(resolve, [], [refute_0_72, refute_0_7])).

cnf(refute_0_74, plain,
    (skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     (skoXC1 *
      (23252229/2000 +
       skoXC1 *
       (skoXC1 *
        (54255201/2500 +
         skoXC1 *
         (skoXC1 *
          (12520431/1000 +
           skoXC1 *
           (skoXC1 *
            (1138221/500 + skoXC1 * (skoXC1 * 126469/2000)))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) |
     skoXC1 <= 86/89), inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     skoXC1 <= 86/89),
    inference(decision, [], [refute_0_70, refute_0_24])).

cnf(refute_0_76, plain,
    ((skoXC1 *
      (23252229/2000 +
       skoXC1 *
       (skoXC1 *
        (54255201/2500 +
         skoXC1 *
         (skoXC1 *
          (12520431/1000 +
           skoXC1 *
           (skoXC1 *
            (1138221/500 + skoXC1 * (skoXC1 * 126469/2000)))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (skoXC1 * 89/10 + sqrt(3) * 3) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) |
     skoXC1 <= 86/89), inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 * 89/10 + sqrt(3) * 3 <= 0 |
     skoXC1 *
     (23252229/2000 +
      skoXC1 *
      (skoXC1 *
       (54255201/2500 +
        skoXC1 *
        (skoXC1 *
         (12520431/1000 +
          skoXC1 *
          (skoXC1 * (1138221/500 + skoXC1 * (skoXC1 * 126469/2000)))))))) +
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-51480 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-96096 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-55440 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-10080 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) *
     (skoXC1 * 89/10 + sqrt(3) * 3) | skoXC1 <= 86/89),
    inference(resolve, [], [refute_0_76, refute_0_6])).

cnf(refute_0_78, plain,
    (sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) |
     skoXC1 <= 86/89 | sqrt(3) <= skoXC1 * -89/30),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) <
     (11234223/1000 +
      skoXC1 *
      (skoXC1 *
       (26213187/1250 +
        skoXC1 *
        (skoXC1 *
         (6049197/500 +
          skoXC1 *
          (skoXC1 * (549927/250 + skoXC1 * (skoXC1 * 61103/1000))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (43/5 + sqrt(3) * 3) | 43/5 + sqrt(3) * 3 <= 0 |
     11234223/1000 +
     skoXC1 *
     (skoXC1 *
      (26213187/1250 +
       skoXC1 *
       (skoXC1 *
        (6049197/500 +
         skoXC1 *
         (skoXC1 * (549927/250 + skoXC1 * (skoXC1 * 61103/1000))))))) +
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-51480 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-96096 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-55440 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-10080 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) *
     (43/5 + sqrt(3) * 3)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_80, plain,
    ((-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3) <
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_81, plain,
    (43/5 < skoXC1 * 89/10 | max(43/5, skoXC1 * 89/10) = 43/5),
    inference(subst, [], [max_1])).

cnf(refute_0_82, plain,
    (sqrt(3) * (skoXC1 * 8) / (max(43/5, skoXC1 * 89/10) + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) | max(43/5, skoXC1 * 89/10) != 43/5 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (43/5 + sqrt(3) * 3)),
    introduced(tautology, [equality])).

cnf(refute_0_83, plain,
    (43/5 < skoXC1 * 89/10 |
     sqrt(3) * (skoXC1 * 8) / (max(43/5, skoXC1 * 89/10) + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (43/5 < skoXC1 * 89/10 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_9, refute_0_83])).

cnf(refute_0_85, plain,
    (86/89 < skoXC1 |
     203/1000 + arctan(skoXC1) <=
     sqrt(3) * (skoXC1 * 8) / (43/5 + sqrt(3) * 3)),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (sqrt(3) * (skoXC1 * 8) / (43/5 + sqrt(3) * 3) <
     203/1000 + arctan(skoXC1) |
     (203/1000 + arctan(skoXC1)) * (43/5 + sqrt(3) * 3) <=
     sqrt(3) * (skoXC1 * 8) | 43/5 + sqrt(3) * 3 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_87, plain,
    (86/89 < skoXC1 |
     (203/1000 + arctan(skoXC1)) * (43/5 + sqrt(3) * 3) <=
     sqrt(3) * (skoXC1 * 8) | 43/5 + sqrt(3) * 3 <= 0),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (86/89 < skoXC1 |
     arctan(skoXC1) * (43/5 + sqrt(3) * 3) <=
     -8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8) | sqrt(3) <= -43/15),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (3 < 0 | -43/15 < sqrt(3) | 3 * 2 / (1 + 3) <= -43/15),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_90, plain, (-43/15 < sqrt(3)),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (86/89 < skoXC1 |
     arctan(skoXC1) * (43/5 + sqrt(3) * 3) <=
     -8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_88, refute_0_90])).

cnf(refute_0_92, plain,
    (86/89 < skoXC1 | 43/5 + sqrt(3) * 3 <= 0 |
     arctan(skoXC1) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3)), inference(split, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (86/89 < skoXC1 |
     arctan(skoXC1) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3) | sqrt(3) <= -43/15),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (86/89 < skoXC1 |
     arctan(skoXC1) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_93, refute_0_90])).

cnf(refute_0_95, plain,
    (skoXC1 < 0 |
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3) < arctan(skoXC1) |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3)), inference(subst, [], [refute_0_34])).

cnf(refute_0_96, plain,
    (skoXC1 < 0 | 86/89 < skoXC1 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_94, refute_0_95])).

cnf(refute_0_97, plain, (86/89 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_70, refute_0_24])).

cnf(refute_0_98, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) /
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 +
          skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_97, refute_0_96])).

cnf(refute_0_99, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (9867 +
        skoXC1 * (skoXC1 * (4213 + skoXC1 * (skoXC1 * 15159/35)))))) <=
     (-8729/5000 + sqrt(3) * (-609/1000 + skoXC1 * 8)) /
     (43/5 + sqrt(3) * 3) *
     (6435 +
      skoXC1 *
      (skoXC1 *
       (12012 +
        skoXC1 *
        (skoXC1 *
         (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35)))))))) |
     6435 +
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     0), inference(resolve, [], [refute_0_98, refute_0_80])).

cnf(refute_0_100, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) <=
     -6435 |
     (11234223/1000 +
      skoXC1 *
      (skoXC1 *
       (26213187/1250 +
        skoXC1 *
        (skoXC1 *
         (6049197/500 +
          skoXC1 *
          (skoXC1 * (549927/250 + skoXC1 * (skoXC1 * 61103/1000))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (43/5 + sqrt(3) * 3) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35))))))),
    inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (-6435 <
     skoXC1 *
     (skoXC1 *
      (12012 +
       skoXC1 *
       (skoXC1 *
        (6930 + skoXC1 * (skoXC1 * (1260 + skoXC1 * (skoXC1 * 35))))))) |
     86/89 < skoXC1), inference(decision, [], [refute_0_70, refute_0_24])).

cnf(refute_0_102, plain,
    (86/89 < skoXC1 |
     (11234223/1000 +
      skoXC1 *
      (skoXC1 *
       (26213187/1250 +
        skoXC1 *
        (skoXC1 *
         (6049197/500 +
          skoXC1 *
          (skoXC1 * (549927/250 + skoXC1 * (skoXC1 * 61103/1000))))))) +
      sqrt(3) *
      (783783/200 +
       skoXC1 *
       (-51480 +
        skoXC1 *
        (1828827/250 +
         skoXC1 *
         (-96096 +
          skoXC1 *
          (422037/100 +
           skoXC1 *
           (-55440 +
            skoXC1 *
            (38367/50 +
             skoXC1 *
             (-10080 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))) /
     (43/5 + sqrt(3) * 3) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35))))))),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (86/89 < skoXC1 | 43/5 + sqrt(3) * 3 <= 0 |
     11234223/1000 +
     skoXC1 *
     (skoXC1 *
      (26213187/1250 +
       skoXC1 *
       (skoXC1 *
        (6049197/500 +
         skoXC1 *
         (skoXC1 * (549927/250 + skoXC1 * (skoXC1 * 61103/1000))))))) +
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-51480 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-96096 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-55440 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-10080 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-6435 +
      skoXC1 *
      (skoXC1 *
       (-9867 +
        skoXC1 * (skoXC1 * (-4213 + skoXC1 * (skoXC1 * -15159/35)))))) *
     (43/5 + sqrt(3) * 3)),
    inference(resolve, [], [refute_0_102, refute_0_79])).

cnf(refute_0_104, plain,
    (86/89 < skoXC1 |
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000))))))) |
     sqrt(3) <= -43/15), inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (86/89 < skoXC1 |
     sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))),
    inference(resolve, [], [refute_0_104, refute_0_90])).

cnf(refute_0_106, plain,
    (783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) < 0 |
     0 <
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     86/89 < skoXC1 |
     0 <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))),
    inference(split, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) |
     86/89 < skoXC1 |
     skoXC1 *
     (55341 +
      skoXC1 *
      (26213187/1250 +
       skoXC1 *
       (424281/5 +
        skoXC1 *
        (6049197/500 +
         skoXC1 *
         (181159/5 +
          skoXC1 *
          (549927/250 +
           skoXC1 * (651837/175 + skoXC1 * 61103/1000))))))) <=
     -11234223/1000), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) |
     -11234223/1000 <
     skoXC1 *
     (55341 +
      skoXC1 *
      (26213187/1250 +
       skoXC1 *
       (424281/5 +
        skoXC1 *
        (6049197/500 +
         skoXC1 *
         (181159/5 +
          skoXC1 *
          (549927/250 + skoXC1 * (651837/175 + skoXC1 * 61103/1000))))))) |
     86/89 < skoXC1), inference(decision, [], [refute_0_24, refute_0_70])).

cnf(refute_0_109, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) |
     86/89 < skoXC1),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) |
     skoXC1 <= 86/89),
    inference(decision, [], [refute_0_24, refute_0_70])).

cnf(refute_0_111, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280))))))))),
    inference(resolve, [], [refute_0_110, refute_0_109])).

cnf(refute_0_112, plain,
    (583057/336224 <
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) |
     0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     583057/336224 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_113, plain,
    (3 < 1 |
     sqrt(3) <
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_114, plain,
    (86/89 < skoXC1 |
     0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     sqrt(3)), inference(split, [], [refute_0_105])).

cnf(refute_0_115, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200 |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     sqrt(3)), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (3 < 1 | 86/89 < skoXC1 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200 |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3)))),
    inference(resolve, [], [refute_0_115, refute_0_113])).

cnf(refute_0_117, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200 |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     583057/336224), inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (86/89 < skoXC1 |
     0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     583057/336224 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000))))))) |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200), inference(resolve, [], [refute_0_117, refute_0_112])).

cnf(refute_0_119, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (-152886591/336224 +
      skoXC1 *
      (2020722072369/60040000 +
       skoXC1 *
       (-51198421131/1681120 +
        skoXC1 *
        (466320478239/24016000 +
         skoXC1 *
         (-63867109669/1681120 +
          skoXC1 *
          (42392770749/12008000 +
           skoXC1 *
           (-676770888567/58839200 +
            skoXC1 *
            (4710307861/48032000 + skoXC1 * -2915285/6004)))))))) <=
     -866023745301/48032000 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200), inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <
     1008672/583057 |
     1008672/583057 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000))))))) |
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_121, plain,
    (X_000078 < sqrt(X_000077) | ~ lgen(1, X_000078, sqrt(X_000077))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_122, plain,
    (X_000077 < 1 |
     ~ lgen(1, X_000078,
            32 * X_000077 * (X_000077 + 16) *
            (X_000077 ^ 2 + 96 * X_000077 + 256) /
            (X_000077 ^ 4 + 448 * X_000077 ^ 3 + 17920 * X_000077 ^ 2 +
             114688 * X_000077 + 65536)) |
     lgen(1, X_000078, sqrt(X_000077))),
    inference(subst, [], [sqrt_lower_bound_2_large])).

cnf(refute_0_123, plain,
    (X_000077 < 1 | X_000078 < sqrt(X_000077) |
     ~ lgen(1, X_000078,
            32 * X_000077 * (X_000077 + 16) *
            (X_000077 ^ 2 + 96 * X_000077 + 256) /
            (X_000077 ^ 4 + 448 * X_000077 ^ 3 + 17920 * X_000077 ^ 2 +
             114688 * X_000077 + 65536))),
    inference(resolve, [], [refute_0_122, refute_0_121])).

cnf(refute_0_124, plain,
    (X_000077 < 1 | X_000078 < sqrt(X_000077) |
     X_000077 *
     (131072 + X_000077 * (57344 + X_000077 * (3584 + X_000077 * 32))) /
     (65536 +
      X_000077 *
      (114688 + X_000077 * (17920 + X_000077 * (448 + X_000077)))) <=
     X_000078), inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (3 < 1 |
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <
     sqrt(3) |
     3 * (131072 + 3 * (57344 + 3 * (3584 + 3 * 32))) /
     (65536 + 3 * (114688 + 3 * (17920 + 3 * (448 + 3)))) <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))),
    inference(subst, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (86/89 < skoXC1 |
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <= 0 |
     sqrt(3) <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))),
    inference(split, [], [refute_0_105])).

cnf(refute_0_127, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 |
     sqrt(3) <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    (skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 | skoXC1 <= 86/89),
    inference(decision, [], [refute_0_24, refute_0_70])).

cnf(refute_0_129, plain,
    (skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 |
     sqrt(3) <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))),
    inference(resolve, [], [refute_0_128, refute_0_127])).

cnf(refute_0_130, plain,
    (3 < 1 |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 |
     3 * (131072 + 3 * (57344 + 3 * (3584 + 3 * 32))) /
     (65536 + 3 * (114688 + 3 * (17920 + 3 * (448 + 3)))) <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))))),
    inference(resolve, [], [refute_0_129, refute_0_125])).

cnf(refute_0_131, plain,
    (1008672/583057 <=
     (-11234223/1000 +
      skoXC1 *
      (-55341 +
       skoXC1 *
       (-26213187/1250 +
        skoXC1 *
        (-424281/5 +
         skoXC1 *
         (-6049197/500 +
          skoXC1 *
          (-181159/5 +
           skoXC1 *
           (-549927/250 +
            skoXC1 * (-651837/175 + skoXC1 * -61103/1000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (1008672/583057 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     -11234223/1000 +
     skoXC1 *
     (-55341 +
      skoXC1 *
      (-26213187/1250 +
       skoXC1 *
       (-424281/5 +
        skoXC1 *
        (-6049197/500 +
         skoXC1 *
         (-181159/5 +
          skoXC1 *
          (-549927/250 +
           skoXC1 * (-651837/175 + skoXC1 * -61103/1000))))))) |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 |
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <= 0),
    inference(resolve, [], [refute_0_131, refute_0_120])).

cnf(refute_0_133, plain,
    (skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200 |
     skoXC1 *
     (-187064163/583057 +
      skoXC1 *
      (24507215111379/728821250 +
       skoXC1 *
       (-87978216183/2915285 +
        skoXC1 *
        (5655511179549/291528500 +
         skoXC1 *
         (-110234828297/2915285 +
          skoXC1 *
          (514137379959/145764250 +
           skoXC1 *
           (-1169882399571/102034975 +
            skoXC1 *
            (57126375551/583057000 + skoXC1 * -282428160/583057)))))))) <=
     -10503092190591/583057000),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (-10503092190591/583057000 <
     skoXC1 *
     (-187064163/583057 +
      skoXC1 *
      (24507215111379/728821250 +
       skoXC1 *
       (-87978216183/2915285 +
        skoXC1 *
        (5655511179549/291528500 +
         skoXC1 *
         (-110234828297/2915285 +
          skoXC1 *
          (514137379959/145764250 +
           skoXC1 *
           (-1169882399571/102034975 +
            skoXC1 *
            (57126375551/583057000 + skoXC1 * -282428160/583057)))))))) |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200), inference(decision, [], [refute_0_24, refute_0_70])).

cnf(refute_0_135, plain,
    (skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200), inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain,
    (-866023745301/48032000 <
     skoXC1 *
     (-152886591/336224 +
      skoXC1 *
      (2020722072369/60040000 +
       skoXC1 *
       (-51198421131/1681120 +
        skoXC1 *
        (466320478239/24016000 +
         skoXC1 *
         (-63867109669/1681120 +
          skoXC1 *
          (42392770749/12008000 +
           skoXC1 *
           (-676770888567/58839200 +
            skoXC1 *
            (4710307861/48032000 + skoXC1 * -2915285/6004)))))))) |
     86/89 < skoXC1 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200),
    inference(decision, [], [refute_0_135, refute_0_24, refute_0_70])).

cnf(refute_0_137, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200), inference(resolve, [], [refute_0_119, refute_0_136])).

cnf(refute_0_138, plain,
    (-783783/200 <
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     86/89 < skoXC1),
    inference(resolve, [], [refute_0_137, refute_0_111])).

cnf(refute_0_139, plain,
    (86/89 < skoXC1 |
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) <=
     -783783/200),
    inference(decision, [],
              [refute_0_137, refute_0_135, refute_0_24, refute_0_70])).

cnf(refute_0_140, plain, (86/89 < skoXC1),
    inference(resolve, [], [refute_0_139, refute_0_138])).

cnf(refute_0_141, plain,
    (sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) |
     sqrt(3) <= skoXC1 * -89/30),
    inference(resolve, [], [refute_0_78, refute_0_140])).

cnf(refute_0_142, plain,
    (sqrt(3) *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000))))))))),
    inference(resolve, [], [refute_0_141, refute_0_26])).

cnf(refute_0_143, plain,
    (0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     sqrt(3)), inference(split, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200 |
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     sqrt(3)), inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (783783/200 <
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280))))))))),
    inference(decision, [],
              [refute_0_70, refute_0_24, refute_0_135, refute_0_137,
               refute_0_140])).

cnf(refute_0_146, plain,
    (skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     sqrt(3)), inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (3 < 1 |
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     (2048 + 3 * (3584 + 3 * (560 + 3 * (14 + 3 * 1/32)))) /
     (4096 + 3 * (1792 + 3 * (112 + 3)))),
    inference(resolve, [], [refute_0_146, refute_0_5])).

cnf(refute_0_148, plain,
    (skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000)))))))) /
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     583057/336224), inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain,
    (0 <=
     783783/200 +
     skoXC1 *
     (-32175 +
      skoXC1 *
      (1828827/250 +
       skoXC1 *
       (-66495 +
        skoXC1 *
        (422037/100 +
         skoXC1 *
         (-42801 +
          skoXC1 *
          (38367/50 +
           skoXC1 *
           (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280)))))))) |
     583057/336224 *
     (783783/200 +
      skoXC1 *
      (-32175 +
       skoXC1 *
       (1828827/250 +
        skoXC1 *
        (-66495 +
         skoXC1 *
         (422037/100 +
          skoXC1 *
          (-42801 +
           skoXC1 *
           (38367/50 +
            skoXC1 *
            (-307323/35 + skoXC1 * (4263/200 + skoXC1 * -280))))))))) <=
     skoXC1 *
     (-23252229/2000 +
      skoXC1 *
      (-114543/2 +
       skoXC1 *
       (-54255201/2500 +
        skoXC1 *
        (-878163/10 +
         skoXC1 *
         (-12520431/1000 +
          skoXC1 *
          (-374957/10 +
           skoXC1 *
           (-1138221/500 +
            skoXC1 * (-1349151/350 + skoXC1 * -126469/2000))))))))),
    inference(resolve, [], [refute_0_148, refute_0_0])).

cnf(refute_0_150, plain,
    (skoXC1 *
     (-1856360031669/42028000 +
      skoXC1 *
      (840046226877/12008000 +
       skoXC1 *
       (-19671009334119/210140000 +
        skoXC1 *
        (456952341747/4803200 +
         skoXC1 *
         (-2593219158057/42028000 +
          skoXC1 *
          (93245408537/2401600 +
           skoXC1 *
           (-3809950529643/294196000 +
            skoXC1 *
            (261695255319/67244800 +
             skoXC1 * -1267812531/3002000)))))))) <= -65284309233/9606400 |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200), inference(arithmetic, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (-65284309233/9606400 <
     skoXC1 *
     (-1856360031669/42028000 +
      skoXC1 *
      (840046226877/12008000 +
       skoXC1 *
       (-19671009334119/210140000 +
        skoXC1 *
        (456952341747/4803200 +
         skoXC1 *
         (-2593219158057/42028000 +
          skoXC1 *
          (93245408537/2401600 +
           skoXC1 *
           (-3809950529643/294196000 +
            skoXC1 *
            (261695255319/67244800 + skoXC1 * -1267812531/3002000)))))))) |
     skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200),
    inference(decision, [],
              [refute_0_70, refute_0_24, refute_0_135, refute_0_137,
               refute_0_140])).

cnf(refute_0_152, plain,
    (skoXC1 *
     (32175 +
      skoXC1 *
      (-1828827/250 +
       skoXC1 *
       (66495 +
        skoXC1 *
        (-422037/100 +
         skoXC1 *
         (42801 +
          skoXC1 *
          (-38367/50 +
           skoXC1 *
           (307323/35 + skoXC1 * (-4263/200 + skoXC1 * 280)))))))) <=
     783783/200), inference(resolve, [], [refute_0_150, refute_0_151])).

cnf(refute_0_153, plain, ($false),
    inference(resolve, [], [refute_0_152, refute_0_145])).
% SZS output end CNFRefutation for atan-problem-1-sqrt-max.tptp
