% SZS output start CNFRefutation for Chua-1-IL-U-sincos.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(sin_lower_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, Y, X) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof('Chua', conjecture,
    (! [X] :
       ((0 <= X & X <= 289) =>
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * cos(19/1000 * X) -
        21/50 * exp(3/12500 * X) * sin(19/1000 * X) < 5))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 289) =>
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * cos(19/1000 * X) -
        21/50 * exp(3/12500 * X) * sin(19/1000 * X) < 5)),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 289) =>
          151/100 - 23/1000 * exp(-19/1000 * X) -
          47/20 * exp(3/12500 * X) * cos(19/1000 * X) -
          21/50 * exp(3/12500 * X) * sin(19/1000 * X) < 5)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 289 & 0 <= X &
        5 <=
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * cos(19/1000 * X) -
        21/50 * exp(3/12500 * X) * sin(19/1000 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * cos(19/1000 * skoXC1) -
     21/50 * exp(3/12500 * skoXC1) * sin(19/1000 * skoXC1) &
     skoXC1 <= 289), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * cos(19/1000 * skoXC1) -
     21/50 * exp(3/12500 * skoXC1) * sin(19/1000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 289),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     -3174517702656/25 +
     skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0 |
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     exp(skoXC1 * -19/1000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 <
     (255483445248/5 +
      skoXC1 *
      (542222843904/3125 +
       skoXC1 *
       (-720543154176/78125 +
        skoXC1 *
        (-509745954816/48828125 +
         skoXC1 *
         (338692810752/1220703125 +
          skoXC1 *
          (718821444096/3814697265625 +
           skoXC1 *
           (-318406522608/95367431640625 +
            skoXC1 *
            (-96538147812/59604644775390625 +
             skoXC1 *
             (7184047166343/333786010742187500000 +
              skoXC1 *
              (968063093337/119209289550781250000000 +
               skoXC1 *
               (-288160114116647/3337860107421875000000000000 +
                skoXC1 *
                -349470776694657/
                13113021850585937500000000000000))))))))))) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     255483445248/5 +
     skoXC1 *
     (542222843904/3125 +
      skoXC1 *
      (-720543154176/78125 +
       skoXC1 *
       (-509745954816/48828125 +
        skoXC1 *
        (338692810752/1220703125 +
         skoXC1 *
         (718821444096/3814697265625 +
          skoXC1 *
          (-318406522608/95367431640625 +
           skoXC1 *
           (-96538147812/59604644775390625 +
            skoXC1 *
            (7184047166343/333786010742187500000 +
             skoXC1 *
             (968063093337/119209289550781250000000 +
              skoXC1 *
              (-288160114116647/3337860107421875000000000000 +
               skoXC1 *
               -349470776694657/
               13113021850585937500000000000000)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_6, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     0 <=
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_7, plain,
    (exp(X_000129) < X_000130 | ~ lgen(1, exp(X_000129), X_000130)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_8, plain,
    (3 < X_000129 |
     ~ lgen(1,
            21743271936 /
            (-(X_000129 ^ 3) + 12 * X_000129 ^ 2 - 96 * X_000129 + 384) ^
            4, X_000130) | lgen(1, exp(X_000129), X_000130)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_9, plain,
    (3 < X_000129 | exp(X_000129) < X_000130 |
     ~ lgen(1,
            21743271936 /
            (-(X_000129 ^ 3) + 12 * X_000129 ^ 2 - 96 * X_000129 + 384) ^
            4, X_000130)),
    inference(resolve, [], [refute_0_8, refute_0_7])).

cnf(refute_0_10, plain,
    (3 < X_000129 | exp(X_000129) < X_000130 |
     X_000130 <=
     21743271936 /
     (21743271936 +
      X_000129 *
      (-21743271936 +
       X_000129 *
       (10871635968 +
        X_000129 *
        (-3623878656 +
         X_000129 *
         (891813888 +
          X_000129 *
          (-169869312 +
           X_000129 *
           (25657344 +
            X_000129 *
            (-3096576 +
             X_000129 *
             (297216 +
              X_000129 *
              (-22272 +
               X_000129 * (1248 + X_000129 * (-48 + X_000129))))))))))))),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (3 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_12, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-361/2000000 +
       skoXC1 *
       (skoXC1 *
        (130321/24000000000000 +
         skoXC1 *
         (skoXC1 *
          (-47045881/720000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (16983563041/40320000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -6131066257801/
              3628800000000000000000000000000000000))))))))) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-361/2000000 +
        skoXC1 *
        (skoXC1 *
         (130321/24000000000000 +
          skoXC1 *
          (skoXC1 *
           (-47045881/720000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (16983563041/40320000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -6131066257801/
               3628800000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 *
      (-399/50000 +
       skoXC1 *
       (skoXC1 *
        (48013/100000000000 +
         skoXC1 *
         (skoXC1 *
          (-17332693/2000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (893871739/12000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_13, plain,
    (X_000035 < cos(X_000034) | ~ lgen(1, X_000035, cos(X_000034))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (~ lgen(1, X_000035,
            1 - X_000034 ^ 2 / 2 + X_000034 ^ 4 / 24 - X_000034 ^ 6 / 720 +
            X_000034 ^ 8 / 40320 - X_000034 ^ 10 / 3628800) |
     lgen(1, X_000035, cos(X_000034))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_15, plain,
    (X_000035 < cos(X_000034) |
     ~ lgen(1, X_000035,
            1 - X_000034 ^ 2 / 2 + X_000034 ^ 4 / 24 - X_000034 ^ 6 / 720 +
            X_000034 ^ 8 / 40320 - X_000034 ^ 10 / 3628800)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (X_000035 < cos(X_000034) |
     1 +
     X_000034 *
     (X_000034 *
      (-1/2 +
       X_000034 *
       (X_000034 *
        (1/24 +
         X_000034 *
         (X_000034 *
          (-1/720 +
           X_000034 *
           (X_000034 *
            (1/40320 + X_000034 * (X_000034 * -1/3628800))))))))) <=
     X_000035), inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20) < cos(skoXC1 * 19/1000) |
     1 +
     skoXC1 * 19/1000 *
     (skoXC1 * 19/1000 *
      (-1/2 +
       skoXC1 * 19/1000 *
       (skoXC1 * 19/1000 *
        (1/24 +
         skoXC1 * 19/1000 *
         (skoXC1 * 19/1000 *
          (-1/720 +
           skoXC1 * 19/1000 *
           (skoXC1 * 19/1000 *
            (1/40320 +
             skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/3628800))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) <
     skoXC1 *
     (19/1000 +
      skoXC1 *
      (skoXC1 *
       (-6859/6000000000 +
        skoXC1 *
        (skoXC1 *
         (2476099/120000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-893871739/5040000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (322687697779/362880000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -116490258898219/
               39916800000000000000000000000000000000000)))))))))) |
     skoXC1 *
     (19/1000 +
      skoXC1 *
      (skoXC1 *
       (-6859/6000000000 +
        skoXC1 *
        (skoXC1 *
         (2476099/120000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-893871739/5040000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (322687697779/362880000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -116490258898219/
               39916800000000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_19, plain,
    (5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * cos(19/1000 * skoXC1) -
     21/50 * exp(3/12500 * skoXC1) * sin(19/1000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_20, plain,
    (sin(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (exp(skoXC1 * 3/12500) * 21/50 <= 0 |
     sin(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(split, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (sin(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (X_000185 < sin(X_000184) | ~ lgen(1, X_000185, sin(X_000184))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_24, plain,
    (X_000184 < 0 |
     ~ lgen(1, X_000185,
            X_000184 - X_000184 ^ 3 / 6 + X_000184 ^ 5 / 120 -
            X_000184 ^ 7 / 5040 + X_000184 ^ 9 / 362880 -
            X_000184 ^ 11 / 39916800) | lgen(1, X_000185, sin(X_000184))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_25, plain,
    (X_000184 < 0 | X_000185 < sin(X_000184) |
     ~ lgen(1, X_000185,
            X_000184 - X_000184 ^ 3 / 6 + X_000184 ^ 5 / 120 -
            X_000184 ^ 7 / 5040 + X_000184 ^ 9 / 362880 -
            X_000184 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (X_000184 < 0 | X_000185 < sin(X_000184) |
     X_000184 *
     (1 +
      X_000184 *
      (X_000184 *
       (-1/6 +
        X_000184 *
        (X_000184 *
         (1/120 +
          X_000184 *
          (X_000184 *
           (-1/5040 +
            X_000184 *
            (X_000184 *
             (1/362880 + X_000184 * (X_000184 * -1/39916800)))))))))) <=
     X_000185), inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (skoXC1 * 19/1000 < 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) < sin(skoXC1 * 19/1000) |
     skoXC1 * 19/1000 *
     (1 +
      skoXC1 * 19/1000 *
      (skoXC1 * 19/1000 *
       (-1/6 +
        skoXC1 * 19/1000 *
        (skoXC1 * 19/1000 *
         (1/120 +
          skoXC1 * 19/1000 *
          (skoXC1 * 19/1000 *
           (-1/5040 +
            skoXC1 * 19/1000 *
            (skoXC1 * 19/1000 *
             (1/362880 +
              skoXC1 * 19/1000 *
              (skoXC1 * 19/1000 * -1/39916800)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 * 19/1000 < 0 |
     skoXC1 * 19/1000 *
     (1 +
      skoXC1 * 19/1000 *
      (skoXC1 * 19/1000 *
       (-1/6 +
        skoXC1 * 19/1000 *
        (skoXC1 * 19/1000 *
         (1/120 +
          skoXC1 * 19/1000 *
          (skoXC1 * 19/1000 *
           (-1/5040 +
            skoXC1 * 19/1000 *
            (skoXC1 * 19/1000 *
             (1/362880 +
              skoXC1 * 19/1000 *
              (skoXC1 * 19/1000 * -1/39916800)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(resolve, [], [refute_0_22, refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (19/1000 +
      skoXC1 *
      (skoXC1 *
       (-6859/6000000000 +
        skoXC1 *
        (skoXC1 *
         (2476099/120000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-893871739/5040000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (322687697779/362880000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -116490258898219/
               39916800000000000000000000000000000000000)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (19/1000 +
      skoXC1 *
      (skoXC1 *
       (-6859/6000000000 +
        skoXC1 *
        (skoXC1 *
         (2476099/120000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-893871739/5040000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (322687697779/362880000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -116490258898219/
               39916800000000000000000000000000000000000)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(resolve, [], [refute_0_30, refute_0_29])).

cnf(refute_0_32, plain,
    (skoXC1 *
     (19/1000 +
      skoXC1 *
      (skoXC1 *
       (-6859/6000000000 +
        skoXC1 *
        (skoXC1 *
         (2476099/120000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-893871739/5040000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (322687697779/362880000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -116490258898219/
               39916800000000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(resolve, [], [refute_0_31, refute_0_18])).

cnf(refute_0_33, plain,
    (cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 *
      (-399/50000 +
       skoXC1 *
       (skoXC1 *
        (48013/100000000000 +
         skoXC1 *
         (skoXC1 *
          (-17332693/2000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (893871739/12000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                116490258898219/
                95040000000000000000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (exp(skoXC1 * 3/12500) * 47/20 <= 0 |
     cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(split, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (1 +
     skoXC1 * 19/1000 *
     (skoXC1 * 19/1000 *
      (-1/2 +
       skoXC1 * 19/1000 *
       (skoXC1 * 19/1000 *
        (1/24 +
         skoXC1 * 19/1000 *
         (skoXC1 * 19/1000 *
          (-1/720 +
           skoXC1 * 19/1000 *
           (skoXC1 * 19/1000 *
            (1/40320 +
             skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/3628800))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(resolve, [], [refute_0_35, refute_0_17])).

cnf(refute_0_37, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-361/2000000 +
       skoXC1 *
       (skoXC1 *
        (130321/24000000000000 +
         skoXC1 *
         (skoXC1 *
          (-47045881/720000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (16983563041/40320000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -6131066257801/
              3628800000000000000000000000000000000))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 *
       (-399/50000 +
        skoXC1 *
        (skoXC1 *
         (48013/100000000000 +
          skoXC1 *
          (skoXC1 *
           (-17332693/2000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (893871739/12000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (-322687697779/864000000000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 116490258898219/
                 95040000000000000000000000000000000000000)))))))))))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-361/2000000 +
        skoXC1 *
        (skoXC1 *
         (130321/24000000000000 +
          skoXC1 *
          (skoXC1 *
           (-47045881/720000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (16983563041/40320000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -6131066257801/
               3628800000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 *
      (-399/50000 +
       skoXC1 *
       (skoXC1 *
        (48013/100000000000 +
         skoXC1 *
         (skoXC1 *
          (-17332693/2000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (893871739/12000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(resolve, [], [refute_0_37, refute_0_12])).

cnf(refute_0_39, plain,
    (exp(skoXC1 * 3/12500) *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (0 <=
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     exp(skoXC1 * 3/12500)), inference(split, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     exp(skoXC1 * 3/12500)), inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (3 < skoXC1 * 3/12500 |
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(resolve, [], [refute_0_41, refute_0_11])).

cnf(refute_0_43, plain,
    (12500 < skoXC1 |
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain, (0 < -3490/23 | -3490/23 < exp(skoXC1 * -19/1000)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_45, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-361/2000000 +
       skoXC1 *
       (skoXC1 *
        (130321/24000000000000 +
         skoXC1 *
         (skoXC1 *
          (-47045881/720000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (16983563041/40320000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -6131066257801/
              3628800000000000000000000000000000000))))))))) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-361/2000000 +
        skoXC1 *
        (skoXC1 *
         (130321/24000000000000 +
          skoXC1 *
          (skoXC1 *
           (-47045881/720000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (16983563041/40320000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -6131066257801/
               3628800000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_46, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20) < cos(skoXC1 * 19/1000) |
     1 +
     skoXC1 * 19/1000 *
     (skoXC1 * 19/1000 *
      (-1/2 +
       skoXC1 * 19/1000 *
       (skoXC1 * 19/1000 *
        (1/24 +
         skoXC1 * 19/1000 *
         (skoXC1 * 19/1000 *
          (-1/720 +
           skoXC1 * 19/1000 *
           (skoXC1 * 19/1000 *
            (1/40320 +
             skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/3628800))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_47, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) <
     skoXC1 * (19/1000 + skoXC1 * (skoXC1 * -6859/6000000000)) |
     skoXC1 * (19/1000 + skoXC1 * (skoXC1 * -6859/6000000000)) *
     (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_48, plain,
    (X_000116 < sin(X_000115) | ~ lgen(1, X_000116, sin(X_000115))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_49, plain,
    (X_000115 < 0 | ~ lgen(1, X_000116, X_000115 - X_000115 ^ 3 / 6) |
     lgen(1, X_000116, sin(X_000115))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_50, plain,
    (X_000115 < 0 | X_000116 < sin(X_000115) |
     ~ lgen(1, X_000116, X_000115 - X_000115 ^ 3 / 6)),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    (X_000115 < 0 | X_000116 < sin(X_000115) |
     X_000115 * (1 + X_000115 * (X_000115 * -1/6)) <= X_000116),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (skoXC1 * 19/1000 < 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) < sin(skoXC1 * 19/1000) |
     skoXC1 * 19/1000 *
     (1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/6)) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(subst, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (skoXC1 * 19/1000 < 0 |
     skoXC1 * 19/1000 *
     (1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/6)) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(resolve, [], [refute_0_22, refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < 0 |
     skoXC1 * (19/1000 + skoXC1 * (skoXC1 * -6859/6000000000)) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 * (19/1000 + skoXC1 * (skoXC1 * -6859/6000000000)) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(resolve, [], [refute_0_30, refute_0_54])).

cnf(refute_0_56, plain,
    (skoXC1 * (19/1000 + skoXC1 * (skoXC1 * -6859/6000000000)) *
     (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(resolve, [], [refute_0_55, refute_0_47])).

cnf(refute_0_57, plain,
    (cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (exp(skoXC1 * 3/12500) * 47/20 <= 0 |
     cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(split, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (1 +
     skoXC1 * 19/1000 *
     (skoXC1 * 19/1000 *
      (-1/2 +
       skoXC1 * 19/1000 *
       (skoXC1 * 19/1000 *
        (1/24 +
         skoXC1 * 19/1000 *
         (skoXC1 * 19/1000 *
          (-1/720 +
           skoXC1 * 19/1000 *
           (skoXC1 * 19/1000 *
            (1/40320 +
             skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/3628800))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(resolve, [], [refute_0_59, refute_0_46])).

cnf(refute_0_61, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-361/2000000 +
       skoXC1 *
       (skoXC1 *
        (130321/24000000000000 +
         skoXC1 *
         (skoXC1 *
          (-47045881/720000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (16983563041/40320000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -6131066257801/
              3628800000000000000000000000000000000))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-361/2000000 +
        skoXC1 *
        (skoXC1 *
         (130321/24000000000000 +
          skoXC1 *
          (skoXC1 *
           (-47045881/720000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (16983563041/40320000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -6131066257801/
               3628800000000000000000000000000000000)))))))))) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(resolve, [], [refute_0_61, refute_0_45])).

cnf(refute_0_63, plain,
    (exp(skoXC1 * 3/12500) *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) < 0 |
     0 <
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) |
     0 <= -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(split, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              288160114116647/
              72576000000000000000000000000000000000))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) |
     exp(skoXC1 * -19/1000) <= -3490/23),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (0 < -3490/23 |
     47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              288160114116647/
              72576000000000000000000000000000000000))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000)))))))))),
    inference(resolve, [], [refute_0_65, refute_0_44])).

cnf(refute_0_67, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              288160114116647/
              72576000000000000000000000000000000000))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000)))))))))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_69, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_70, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_69, refute_0_68])).

cnf(refute_0_71, plain,
    (0 < -3490/23 | -3490/23 <= exp(skoXC1 * -19/1000)),
    inference(subst, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_73, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) <
     exp(skoXC1 * 3/12500)), inference(subst, [], [refute_0_3])).

cnf(refute_0_74, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000))))))))))),
    inference(split, [], [refute_0_63])).

cnf(refute_0_75, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000))))))))))),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20),
    inference(resolve, [], [refute_0_75, refute_0_73])).

cnf(refute_0_77, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -288160114116647/
               72576000000000000000000000000000000000)))))))))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= 0),
    inference(resolve, [], [refute_0_72, refute_0_76])).

cnf(refute_0_78, plain,
    (exp(skoXC1 * -19/1000) < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (0 < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20),
    inference(resolve, [], [refute_0_71, refute_0_78])).

cnf(refute_0_80, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -288160114116647/
              72576000000000000000000000000000000000))))))))) <= -47/20),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              288160114116647/
              72576000000000000000000000000000000000)))))))))),
    inference(resolve, [], [refute_0_80, refute_0_67])).

cnf(refute_0_82, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <= 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_83, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) <
     exp(skoXC1 * 3/12500)), inference(subst, [], [refute_0_3])).

cnf(refute_0_84, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20) <
     1 + skoXC1 * (skoXC1 * -361/2000000) |
     (1 + skoXC1 * (skoXC1 * -361/2000000)) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_85, plain,
    (X_000017 < cos(X_000016) | ~ lgen(1, X_000017, cos(X_000016))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_86, plain,
    (~ lgen(1, X_000017, 1 - X_000016 ^ 2 / 2) |
     lgen(1, X_000017, cos(X_000016))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_87, plain,
    (X_000017 < cos(X_000016) | ~ lgen(1, X_000017, 1 - X_000016 ^ 2 / 2)),
    inference(resolve, [], [refute_0_86, refute_0_85])).

cnf(refute_0_88, plain,
    (X_000017 < cos(X_000016) |
     1 + X_000016 * (X_000016 * -1/2) <= X_000017),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20) < cos(skoXC1 * 19/1000) |
     1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/2) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/2) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(resolve, [], [refute_0_59, refute_0_89])).

cnf(refute_0_91, plain,
    (1 + skoXC1 * (skoXC1 * -361/2000000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) *
      (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000)))) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    ((1 + skoXC1 * (skoXC1 * -361/2000000)) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) *
     (skoXC1 * (-399/50000 + skoXC1 * (skoXC1 * 48013/100000000000))) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(resolve, [], [refute_0_91, refute_0_84])).

cnf(refute_0_93, plain,
    (exp(skoXC1 * 3/12500) *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)))),
    inference(split, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)))),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(resolve, [], [refute_0_95, refute_0_83])).

cnf(refute_0_97, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <= 0),
    inference(resolve, [], [refute_0_82, refute_0_96])).

cnf(refute_0_98, plain,
    (exp(skoXC1 * -19/1000) < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (0 < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(resolve, [], [refute_0_71, refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (~ lgen(0, X_000102, exp(X_000101)) | X_000102 <= exp(X_000101)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_102, plain,
    (X_000101 < -1 | ~ lgen(0, X_000102, 1 + X_000101) |
     lgen(0, X_000102, exp(X_000101))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_103, plain,
    (X_000101 < -1 | ~ lgen(0, X_000102, 1 + X_000101) |
     X_000102 <= exp(X_000101)),
    inference(resolve, [], [refute_0_102, refute_0_101])).

cnf(refute_0_104, plain,
    (1 + X_000101 < X_000102 | X_000101 < -1 | X_000102 <= exp(X_000101)),
    inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (1 + skoXC1 * -19/1000 < -3490/23 | skoXC1 * -19/1000 < -1 |
     -3490/23 <= exp(skoXC1 * -19/1000)),
    inference(subst, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (1 + skoXC1 * -19/1000 < -3490/23 | skoXC1 * -19/1000 < -1 |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(resolve, [], [refute_0_105, refute_0_98])).

cnf(refute_0_107, plain,
    (1000/19 < skoXC1 | 3513000/437 < skoXC1 |
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) <=
     -47/20), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) < 0 |
     0 <
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) |
     0 <= -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(split, [], [refute_0_93])).

cnf(refute_0_109, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 * (16967/40000000 + skoXC1 * 48013/100000000000)) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) |
     exp(skoXC1 * -19/1000) <= -3490/23),
    inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (0 < -3490/23 |
     47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 * (16967/40000000 + skoXC1 * 48013/100000000000)) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))),
    inference(resolve, [], [refute_0_109, refute_0_44])).

cnf(refute_0_111, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 * (16967/40000000 + skoXC1 * 48013/100000000000)) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000))),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (0 * (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000) <= 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_113, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) <
     exp(skoXC1 * 3/12500)), inference(subst, [], [refute_0_3])).

cnf(refute_0_114, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20) <
     1 + skoXC1 * (skoXC1 * -361/2000000) |
     (1 + skoXC1 * (skoXC1 * -361/2000000)) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_115, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20) < cos(skoXC1 * 19/1000) |
     1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/2) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_116, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) < skoXC1 * 19/1000 |
     skoXC1 * 19/1000 * (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_117, plain,
    (X_000095 < sin(X_000094) | ~ lgen(1, X_000095, sin(X_000094))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_118, plain,
    (0 < X_000094 | ~ lgen(1, X_000095, X_000094) |
     lgen(1, X_000095, sin(X_000094))),
    inference(subst, [], [sin_lower_bound_0_neg])).

cnf(refute_0_119, plain,
    (0 < X_000094 | X_000095 < sin(X_000094) |
     ~ lgen(1, X_000095, X_000094)),
    inference(resolve, [], [refute_0_118, refute_0_117])).

cnf(refute_0_120, plain,
    (0 < X_000094 | X_000095 < sin(X_000094) | X_000094 <= X_000095),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (0 < skoXC1 * 19/1000 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50) < sin(skoXC1 * 19/1000) |
     skoXC1 * 19/1000 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(subst, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (0 < skoXC1 * 19/1000 |
     skoXC1 * 19/1000 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(resolve, [], [refute_0_22, refute_0_121])).

cnf(refute_0_123, plain,
    (0 < skoXC1 |
     skoXC1 * 19/1000 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20)) /
     (exp(skoXC1 * 3/12500) * 21/50)),
    inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (0 < skoXC1 |
     skoXC1 * 19/1000 * (exp(skoXC1 * 3/12500) * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * -47/20) |
     exp(skoXC1 * 3/12500) * 21/50 <= 0),
    inference(resolve, [], [refute_0_123, refute_0_116])).

cnf(refute_0_125, plain,
    (0 < skoXC1 |
     cos(skoXC1 * 19/1000) * (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (0 < skoXC1 | exp(skoXC1 * 3/12500) * 47/20 <= 0 |
     cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(split, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (0 < skoXC1 |
     cos(skoXC1 * 19/1000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    (0 < skoXC1 |
     1 + skoXC1 * 19/1000 * (skoXC1 * 19/1000 * -1/2) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(resolve, [], [refute_0_127, refute_0_115])).

cnf(refute_0_129, plain,
    (0 < skoXC1 |
     1 + skoXC1 * (skoXC1 * -361/2000000) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
      exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000)) /
     (exp(skoXC1 * 3/12500) * 47/20)),
    inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (0 < skoXC1 |
     (1 + skoXC1 * (skoXC1 * -361/2000000)) *
     (exp(skoXC1 * 3/12500) * 47/20) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 +
     exp(skoXC1 * 3/12500) * (skoXC1 * -399/50000) |
     exp(skoXC1 * 3/12500) * 47/20 <= 0),
    inference(resolve, [], [refute_0_129, refute_0_114])).

cnf(refute_0_131, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 3/12500) *
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (0 < skoXC1 |
     47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000) <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000))),
    inference(split, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (0 < skoXC1 |
     skoXC1 * (399/50000 + skoXC1 * -16967/40000000) <= -47/20 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000))),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 <= 289),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_135, plain,
    (-47/20 < skoXC1 * (399/50000 + skoXC1 * -16967/40000000) |
     0 < skoXC1), inference(decision, [], [refute_0_134, refute_0_30])).

cnf(refute_0_136, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000))),
    inference(resolve, [], [refute_0_133, refute_0_135])).

cnf(refute_0_137, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) |
     0 < skoXC1), inference(resolve, [], [refute_0_136, refute_0_113])).

cnf(refute_0_138, plain,
    (0 * (47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000)) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 | 0 < skoXC1 |
     47/20 + skoXC1 * (399/50000 + skoXC1 * -16967/40000000) <= 0),
    inference(resolve, [], [refute_0_112, refute_0_137])).

cnf(refute_0_139, plain,
    (exp(skoXC1 * -19/1000) < -3490/23 | 0 < skoXC1 |
     skoXC1 * (399/50000 + skoXC1 * -16967/40000000) <= -47/20),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain, (exp(skoXC1 * -19/1000) < -3490/23 | 0 < skoXC1),
    inference(resolve, [], [refute_0_139, refute_0_135])).

cnf(refute_0_141, plain, (0 < -3490/23 | 0 < skoXC1),
    inference(resolve, [], [refute_0_71, refute_0_140])).

cnf(refute_0_142, plain, (0 < skoXC1),
    inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (-47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 * (-16967/40000000 + skoXC1 * -48013/100000000000)) |
     1000/19 < skoXC1 | 3513000/437 < skoXC1),
    inference(decision, [],
              [refute_0_111, refute_0_142, refute_0_134, refute_0_30])).

cnf(refute_0_144, plain, (1000/19 < skoXC1 | 3513000/437 < skoXC1),
    inference(resolve, [], [refute_0_107, refute_0_143])).

cnf(refute_0_145, plain, (1000/19 < skoXC1 | skoXC1 <= 3513000/437),
    inference(decision, [],
              [refute_0_111, refute_0_142, refute_0_134, refute_0_30])).

cnf(refute_0_146, plain, (1000/19 < skoXC1),
    inference(resolve, [], [refute_0_145, refute_0_144])).

cnf(refute_0_147, plain,
    (skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 | skoXC1 <= 12500),
    inference(decision, [],
              [refute_0_81, refute_0_80, refute_0_100, refute_0_146,
               refute_0_111, refute_0_142, refute_0_134, refute_0_30])).

cnf(refute_0_148, plain,
    (skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_147, refute_0_43])).

cnf(refute_0_149, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) < 0 |
     0 <
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     0 <= -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(split, [], [refute_0_39])).

cnf(refute_0_150, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     exp(skoXC1 * -19/1000) <= -3490/23),
    inference(arithmetic, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (0 < -3490/23 |
     47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000))))))))))),
    inference(resolve, [], [refute_0_150, refute_0_44])).

cnf(refute_0_152, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     -47/20 <
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000))))))))))),
    inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <= 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_154, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <
     exp(skoXC1 * 3/12500)), inference(subst, [], [refute_0_3])).

cnf(refute_0_155, plain,
    (47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000)))))))))))),
    inference(split, [], [refute_0_39])).

cnf(refute_0_156, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20), inference(resolve, [], [refute_0_156, refute_0_154])).

cnf(refute_0_158, plain,
    (0 *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20 |
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <= 0),
    inference(resolve, [], [refute_0_153, refute_0_157])).

cnf(refute_0_159, plain,
    (exp(skoXC1 * -19/1000) < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20), inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (0 < -3490/23 |
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20), inference(resolve, [], [refute_0_71, refute_0_159])).

cnf(refute_0_161, plain,
    (skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     -47/20), inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (47/20 <
     skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000))))))))))),
    inference(resolve, [], [refute_0_161, refute_0_152])).

cnf(refute_0_163, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_148, refute_0_162])).

cnf(refute_0_164, plain,
    (0 <=
     47/20 +
     skoXC1 *
     (399/50000 +
      skoXC1 *
      (-16967/40000000 +
       skoXC1 *
       (-48013/100000000000 +
        skoXC1 *
        (6125087/480000000000000 +
         skoXC1 *
         (17332693/2000000000000000000 +
          skoXC1 *
          (-2211156407/14400000000000000000000 +
           skoXC1 *
           (-893871739/12000000000000000000000000 +
            skoXC1 *
            (798227462927/806400000000000000000000000000 +
             skoXC1 *
             (322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (-288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               -116490258898219/
               95040000000000000000000000000000000000000)))))))))) |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (47/20 +
      skoXC1 *
      (399/50000 +
       skoXC1 *
       (-16967/40000000 +
        skoXC1 *
        (-48013/100000000000 +
         skoXC1 *
         (6125087/480000000000000 +
          skoXC1 *
          (17332693/2000000000000000000 +
           skoXC1 *
           (-2211156407/14400000000000000000000 +
            skoXC1 *
            (-893871739/12000000000000000000000000 +
             skoXC1 *
             (798227462927/806400000000000000000000000000 +
              skoXC1 *
              (322687697779/864000000000000000000000000000000 +
               skoXC1 *
               (-288160114116647/72576000000000000000000000000000000000 +
                skoXC1 *
                -116490258898219/
                95040000000000000000000000000000000000000))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_163, refute_0_6])).

cnf(refute_0_165, plain,
    (skoXC1 *
     (-399/50000 +
      skoXC1 *
      (16967/40000000 +
       skoXC1 *
       (48013/100000000000 +
        skoXC1 *
        (-6125087/480000000000000 +
         skoXC1 *
         (-17332693/2000000000000000000 +
          skoXC1 *
          (2211156407/14400000000000000000000 +
           skoXC1 *
           (893871739/12000000000000000000000000 +
            skoXC1 *
            (-798227462927/806400000000000000000000000000 +
             skoXC1 *
             (-322687697779/864000000000000000000000000000000 +
              skoXC1 *
              (288160114116647/72576000000000000000000000000000000000 +
               skoXC1 *
               116490258898219/
               95040000000000000000000000000000000000000)))))))))) <=
     47/20 |
     (255483445248/5 +
      skoXC1 *
      (542222843904/3125 +
       skoXC1 *
       (-720543154176/78125 +
        skoXC1 *
        (-509745954816/48828125 +
         skoXC1 *
         (338692810752/1220703125 +
          skoXC1 *
          (718821444096/3814697265625 +
           skoXC1 *
           (-318406522608/95367431640625 +
            skoXC1 *
            (-96538147812/59604644775390625 +
             skoXC1 *
             (7184047166343/333786010742187500000 +
              skoXC1 *
              (968063093337/119209289550781250000000 +
               skoXC1 *
               (-288160114116647/3337860107421875000000000000 +
                skoXC1 *
                -349470776694657/
                13113021850585937500000000000000))))))))))) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_164])).

cnf(refute_0_166, plain,
    ((255483445248/5 +
      skoXC1 *
      (542222843904/3125 +
       skoXC1 *
       (-720543154176/78125 +
        skoXC1 *
        (-509745954816/48828125 +
         skoXC1 *
         (338692810752/1220703125 +
          skoXC1 *
          (718821444096/3814697265625 +
           skoXC1 *
           (-318406522608/95367431640625 +
            skoXC1 *
            (-96538147812/59604644775390625 +
             skoXC1 *
             (7184047166343/333786010742187500000 +
              skoXC1 *
              (968063093337/119209289550781250000000 +
               skoXC1 *
               (-288160114116647/3337860107421875000000000000 +
                skoXC1 *
                -349470776694657/
                13113021850585937500000000000000))))))))))) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_165, refute_0_162])).

cnf(refute_0_167, plain,
    (21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     255483445248/5 +
     skoXC1 *
     (542222843904/3125 +
      skoXC1 *
      (-720543154176/78125 +
       skoXC1 *
       (-509745954816/48828125 +
        skoXC1 *
        (338692810752/1220703125 +
         skoXC1 *
         (718821444096/3814697265625 +
          skoXC1 *
          (-318406522608/95367431640625 +
           skoXC1 *
           (-96538147812/59604644775390625 +
            skoXC1 *
            (7184047166343/333786010742187500000 +
             skoXC1 *
             (968063093337/119209289550781250000000 +
              skoXC1 *
              (-288160114116647/3337860107421875000000000000 +
               skoXC1 *
               -349470776694657/
               13113021850585937500000000000000)))))))))) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_166, refute_0_5])).

cnf(refute_0_168, plain,
    (exp(skoXC1 * -19/1000) *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -3174517702656/25 +
     skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_167])).

cnf(refute_0_169, plain,
    (-21743271936 <
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_162, refute_0_161, refute_0_81, refute_0_80,
               refute_0_100, refute_0_146, refute_0_111, refute_0_142,
               refute_0_134, refute_0_30])).

cnf(refute_0_170, plain,
    (exp(skoXC1 * -19/1000) *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -3174517702656/25 +
     skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0 |
     exp(skoXC1 * -19/1000) <=
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(split, [], [refute_0_170])).

cnf(refute_0_172, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125 |
     exp(skoXC1 * -19/1000) <=
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_171])).

cnf(refute_0_173, plain,
    (-62511906816/125 <
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_162, refute_0_161, refute_0_81, refute_0_80,
               refute_0_100, refute_0_146, refute_0_111, refute_0_142,
               refute_0_134, refute_0_30])).

cnf(refute_0_174, plain,
    (exp(skoXC1 * -19/1000) <=
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (0 <
     (-3174517702656/25 +
      skoXC1 *
      (-12132745740288/78125 +
       skoXC1 *
       (2251163797291008/244140625 +
        skoXC1 *
        (7964913933877248/762939453125 +
         skoXC1 *
         (-661509420619811328/2384185791015625 +
          skoXC1 *
          (-1403948129482884096/7450580596923828125 +
           skoXC1 *
           (77735967433195326714/23283064365386962890625 +
            skoXC1 *
            (471377674863425506707/291038304567337036132812500 +
             skoXC1 *
             (-35078355304409198296615203/
              1629814505577087402343750000000000 +
              skoXC1 *
              (-23634352864672850964865071/
               2910383045673370361328125000000000000 +
               skoXC1 *
               (175878975901273803705311488427/
                2037268131971359252929687500000000000000000 +
                skoXC1 *
                (2133000346036724853517665201999/
                 80035533756017684936523437500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_174, refute_0_4])).

cnf(refute_0_176, plain,
    (0 *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     -3174517702656/25 +
     skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_175])).

cnf(refute_0_177, plain,
    (3174517702656/25 <
     skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 *
     (-12132745740288/78125 +
      skoXC1 *
      (2251163797291008/244140625 +
       skoXC1 *
       (7964913933877248/762939453125 +
        skoXC1 *
        (-661509420619811328/2384185791015625 +
         skoXC1 *
         (-1403948129482884096/7450580596923828125 +
          skoXC1 *
          (77735967433195326714/23283064365386962890625 +
           skoXC1 *
           (471377674863425506707/291038304567337036132812500 +
            skoXC1 *
            (-35078355304409198296615203/
             1629814505577087402343750000000000 +
             skoXC1 *
             (-23634352864672850964865071/
              2910383045673370361328125000000000000 +
              skoXC1 *
              (175878975901273803705311488427/
               2037268131971359252929687500000000000000000 +
               skoXC1 *
               (2133000346036724853517665201999/
                80035533756017684936523437500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) <=
     3174517702656/25 |
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125),
    inference(decision, [],
              [refute_0_162, refute_0_161, refute_0_81, refute_0_80,
               refute_0_100, refute_0_146, refute_0_111, refute_0_142,
               refute_0_134, refute_0_30])).

cnf(refute_0_179, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125),
    inference(resolve, [], [refute_0_178, refute_0_177])).

cnf(refute_0_180, plain, ($false),
    inference(resolve, [], [refute_0_179, refute_0_173])).
% SZS output end CNFRefutation for Chua-1-IL-U-sincos.tptp
