% SZS output start CNFRefutation for Chua-1-VC2-U.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Chua', conjecture,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        189/10000 + 77/10000 * exp(-19/1000 * X) -
        183/1000 * exp(3/12500 * X) * C + 793/1000 * exp(3/12500 * X) * S <
        2))).

fof(subgoal_0, plain,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        189/10000 + 77/10000 * exp(-19/1000 * X) -
        183/1000 * exp(3/12500 * X) * C + 793/1000 * exp(3/12500 * X) * S <
        2)), inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X, S, C] :
         ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
          189/10000 + 77/10000 * exp(-19/1000 * X) -
          183/1000 * exp(3/12500 * X) * C +
          793/1000 * exp(3/12500 * X) * S < 2)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, S, X] :
       (X <= 289 & 0 <= X &
        2 <=
        189/10000 + 77/10000 * exp(-19/1000 * X) -
        183/1000 * exp(3/12500 * X) * C + 793/1000 * exp(3/12500 * X) * S &
        S ^ 2 + C ^ 2 = 1)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     2 <=
     189/10000 + 77/10000 * exp(-19/1000 * skoXC1) -
     183/1000 * exp(3/12500 * skoXC1) * skoCC1 +
     793/1000 * exp(3/12500 * skoXC1) * skoSC1 & skoXC1 <= 289 &
     skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (2 <=
     189/10000 + 77/10000 * exp(-19/1000 * skoXC1) -
     183/1000 * exp(3/12500 * skoXC1) * skoCC1 +
     793/1000 * exp(3/12500 * skoXC1) * skoSC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 289),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000))))))))))) <
     (-1255673954304/625 +
      skoXC1 *
      (38219575984128/1953125 +
       skoXC1 *
       (-397206622568448/6103515625 +
        skoXC1 *
        (287244000755712/19073486328125 +
         skoXC1 *
         (-106389138949632/59604644775390625 +
          skoXC1 *
          (26424041234976/186264514923095703125 +
           skoXC1 *
           (-4856747733684/582076609134674072265625 +
            skoXC1 *
            (5534106272541/14551915228366851806640625000 +
             skoXC1 *
             (-80061035466951/5820766091346740722656250000000000 +
              skoXC1 *
              (115734346533729/291038304567337036132812500000000000000 +
               skoXC1 *
               (-133089315419907/
                14551915228366851806640625000000000000000000 +
                skoXC1 *
                (59753140606467/
                 363797880709171295166015625000000000000000000000 +
                 skoXC1 *
                 (-80250893551521/
                  36379788070917129516601562500000000000000000000000000 +
                  skoXC1 *
                  (1479579042249/
                   72759576141834259033203125000000000000000000000000000000 +
                   skoXC1 *
                   -14772465477/
                   145519152283668518066406250000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     -1255673954304/625 +
     skoXC1 *
     (38219575984128/1953125 +
      skoXC1 *
      (-397206622568448/6103515625 +
       skoXC1 *
       (287244000755712/19073486328125 +
        skoXC1 *
        (-106389138949632/59604644775390625 +
         skoXC1 *
         (26424041234976/186264514923095703125 +
          skoXC1 *
          (-4856747733684/582076609134674072265625 +
           skoXC1 *
           (5534106272541/14551915228366851806640625000 +
            skoXC1 *
            (-80061035466951/5820766091346740722656250000000000 +
             skoXC1 *
             (115734346533729/291038304567337036132812500000000000000 +
              skoXC1 *
              (-133089315419907/
               14551915228366851806640625000000000000000000 +
               skoXC1 *
               (59753140606467/
                363797880709171295166015625000000000000000000000 +
                skoXC1 *
                (-80250893551521/
                 36379788070917129516601562500000000000000000000000000 +
                 skoXC1 *
                 (1479579042249/
                  72759576141834259033203125000000000000000000000000000000 +
                  skoXC1 *
                  -14772465477/
                  145519152283668518066406250000000000000000000000000000000000))))))))))))) <=
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) |
     0 <=
     -104639496192/625 +
     skoXC1 *
     (78479622144/1953125 +
      skoXC1 *
      (-29429858304/6103515625 +
       skoXC1 *
       (7357464576/19073486328125 +
        skoXC1 *
        (-1357969536/59604644775390625 +
         skoXC1 *
         (193995648/186264514923095703125 +
          skoXC1 *
          (-43952139/1164153218269348144531250 +
           skoXC1 *
           (31827411/29103830456733703613281250000 +
            skoXC1 *
            (-586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (-177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (40920957/727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -40920957/
                145519152283668518066406250000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    (exp(X_000064) < X_000065 | ~ lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064 |
     lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_4, plain,
    (exp(X_000064) < X_000065 |
     ~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (exp(X_000064) < X_000065 |
     X_000065 <=
     (12 + X_000064 * (6 + X_000064)) / (12 + X_000064 * (-6 + X_000064)) |
     0 <= X_000064), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (exp(skoXC1 * -19/1000) <
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) |
     0 <= skoXC1 * -19/1000 |
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 <
     (skoCC1 * 497377345536/125 + skoSC1 * -2155301830656/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * 497377345536/125 + skoSC1 * -2155301830656/125 <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) |
     0 <= skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_9, plain,
    (2 <=
     189/10000 + 77/10000 * exp(-19/1000 * skoXC1) -
     183/1000 * exp(3/12500 * skoXC1) * skoCC1 +
     793/1000 * exp(3/12500 * skoXC1) * skoSC1),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    (exp(skoXC1 * 3/12500) * (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (0 <= skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <= exp(skoXC1 * 3/12500)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_12, plain,
    ((-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <= exp(skoXC1 * 3/12500) |
     skoSC1 <= skoCC1 * 3/13), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (exp(X_000061) < X_000062 | ~ lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (3 < X_000061 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_15, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     X_000062 <=
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061))))))))))))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (3 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (3 < skoXC1 * 3/12500 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500)))))))))))) |
     skoSC1 <= skoCC1 * 3/13),
    inference(resolve, [], [refute_0_12, refute_0_17])).

cnf(refute_0_19, plain,
    (12500 < skoXC1 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     skoSC1 <= skoCC1 * 3/13), inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000))))) <
     (-25546752/625 +
      skoXC1 *
      (19160064/1953125 +
       skoXC1 *
       (-7185024/6103515625 +
        skoXC1 *
        (1796256/19073486328125 +
         skoXC1 *
         (-1178793/238418579101562500 +
          skoXC1 *
          (505197/2980232238769531250000 +
           skoXC1 * -505197/149011611938476562500000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     -25546752/625 +
     skoXC1 *
     (19160064/1953125 +
      skoXC1 *
      (-7185024/6103515625 +
       skoXC1 *
       (1796256/19073486328125 +
        skoXC1 *
        (-1178793/238418579101562500 +
         skoXC1 *
         (505197/2980232238769531250000 +
          skoXC1 * -505197/149011611938476562500000000))))) <=
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_21, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) |
     0 <=
     -11088/625 +
     skoXC1 *
     (8316/1953125 +
      skoXC1 *
      (-6237/12207031250 +
       skoXC1 *
       (6237/152587890625000 +
        skoXC1 *
        (-130977/61035156250000000000 +
         skoXC1 *
         (56133/762939453125000000000000 +
          skoXC1 * -56133/38146972656250000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000)))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_22, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_23, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_24, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_23, refute_0_22])).

cnf(refute_0_25, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (0 < skoXC1 * -19/1000 |
     exp(skoXC1 * -19/1000) <
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 <
     (skoCC1 * 52704/125 + skoSC1 * -228384/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * 52704/125 + skoSC1 * -228384/125 <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_28, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) |
     0 <= skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_29, plain,
    (0 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_30, plain,
    (0 < skoXC1 * 3/12500 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500)))))) |
     skoSC1 <= skoCC1 * 3/13),
    inference(resolve, [], [refute_0_12, refute_0_29])).

cnf(refute_0_31, plain,
    (0 < skoXC1 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     skoSC1 <= skoCC1 * 3/13), inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (0 < skoXC1 | 0 <= skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 |
     skoSC1 <= skoCC1 * 3/13),
    inference(resolve, [], [refute_0_31, refute_0_28])).

cnf(refute_0_33, plain,
    (0 < skoXC1 |
     (skoCC1 * 52704/125 + skoSC1 * -228384/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 |
     skoSC1 <= skoCC1 * 3/13), inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 1801/7 |
     1801/7 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_35, plain,
    (skoCC1 * 183/1000 + skoSC1 * -793/1000 < 0 |
     0 < skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     0 <= -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(split, [], [refute_0_10])).

cnf(refute_0_36, plain,
    (skoSC1 < skoCC1 * 3/13 | skoCC1 * 3/13 < skoSC1 |
     1801/7 <= exp(skoXC1 * -19/1000)),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (exp(skoXC1 * -19/1000) < 1801/7 | 0 < skoXC1 * -19/1000 |
     1801/7 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_38, plain,
    (skoSC1 < skoCC1 * 3/13 | 0 < skoXC1 * -19/1000 |
     skoCC1 * 3/13 < skoSC1 |
     1801/7 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 < 0 | skoSC1 < skoCC1 * 3/13 | skoCC1 * 3/13 < skoSC1 |
     1801/7 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (1801/7 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 1801/7),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (skoCC1 * 183/1000 + skoSC1 * -793/1000 <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_42, plain,
    (skoCC1 * 3/13 <= skoSC1 |
     exp(skoXC1 * 3/12500) <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_45, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (0 <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) < exp(skoXC1 * 3/12500)),
    inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (0 <
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) | skoCC1 * 3/13 <= skoSC1),
    inference(resolve, [], [refute_0_42, refute_0_46])).

cnf(refute_0_48, plain,
    (0 * (skoCC1 * 183/1000 + skoSC1 * -793/1000) <
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 |
     skoCC1 * 183/1000 + skoSC1 * -793/1000 <= 0 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_49, plain,
    (0 * (skoCC1 * 183/1000 + skoSC1 * -793/1000) <
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 |
     skoCC1 * 3/13 <= skoSC1 |
     skoCC1 * 183/1000 + skoSC1 * -793/1000 <= 0),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    (1801/7 < exp(skoXC1 * -19/1000) | skoCC1 * 3/13 <= skoSC1),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_52, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_53, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_52, refute_0_51])).

cnf(refute_0_54, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 * -19/1000 |
     1801/7 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     exp(skoXC1 * -19/1000) <= 1801/7),
    inference(subst, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (0 < skoXC1 * -19/1000 |
     1801/7 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     skoCC1 * 3/13 <= skoSC1),
    inference(resolve, [], [refute_0_55, refute_0_50])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 |
     1801/7 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoCC1 * 3/13 <= skoSC1), inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_59, plain, (skoSC1 * skoSC1 = 1 + skoCC1 * (skoCC1 * -1)),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain, (skoXC1 <= 289),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_61, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_62, plain, (0 <= skoXC1 | skoCC1 * 3/13 <= skoSC1),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_63, plain,
    (1801/7 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoCC1 * 3/13 <= skoSC1),
    inference(resolve, [], [refute_0_62, refute_0_57])).

cnf(refute_0_64, plain,
    (1801/7 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     skoCC1 * 3/13 <= skoSC1 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_40, refute_0_63])).

cnf(refute_0_65, plain,
    (4133376/7 <
     skoXC1 *
     (-9855072/875 +
      skoXC1 *
      (-11702898/109375 +
       skoXC1 *
       (-37059177/54687500 +
        skoXC1 *
        (-704124363/250000000000 +
         skoXC1 *
         (-13378362897/1750000000000000 +
          skoXC1 * -84729631681/7000000000000000000))))) |
     skoCC1 * 3/13 <= skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoCC1 * 3/13 <= skoSC1 |
     skoXC1 *
     (-9855072/875 +
      skoXC1 *
      (-11702898/109375 +
       skoXC1 *
       (-37059177/54687500 +
        skoXC1 *
        (-704124363/250000000000 +
         skoXC1 *
         (-13378362897/1750000000000000 +
          skoXC1 * -84729631681/7000000000000000000))))) <= 4133376/7 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_67, plain,
    (skoCC1 * 3/13 <= skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_66, refute_0_65])).

cnf(refute_0_68, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoCC1 * 3/13 <= skoSC1),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_69, plain, (skoCC1 * 3/13 <= skoSC1),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 | skoCC1 * 3/13 < skoSC1 | skoCC1 * 3/13 <= skoSC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 | skoCC1 * 3/13 < skoSC1 |
     1801/7 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_70, refute_0_39])).

cnf(refute_0_72, plain, (skoCC1 * 3/13 < skoSC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_73, plain,
    (skoCC1 * 3/13 < skoSC1 |
     1801/7 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_72, refute_0_71])).

cnf(refute_0_74, plain,
    (skoCC1 * 3/13 < skoSC1 |
     1801/7 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_73, refute_0_34])).

cnf(refute_0_75, plain,
    (skoCC1 * 3/13 < skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (9855072/875 +
      skoXC1 *
      (11702898/109375 +
       skoXC1 *
       (37059177/54687500 +
        skoXC1 *
        (704124363/250000000000 +
         skoXC1 *
         (13378362897/1750000000000000 +
          skoXC1 * 84729631681/7000000000000000000))))) <= -4133376/7),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoCC1 * 3/13 < skoSC1 |
     skoXC1 *
     (9855072/875 +
      skoXC1 *
      (11702898/109375 +
       skoXC1 *
       (37059177/54687500 +
        skoXC1 *
        (704124363/250000000000 +
         skoXC1 *
         (13378362897/1750000000000000 +
          skoXC1 * 84729631681/7000000000000000000))))) <= -4133376/7),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_77, plain,
    (skoCC1 * 3/13 < skoSC1 |
     skoXC1 *
     (9855072/875 +
      skoXC1 *
      (11702898/109375 +
       skoXC1 *
       (37059177/54687500 +
        skoXC1 *
        (704124363/250000000000 +
         skoXC1 *
         (13378362897/1750000000000000 +
          skoXC1 * 84729631681/7000000000000000000))))) <= -4133376/7),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (-4133376/7 <
     skoXC1 *
     (9855072/875 +
      skoXC1 *
      (11702898/109375 +
       skoXC1 *
       (37059177/54687500 +
        skoXC1 *
        (704124363/250000000000 +
         skoXC1 *
         (13378362897/1750000000000000 +
          skoXC1 * 84729631681/7000000000000000000))))) |
     skoCC1 * 3/13 < skoSC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_79, plain, (skoCC1 * 3/13 < skoSC1),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain, (skoCC1 * 3/13 < skoSC1 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_81, plain,
    (0 < skoXC1 |
     (skoCC1 * 52704/125 + skoSC1 * -228384/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(resolve, [], [refute_0_33, refute_0_80])).

cnf(refute_0_82, plain,
    (0 < skoXC1 |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * 52704/125 + skoSC1 * -228384/125 <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(resolve, [], [refute_0_81, refute_0_27])).

cnf(refute_0_83, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000))))) |
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-2304 <
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_85, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000)))))),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (0 < skoXC1 |
     0 <=
     -11088/625 +
     skoXC1 *
     (8316/1953125 +
      skoXC1 *
      (-6237/12207031250 +
       skoXC1 *
       (6237/152587890625000 +
        skoXC1 *
        (-130977/61035156250000000000 +
         skoXC1 *
         (56133/762939453125000000000000 +
          skoXC1 * -56133/38146972656250000000000000000))))) |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-8316/1953125 +
      skoXC1 *
      (6237/12207031250 +
       skoXC1 *
       (-6237/152587890625000 +
        skoXC1 *
        (130977/61035156250000000000 +
         skoXC1 *
         (-56133/762939453125000000000000 +
          skoXC1 * 56133/38146972656250000000000000000))))) <= -11088/625 |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (-11088/625 <
     skoXC1 *
     (-8316/1953125 +
      skoXC1 *
      (6237/12207031250 +
       skoXC1 *
       (-6237/152587890625000 +
        skoXC1 *
        (130977/61035156250000000000 +
         skoXC1 *
         (-56133/762939453125000000000000 +
          skoXC1 * 56133/38146972656250000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_89, plain,
    (0 < skoXC1 |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (0 < skoXC1 * -19/1000 | 0 < skoXC1 |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_89, refute_0_26])).

cnf(refute_0_91, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain, (0 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_93, plain,
    (0 < skoXC1 |
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) /
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (0 < skoXC1 |
     0 <=
     -11088/625 +
     skoXC1 *
     (8316/1953125 +
      skoXC1 *
      (-6237/12207031250 +
       skoXC1 *
       (6237/152587890625000 +
        skoXC1 *
        (-130977/61035156250000000000 +
         skoXC1 *
         (56133/762939453125000000000000 +
          skoXC1 * -56133/38146972656250000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-11088/625 +
      skoXC1 *
      (8316/1953125 +
       skoXC1 *
       (-6237/12207031250 +
        skoXC1 *
        (6237/152587890625000 +
         skoXC1 *
         (-130977/61035156250000000000 +
          skoXC1 *
          (56133/762939453125000000000000 +
           skoXC1 * -56133/38146972656250000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000)))))),
    inference(resolve, [], [refute_0_93, refute_0_21])).

cnf(refute_0_95, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-8316/1953125 +
      skoXC1 *
      (6237/12207031250 +
       skoXC1 *
       (-6237/152587890625000 +
        skoXC1 *
        (130977/61035156250000000000 +
         skoXC1 *
         (-56133/762939453125000000000000 +
          skoXC1 * 56133/38146972656250000000000000000))))) <= -11088/625 |
     (-25546752/625 +
      skoXC1 *
      (19160064/1953125 +
       skoXC1 *
       (-7185024/6103515625 +
        skoXC1 *
        (1796256/19073486328125 +
         skoXC1 *
         (-1178793/238418579101562500 +
          skoXC1 *
          (505197/2980232238769531250000 +
           skoXC1 * -505197/149011611938476562500000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000)))))),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 < skoXC1 |
     (-25546752/625 +
      skoXC1 *
      (19160064/1953125 +
       skoXC1 *
       (-7185024/6103515625 +
        skoXC1 *
        (1796256/19073486328125 +
         skoXC1 *
         (-1178793/238418579101562500 +
          skoXC1 *
          (505197/2980232238769531250000 +
           skoXC1 * -505197/149011611938476562500000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
     skoXC1 *
     (2139588/1953125 +
      skoXC1 *
      (-1604691/12207031250 +
       skoXC1 *
       (1604691/152587890625000 +
        skoXC1 *
        (-33698511/61035156250000000000 +
         skoXC1 *
         (14442219/762939453125000000000000 +
          skoXC1 * -14442219/38146972656250000000000000000)))))),
    inference(resolve, [], [refute_0_95, refute_0_88])).

cnf(refute_0_97, plain,
    (0 < skoXC1 |
     -25546752/625 +
     skoXC1 *
     (19160064/1953125 +
      skoXC1 *
      (-7185024/6103515625 +
       skoXC1 *
       (1796256/19073486328125 +
        skoXC1 *
        (-1178793/238418579101562500 +
         skoXC1 *
         (505197/2980232238769531250000 +
          skoXC1 * -505197/149011611938476562500000000))))) <=
     (-2852784/625 + skoCC1 * -52704/125 + skoSC1 * 228384/125 +
      skoXC1 *
      (2139588/1953125 +
       skoXC1 *
       (-1604691/12207031250 +
        skoXC1 *
        (1604691/152587890625000 +
         skoXC1 *
         (-33698511/61035156250000000000 +
          skoXC1 *
          (14442219/762939453125000000000000 +
           skoXC1 * -14442219/38146972656250000000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_20])).

cnf(refute_0_98, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (385350400512/1953125 + skoCC1 * 288396288/15625 +
      skoSC1 * -1249717248/15625 +
      skoXC1 *
      (11295014800608/6103515625 + skoCC1 * 342470592/1953125 +
       skoSC1 * -1484039232/1953125 +
       skoXC1 *
       (220723556432148/19073486328125 + skoCC1 * 271122552/244140625 +
        skoSC1 * -1174864392/244140625 +
        skoXC1 *
        (179808884185988871/3814697265625000000 +
         skoCC1 * 4507412427/976562500000 +
         skoSC1 * -19532120517/976562500000 +
         skoXC1 *
         (11823112384976380257/95367431640625000000000 +
          skoCC1 * 12234405159/976562500000000 +
          skoSC1 * -53015755689/976562500000000 +
          skoXC1 *
          (1750950599566531600611/9536743164062500000000000000 +
           skoCC1 * 77484566007/3906250000000000000 +
           skoSC1 * -335766452697/3906250000000000000 +
           skoXC1 *
           (-227665896064691410557/4768371582031250000000000000000 +
            skoXC1 *
            (56056902795536294817/9536743164062500000000000000000000 +
             skoXC1 *
             (-4563896492594205459/9536743164062500000000000000000000000 +
              skoXC1 *
              (969561966526758891/
               38146972656250000000000000000000000000000 +
               skoXC1 *
               (-16771610727106389/
                19073486328125000000000000000000000000000000 +
                skoXC1 *
                679446916449939/
                38146972656250000000000000000000000000000000000))))))))))) <=
     -6547267584/625 + skoCC1 * -121430016/125 + skoSC1 * 526196736/125),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (-6547267584/625 + skoCC1 * -121430016/125 + skoSC1 * 526196736/125 <
     skoXC1 *
     (385350400512/1953125 + skoCC1 * 288396288/15625 +
      skoSC1 * -1249717248/15625 +
      skoXC1 *
      (11295014800608/6103515625 + skoCC1 * 342470592/1953125 +
       skoSC1 * -1484039232/1953125 +
       skoXC1 *
       (220723556432148/19073486328125 + skoCC1 * 271122552/244140625 +
        skoSC1 * -1174864392/244140625 +
        skoXC1 *
        (179808884185988871/3814697265625000000 +
         skoCC1 * 4507412427/976562500000 +
         skoSC1 * -19532120517/976562500000 +
         skoXC1 *
         (11823112384976380257/95367431640625000000000 +
          skoCC1 * 12234405159/976562500000000 +
          skoSC1 * -53015755689/976562500000000 +
          skoXC1 *
          (1750950599566531600611/9536743164062500000000000000 +
           skoCC1 * 77484566007/3906250000000000000 +
           skoSC1 * -335766452697/3906250000000000000 +
           skoXC1 *
           (-227665896064691410557/4768371582031250000000000000000 +
            skoXC1 *
            (56056902795536294817/9536743164062500000000000000000000 +
             skoXC1 *
             (-4563896492594205459/9536743164062500000000000000000000000 +
              skoXC1 *
              (969561966526758891/
               38146972656250000000000000000000000000000 +
               skoXC1 *
               (-16771610727106389/
                19073486328125000000000000000000000000000000 +
                skoXC1 *
                679446916449939/
                38146972656250000000000000000000000000000000000))))))))))) |
     0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_100, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_102, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain, (skoCC1 * 3/13 < skoSC1 | 12500 < skoXC1),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_104, plain,
    (12500 < skoXC1 |
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_19, refute_0_103])).

cnf(refute_0_105, plain, (skoXC1 <= 12500),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_106, plain,
    ((-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) /
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    (0 <= skoCC1 * 183/1000 + skoSC1 * -793/1000 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * 183/1000 + skoSC1 * -793/1000) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(resolve, [], [refute_0_106, refute_0_8])).

cnf(refute_0_108, plain,
    ((skoCC1 * 497377345536/125 + skoSC1 * -2155301830656/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000 |
     skoSC1 <= skoCC1 * 3/13), inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    ((skoCC1 * 497377345536/125 + skoSC1 * -2155301830656/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -19811/10000 + exp(skoXC1 * -19/1000) * 77/10000),
    inference(resolve, [], [refute_0_108, refute_0_79])).

cnf(refute_0_110, plain,
    (21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * 497377345536/125 + skoSC1 * -2155301830656/125 <=
     (-19811/10000 + exp(skoXC1 * -19/1000) * 77/10000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_109, refute_0_7])).

cnf(refute_0_111, plain,
    (exp(skoXC1 * -19/1000) *
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000))))))))))) |
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-21743271936 <
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_113, plain,
    (exp(skoXC1 * -19/1000) *
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (0 <=
     -104639496192/625 +
     skoXC1 *
     (78479622144/1953125 +
      skoXC1 *
      (-29429858304/6103515625 +
       skoXC1 *
       (7357464576/19073486328125 +
        skoXC1 *
        (-1357969536/59604644775390625 +
         skoXC1 *
         (193995648/186264514923095703125 +
          skoXC1 *
          (-43952139/1164153218269348144531250 +
           skoXC1 *
           (31827411/29103830456733703613281250000 +
            skoXC1 *
            (-586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (-177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (40920957/727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -40920957/
                145519152283668518066406250000000000000000000000000000))))))))))) |
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (-78479622144/1953125 +
      skoXC1 *
      (29429858304/6103515625 +
       skoXC1 *
       (-7357464576/19073486328125 +
        skoXC1 *
        (1357969536/59604644775390625 +
         skoXC1 *
         (-193995648/186264514923095703125 +
          skoXC1 *
          (43952139/1164153218269348144531250 +
           skoXC1 *
           (-31827411/29103830456733703613281250000 +
            skoXC1 *
            (586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (-131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (-40920957/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                40920957/
                145519152283668518066406250000000000000000000000000000))))))))))) <=
     -104639496192/625 |
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (-104639496192/625 <
     skoXC1 *
     (-78479622144/1953125 +
      skoXC1 *
      (29429858304/6103515625 +
       skoXC1 *
       (-7357464576/19073486328125 +
        skoXC1 *
        (1357969536/59604644775390625 +
         skoXC1 *
         (-193995648/186264514923095703125 +
          skoXC1 *
          (43952139/1164153218269348144531250 +
           skoXC1 *
           (-31827411/29103830456733703613281250000 +
            skoXC1 *
            (586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (-131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (-40920957/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                40920957/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_117, plain,
    ((-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (0 <= skoXC1 * -19/1000 |
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(resolve, [], [refute_0_117, refute_0_6])).

cnf(refute_0_119, plain,
    ((-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) /
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000))),
    inference(resolve, [], [refute_0_119, refute_0_102])).

cnf(refute_0_121, plain,
    (0 <=
     -104639496192/625 +
     skoXC1 *
     (78479622144/1953125 +
      skoXC1 *
      (-29429858304/6103515625 +
       skoXC1 *
       (7357464576/19073486328125 +
        skoXC1 *
        (-1357969536/59604644775390625 +
         skoXC1 *
         (193995648/186264514923095703125 +
          skoXC1 *
          (-43952139/1164153218269348144531250 +
           skoXC1 *
           (31827411/29103830456733703613281250000 +
            skoXC1 *
            (-586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (-177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (40920957/727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -40920957/
                145519152283668518066406250000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-104639496192/625 +
      skoXC1 *
      (78479622144/1953125 +
       skoXC1 *
       (-29429858304/6103515625 +
        skoXC1 *
        (7357464576/19073486328125 +
         skoXC1 *
         (-1357969536/59604644775390625 +
          skoXC1 *
          (193995648/186264514923095703125 +
           skoXC1 *
           (-43952139/1164153218269348144531250 +
            skoXC1 *
            (31827411/29103830456733703613281250000 +
             skoXC1 *
             (-586533717/23283064365386962890625000000000000 +
              skoXC1 *
              (131856417/291038304567337036132812500000000000000 +
               skoXC1 *
               (-177324147/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (40920957/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -40920957/
                 145519152283668518066406250000000000000000000000000000)))))))))))) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_120, refute_0_1])).

cnf(refute_0_122, plain,
    (skoXC1 *
     (-78479622144/1953125 +
      skoXC1 *
      (29429858304/6103515625 +
       skoXC1 *
       (-7357464576/19073486328125 +
        skoXC1 *
        (1357969536/59604644775390625 +
         skoXC1 *
         (-193995648/186264514923095703125 +
          skoXC1 *
          (43952139/1164153218269348144531250 +
           skoXC1 *
           (-31827411/29103830456733703613281250000 +
            skoXC1 *
            (586533717/23283064365386962890625000000000000 +
             skoXC1 *
             (-131856417/291038304567337036132812500000000000000 +
              skoXC1 *
              (177324147/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (-40920957/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                40920957/
                145519152283668518066406250000000000000000000000000000))))))))))) <=
     -104639496192/625 |
     (-1255673954304/625 +
      skoXC1 *
      (38219575984128/1953125 +
       skoXC1 *
       (-397206622568448/6103515625 +
        skoXC1 *
        (287244000755712/19073486328125 +
         skoXC1 *
         (-106389138949632/59604644775390625 +
          skoXC1 *
          (26424041234976/186264514923095703125 +
           skoXC1 *
           (-4856747733684/582076609134674072265625 +
            skoXC1 *
            (5534106272541/14551915228366851806640625000 +
             skoXC1 *
             (-80061035466951/5820766091346740722656250000000000 +
              skoXC1 *
              (115734346533729/291038304567337036132812500000000000000 +
               skoXC1 *
               (-133089315419907/
                14551915228366851806640625000000000000000000 +
                skoXC1 *
                (59753140606467/
                 363797880709171295166015625000000000000000000000 +
                 skoXC1 *
                 (-80250893551521/
                  36379788070917129516601562500000000000000000000000000 +
                  skoXC1 *
                  (1479579042249/
                   72759576141834259033203125000000000000000000000000000000 +
                   skoXC1 *
                   -14772465477/
                   145519152283668518066406250000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    ((-1255673954304/625 +
      skoXC1 *
      (38219575984128/1953125 +
       skoXC1 *
       (-397206622568448/6103515625 +
        skoXC1 *
        (287244000755712/19073486328125 +
         skoXC1 *
         (-106389138949632/59604644775390625 +
          skoXC1 *
          (26424041234976/186264514923095703125 +
           skoXC1 *
           (-4856747733684/582076609134674072265625 +
            skoXC1 *
            (5534106272541/14551915228366851806640625000 +
             skoXC1 *
             (-80061035466951/5820766091346740722656250000000000 +
              skoXC1 *
              (115734346533729/291038304567337036132812500000000000000 +
               skoXC1 *
               (-133089315419907/
                14551915228366851806640625000000000000000000 +
                skoXC1 *
                (59753140606467/
                 363797880709171295166015625000000000000000000000 +
                 skoXC1 *
                 (-80250893551521/
                  36379788070917129516601562500000000000000000000000000 +
                  skoXC1 *
                  (1479579042249/
                   72759576141834259033203125000000000000000000000000000000 +
                   skoXC1 *
                   -14772465477/
                   145519152283668518066406250000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -26922247520256/625 + skoCC1 * -497377345536/125 +
     skoSC1 * 2155301830656/125 +
     skoXC1 *
     (20191685640192/1953125 +
      skoXC1 *
      (-7571882115072/6103515625 +
       skoXC1 *
       (1892970528768/19073486328125 +
        skoXC1 *
        (-349386162048/59604644775390625 +
         skoXC1 *
         (49912308864/186264514923095703125 +
          skoXC1 *
          (-11308257477/1164153218269348144531250 +
           skoXC1 *
           (8188738173/29103830456733703613281250000 +
            skoXC1 *
            (-150906746331/23283064365386962890625000000000000 +
             skoXC1 *
             (33924772431/291038304567337036132812500000000000000 +
              skoXC1 *
              (-45622969821/29103830456733703613281250000000000000000000 +
               skoXC1 *
               (10528377651/
                727595761418342590332031250000000000000000000000 +
                skoXC1 *
                -10528377651/
                145519152283668518066406250000000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_122, refute_0_116])).

cnf(refute_0_124, plain,
    (-1255673954304/625 +
     skoXC1 *
     (38219575984128/1953125 +
      skoXC1 *
      (-397206622568448/6103515625 +
       skoXC1 *
       (287244000755712/19073486328125 +
        skoXC1 *
        (-106389138949632/59604644775390625 +
         skoXC1 *
         (26424041234976/186264514923095703125 +
          skoXC1 *
          (-4856747733684/582076609134674072265625 +
           skoXC1 *
           (5534106272541/14551915228366851806640625000 +
            skoXC1 *
            (-80061035466951/5820766091346740722656250000000000 +
             skoXC1 *
             (115734346533729/291038304567337036132812500000000000000 +
              skoXC1 *
              (-133089315419907/
               14551915228366851806640625000000000000000000 +
               skoXC1 *
               (59753140606467/
                363797880709171295166015625000000000000000000000 +
                skoXC1 *
                (-80250893551521/
                 36379788070917129516601562500000000000000000000000000 +
                 skoXC1 *
                 (1479579042249/
                  72759576141834259033203125000000000000000000000000000000 +
                  skoXC1 *
                  -14772465477/
                  145519152283668518066406250000000000000000000000000000000000))))))))))))) <=
     (-26922247520256/625 + skoCC1 * -497377345536/125 +
      skoSC1 * 2155301830656/125 +
      skoXC1 *
      (20191685640192/1953125 +
       skoXC1 *
       (-7571882115072/6103515625 +
        skoXC1 *
        (1892970528768/19073486328125 +
         skoXC1 *
         (-349386162048/59604644775390625 +
          skoXC1 *
          (49912308864/186264514923095703125 +
           skoXC1 *
           (-11308257477/1164153218269348144531250 +
            skoXC1 *
            (8188738173/29103830456733703613281250000 +
             skoXC1 *
             (-150906746331/23283064365386962890625000000000000 +
              skoXC1 *
              (33924772431/291038304567337036132812500000000000000 +
               skoXC1 *
               (-45622969821/29103830456733703613281250000000000000000000 +
                skoXC1 *
                (10528377651/
                 727595761418342590332031250000000000000000000000 +
                 skoXC1 *
                 -10528377651/
                 145519152283668518066406250000000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_0])).

cnf(refute_0_125, plain,
    (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12 |
     skoXC1 *
     (9386970027393024/1953125 + skoCC1 * 7087627173888/15625 +
      skoSC1 * -30713051086848/15625 +
      skoXC1 *
      (87411806965334016/6103515625 + skoCC1 * 2805519089664/1953125 +
       skoSC1 * -12157249388544/1953125 +
       skoXC1 *
       (-68221567900975104/19073486328125 +
        skoXC1 *
        (25917274966201344/59604644775390625 +
         skoXC1 *
         (-6523166641761792/186264514923095703125 +
          skoXC1 *
          (1209152541967128/582076609134674072265625 +
           skoXC1 *
           (-346520961295443/3637978807091712951660156250 +
            skoXC1 *
            (10072546600206321/2910383045673370361328125000000000 +
             skoXC1 *
             (-3653871930217971/36379788070917129516601562500000000000 +
              skoXC1 *
              (16861073930237097/
               7275957614183425903320312500000000000000000 +
               skoXC1 *
               (-948990678572079/
                22737367544323205947875976562500000000000000000 +
                skoXC1 *
                (10223567035348491/
                 18189894035458564758300781250000000000000000000000000 +
                 skoXC1 *
                 (-47249342079093/
                  9094947017729282379150390625000000000000000000000000000 +
                  skoXC1 *
                  1892985933267/
                  72759576141834259033203125000000000000000000000000000000000))))))))))))) <=
     -321811296288768/625 + skoCC1 * -5968528146432/125 +
     skoSC1 * 25863621967872/125),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (-321811296288768/625 + skoCC1 * -5968528146432/125 +
     skoSC1 * 25863621967872/125 <
     skoXC1 *
     (9386970027393024/1953125 + skoCC1 * 7087627173888/15625 +
      skoSC1 * -30713051086848/15625 +
      skoXC1 *
      (87411806965334016/6103515625 + skoCC1 * 2805519089664/1953125 +
       skoSC1 * -12157249388544/1953125 +
       skoXC1 *
       (-68221567900975104/19073486328125 +
        skoXC1 *
        (25917274966201344/59604644775390625 +
         skoXC1 *
         (-6523166641761792/186264514923095703125 +
          skoXC1 *
          (1209152541967128/582076609134674072265625 +
           skoXC1 *
           (-346520961295443/3637978807091712951660156250 +
            skoXC1 *
            (10072546600206321/2910383045673370361328125000000000 +
             skoXC1 *
             (-3653871930217971/36379788070917129516601562500000000000 +
              skoXC1 *
              (16861073930237097/
               7275957614183425903320312500000000000000000 +
               skoXC1 *
               (-948990678572079/
                22737367544323205947875976562500000000000000000 +
                skoXC1 *
                (10223567035348491/
                 18189894035458564758300781250000000000000000000000000 +
                 skoXC1 *
                 (-47249342079093/
                  9094947017729282379150390625000000000000000000000000000 +
                  skoXC1 *
                  1892985933267/
                  72759576141834259033203125000000000000000000000000000000000))))))))))))) |
     skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_127, plain, (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain, (-12 < skoXC1 * (57/500 + skoXC1 * 361/1000000)),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_129, plain, ($false),
    inference(resolve, [], [refute_0_127, refute_0_128])).
% SZS output end CNFRefutation for Chua-1-VC2-U.tptp
