% SZS output start CNFRefutation for Chua.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

fof('Chua', conjecture,
    (! [X] :
       ((0 <= X & X <= 2) =>
        6/5 + 279/100 * exp(-983/10000 * X) +
        117/10000 * exp(-17/6250 * X) * cos(117/5000 * X) +
        27/500 * exp(-17/6250 * X) * sin(117/5000 * X) <= 407/100))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 2) =>
        6/5 + 279/100 * exp(-983/10000 * X) +
        117/10000 * exp(-17/6250 * X) * cos(117/5000 * X) +
        27/500 * exp(-17/6250 * X) * sin(117/5000 * X) <= 407/100)),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 2) =>
          6/5 + 279/100 * exp(-983/10000 * X) +
          117/10000 * exp(-17/6250 * X) * cos(117/5000 * X) +
          27/500 * exp(-17/6250 * X) * sin(117/5000 * X) <= 407/100)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (407/100 <
        6/5 + 279/100 * exp(-983/10000 * X) +
        117/10000 * exp(-17/6250 * X) * cos(117/5000 * X) +
        27/500 * exp(-17/6250 * X) * sin(117/5000 * X) & X <= 2 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (407/100 <
     6/5 + 279/100 * exp(-983/10000 * skoXC1) +
     117/10000 * exp(-17/6250 * skoXC1) * cos(117/5000 * skoXC1) +
     27/500 * exp(-17/6250 * skoXC1) * sin(117/5000 * skoXC1) &
     0 <= skoXC1 & skoXC1 <= 2),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (407/100 <
     6/5 + 279/100 * exp(-983/10000 * skoXC1) +
     117/10000 * exp(-17/6250 * skoXC1) * cos(117/5000 * skoXC1) +
     27/500 * exp(-17/6250 * skoXC1) * sin(117/5000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 2),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) *
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) <
     370262016/25 +
     skoXC1 *
     (3147227136/78125 +
      skoXC1 *
      (13375715328/244140625 +
       skoXC1 *
       (37897860096/762939453125 +
        skoXC1 *
        (70466333616/2384185791015625 +
         skoXC1 *
         (85566262248/7450580596923828125 +
          skoXC1 * 60609435759/23283064365386962890625))))) |
     2304 +
     skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000))))) <= 0 |
     (370262016/25 +
      skoXC1 *
      (3147227136/78125 +
       skoXC1 *
       (13375715328/244140625 +
        skoXC1 *
        (37897860096/762939453125 +
         skoXC1 *
         (70466333616/2384185791015625 +
          skoXC1 *
          (85566262248/7450580596923828125 +
           skoXC1 * 60609435759/23283064365386962890625)))))) /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) <=
     4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     2304 /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) *
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) |
     160704/25 +
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) <=
     (4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (~ lgen(0, exp(X_000122), X_000123) | exp(X_000122) <= X_000123),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (0 < X_000122 |
     ~ lgen(0,
            2304 /
            (-(X_000122 ^ 3) + 6 * X_000122 ^ 2 - 24 * X_000122 + 48) ^ 2,
            X_000123) | lgen(0, exp(X_000122), X_000123)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_4, plain,
    (0 < X_000122 |
     ~ lgen(0,
            2304 /
            (-(X_000122 ^ 3) + 6 * X_000122 ^ 2 - 24 * X_000122 + 48) ^ 2,
            X_000123) | exp(X_000122) <= X_000123),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (0 < X_000122 |
     X_000123 <
     2304 /
     (2304 +
      X_000122 *
      (-2304 +
       X_000122 *
       (1152 +
        X_000122 *
        (-384 + X_000122 * (84 + X_000122 * (-12 + X_000122)))))) |
     exp(X_000122) <= X_000123), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (0 < skoXC1 * -983/10000 |
     (4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     2304 /
     (2304 +
      skoXC1 * -983/10000 *
      (-2304 +
       skoXC1 * -983/10000 *
       (1152 +
        skoXC1 * -983/10000 *
        (-384 +
         skoXC1 * -983/10000 *
         (84 + skoXC1 * -983/10000 * (-12 + skoXC1 * -983/10000)))))) |
     exp(skoXC1 * -983/10000) <=
     (4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) *
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) <
     16848/625 +
     skoXC1 *
     (227448/78125 +
      skoXC1 *
      (-14414517/1953125000 +
       skoXC1 * (skoXC1 * 65773441071/195312500000000000))) |
     2304 +
     skoXC1 *
     (19584/3125 +
      skoXC1 *
      (83232/9765625 +
       skoXC1 *
       (235824/30517578125 +
        skoXC1 *
        (1753941/381469726562500 +
         skoXC1 *
         (4259571/2384185791015625000 +
          skoXC1 * 24137569/59604644775390625000000))))) <= 0 |
     (16848/625 +
      skoXC1 *
      (227448/78125 +
       skoXC1 *
       (-14414517/1953125000 +
        skoXC1 * (skoXC1 * 65773441071/195312500000000000)))) /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) <=
     287/100 + exp(skoXC1 * -983/10000) * -279/100),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_8, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 <
     2304 /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) *
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) |
     117/10000 +
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_9, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-13689/50000000 +
        skoXC1 * (skoXC1 * 62462907/5000000000000000)))) *
     (exp(skoXC1 * -17/6250) * 117/10000) |
     exp(skoXC1 * -17/6250) * 117/10000 <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-13689/50000000 + skoXC1 * (skoXC1 * 62462907/5000000000000000))) <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_10, plain,
    (~ lgen(0, cos(X_000027), X_000028) | cos(X_000027) <= X_000028),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_11, plain,
    (~ lgen(0, 1 - X_000027 ^ 2 / 2 + X_000027 ^ 4 / 24, X_000028) |
     lgen(0, cos(X_000027), X_000028)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_12, plain,
    (~ lgen(0, 1 - X_000027 ^ 2 / 2 + X_000027 ^ 4 / 24, X_000028) |
     cos(X_000027) <= X_000028),
    inference(resolve, [], [refute_0_11, refute_0_10])).

cnf(refute_0_13, plain,
    (X_000028 <
     1 + X_000027 * (X_000027 * (-1/2 + X_000027 * (X_000027 * 1/24))) |
     cos(X_000027) <= X_000028), inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000) <
     1 +
     skoXC1 * 117/5000 *
     (skoXC1 * 117/5000 *
      (-1/2 + skoXC1 * 117/5000 * (skoXC1 * 117/5000 * 1/24))) |
     cos(skoXC1 * 117/5000) <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000)),
    inference(subst, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000) <
     skoXC1 * 117/5000 * (exp(skoXC1 * -17/6250) * 27/500) |
     exp(skoXC1 * -17/6250) * 27/500 <= 0 |
     skoXC1 * 117/5000 <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_16, plain,
    (407/100 <
     6/5 + 279/100 * exp(-983/10000 * skoXC1) +
     117/10000 * exp(-17/6250 * skoXC1) * cos(117/5000 * skoXC1) +
     27/500 * exp(-17/6250 * skoXC1) * sin(117/5000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_17, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000) <
     sin(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * 27/500)),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < sin(skoXC1 * 117/5000) |
     exp(skoXC1 * -17/6250) * 27/500 <= 0),
    inference(split, [], [refute_0_17])).

cnf(refute_0_19, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < sin(skoXC1 * 117/5000)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (~ lgen(0, sin(X_000087), X_000088) | sin(X_000087) <= X_000088),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_21, plain,
    (X_000087 < 0 | ~ lgen(0, X_000087, X_000088) |
     lgen(0, sin(X_000087), X_000088)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_22, plain,
    (X_000087 < 0 | ~ lgen(0, X_000087, X_000088) |
     sin(X_000087) <= X_000088),
    inference(resolve, [], [refute_0_21, refute_0_20])).

cnf(refute_0_23, plain,
    (X_000088 < X_000087 | X_000087 < 0 | sin(X_000087) <= X_000088),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * 117/5000 < 0 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < skoXC1 * 117/5000 |
     sin(skoXC1 * 117/5000) <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500)),
    inference(subst, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 * 117/5000 < 0 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < skoXC1 * 117/5000),
    inference(resolve, [], [refute_0_24, refute_0_19])).

cnf(refute_0_26, plain,
    (skoXC1 < 0 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < skoXC1 * 117/5000),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_28, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000)) /
     (exp(skoXC1 * -17/6250) * 27/500) < skoXC1 * 117/5000),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * -117/10000) <
     skoXC1 * 117/5000 * (exp(skoXC1 * -17/6250) * 27/500) |
     exp(skoXC1 * -17/6250) * 27/500 <= 0),
    inference(resolve, [], [refute_0_15, refute_0_28])).

cnf(refute_0_30, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000) <
     cos(skoXC1 * 117/5000) * (exp(skoXC1 * -17/6250) * 117/10000)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000) < cos(skoXC1 * 117/5000) |
     exp(skoXC1 * -17/6250) * 117/10000 <= 0),
    inference(split, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000) < cos(skoXC1 * 117/5000)),
    inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000) <
     1 +
     skoXC1 * 117/5000 *
     (skoXC1 * 117/5000 *
      (-1/2 + skoXC1 * 117/5000 * (skoXC1 * 117/5000 * 1/24)))),
    inference(resolve, [], [refute_0_14, refute_0_32])).

cnf(refute_0_34, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100 +
      exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000)) /
     (exp(skoXC1 * -17/6250) * 117/10000) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-13689/50000000 + skoXC1 * (skoXC1 * 62462907/5000000000000000)))),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 +
     exp(skoXC1 * -17/6250) * (skoXC1 * -3159/2500000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-13689/50000000 +
        skoXC1 * (skoXC1 * 62462907/5000000000000000)))) *
     (exp(skoXC1 * -17/6250) * 117/10000) |
     exp(skoXC1 * -17/6250) * 117/10000 <= 0),
    inference(resolve, [], [refute_0_9, refute_0_34])).

cnf(refute_0_36, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 <
     exp(skoXC1 * -17/6250) *
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000))))),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     exp(skoXC1 * -17/6250) |
     117/10000 +
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000))) <= 0),
    inference(split, [], [refute_0_36])).

cnf(refute_0_38, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     exp(skoXC1 * -17/6250) |
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000))) <=
     -117/10000), inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain, (skoXC1 <= 2),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_40, plain,
    (-117/10000 <
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))),
    inference(decision, [], [refute_0_39, refute_0_27])).

cnf(refute_0_41, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     exp(skoXC1 * -17/6250)),
    inference(resolve, [], [refute_0_38, refute_0_40])).

cnf(refute_0_42, plain,
    (0 < skoXC1 * -17/6250 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     2304 /
     (2304 +
      skoXC1 * -17/6250 *
      (-2304 +
       skoXC1 * -17/6250 *
       (1152 +
        skoXC1 * -17/6250 *
        (-384 +
         skoXC1 * -17/6250 *
         (84 + skoXC1 * -17/6250 * (-12 + skoXC1 * -17/6250)))))) |
     exp(skoXC1 * -17/6250) <=
     (287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_43, plain,
    (0 < skoXC1 * -17/6250 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     2304 /
     (2304 +
      skoXC1 * -17/6250 *
      (-2304 +
       skoXC1 * -17/6250 *
       (1152 +
        skoXC1 * -17/6250 *
        (-384 +
         skoXC1 * -17/6250 *
         (84 + skoXC1 * -17/6250 * (-12 + skoXC1 * -17/6250))))))),
    inference(resolve, [], [refute_0_42, refute_0_41])).

cnf(refute_0_44, plain,
    (skoXC1 < 0 |
     (287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     2304 /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000))))))),
    inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) /
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) <
     2304 /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000))))))),
    inference(resolve, [], [refute_0_27, refute_0_44])).

cnf(refute_0_46, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 <
     2304 /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) *
     (117/10000 +
      skoXC1 *
      (3159/2500000 +
       skoXC1 *
       (-1601613/500000000000 +
        skoXC1 * (skoXC1 * 7308160119/50000000000000000000)))) |
     117/10000 +
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000))) <= 0),
    inference(resolve, [], [refute_0_8, refute_0_45])).

cnf(refute_0_47, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 <
     (16848/625 +
      skoXC1 *
      (227448/78125 +
       skoXC1 *
       (-14414517/1953125000 +
        skoXC1 * (skoXC1 * 65773441071/195312500000000000)))) /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) |
     skoXC1 *
     (3159/2500000 +
      skoXC1 *
      (-1601613/500000000000 +
       skoXC1 * (skoXC1 * 7308160119/50000000000000000000))) <=
     -117/10000), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (287/100 + exp(skoXC1 * -983/10000) * -279/100 <
     (16848/625 +
      skoXC1 *
      (227448/78125 +
       skoXC1 *
       (-14414517/1953125000 +
        skoXC1 * (skoXC1 * 65773441071/195312500000000000)))) /
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000))))))),
    inference(resolve, [], [refute_0_47, refute_0_40])).

cnf(refute_0_49, plain,
    ((287/100 + exp(skoXC1 * -983/10000) * -279/100) *
     (2304 +
      skoXC1 *
      (19584/3125 +
       skoXC1 *
       (83232/9765625 +
        skoXC1 *
        (235824/30517578125 +
         skoXC1 *
         (1753941/381469726562500 +
          skoXC1 *
          (4259571/2384185791015625000 +
           skoXC1 * 24137569/59604644775390625000000)))))) <
     16848/625 +
     skoXC1 *
     (227448/78125 +
      skoXC1 *
      (-14414517/1953125000 +
       skoXC1 * (skoXC1 * 65773441071/195312500000000000))) |
     2304 +
     skoXC1 *
     (19584/3125 +
      skoXC1 *
      (83232/9765625 +
       skoXC1 *
       (235824/30517578125 +
        skoXC1 *
        (1753941/381469726562500 +
         skoXC1 *
         (4259571/2384185791015625000 +
          skoXC1 * 24137569/59604644775390625000000))))) <= 0),
    inference(resolve, [], [refute_0_7, refute_0_48])).

cnf(refute_0_50, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     exp(skoXC1 * -983/10000) *
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) |
     skoXC1 *
     (19584/3125 +
      skoXC1 *
      (83232/9765625 +
       skoXC1 *
       (235824/30517578125 +
        skoXC1 *
        (1753941/381469726562500 +
         skoXC1 *
         (4259571/2384185791015625000 +
          skoXC1 * 24137569/59604644775390625000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (-2304 <
     skoXC1 *
     (19584/3125 +
      skoXC1 *
      (83232/9765625 +
       skoXC1 *
       (235824/30517578125 +
        skoXC1 *
        (1753941/381469726562500 +
         skoXC1 *
         (4259571/2384185791015625000 +
          skoXC1 * 24137569/59604644775390625000000)))))),
    inference(decision, [], [refute_0_39, refute_0_27])).

cnf(refute_0_52, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     exp(skoXC1 * -983/10000) *
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000))))))),
    inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     exp(skoXC1 * -983/10000) |
     160704/25 +
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000))))) <= 0),
    inference(split, [], [refute_0_52])).

cnf(refute_0_54, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     exp(skoXC1 * -983/10000) |
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000))))) <=
     -160704/25), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-160704/25 <
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000)))))),
    inference(decision, [], [refute_0_39, refute_0_27])).

cnf(refute_0_56, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     exp(skoXC1 * -983/10000)),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (0 < skoXC1 * -983/10000 |
     (4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     2304 /
     (2304 +
      skoXC1 * -983/10000 *
      (-2304 +
       skoXC1 * -983/10000 *
       (1152 +
        skoXC1 * -983/10000 *
        (-384 +
         skoXC1 * -983/10000 *
         (84 + skoXC1 * -983/10000 * (-12 + skoXC1 * -983/10000))))))),
    inference(resolve, [], [refute_0_6, refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 < 0 |
     (4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) /
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000))))))),
    inference(resolve, [], [refute_0_27, refute_0_58])).

cnf(refute_0_60, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     2304 /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) *
     (160704/25 +
      skoXC1 *
      (1365984/78125 +
       skoXC1 *
       (5805432/244140625 +
        skoXC1 *
        (16448724/762939453125 +
         skoXC1 *
         (489349539/38146972656250000 +
          skoXC1 *
          (1188420309/238418579101562500000 +
           skoXC1 * 6734381751/5960464477539062500000000)))))) |
     160704/25 +
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000))))) <= 0),
    inference(resolve, [], [refute_0_1, refute_0_59])).

cnf(refute_0_61, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     (370262016/25 +
      skoXC1 *
      (3147227136/78125 +
       skoXC1 *
       (13375715328/244140625 +
        skoXC1 *
        (37897860096/762939453125 +
         skoXC1 *
         (70466333616/2384185791015625 +
          skoXC1 *
          (85566262248/7450580596923828125 +
           skoXC1 * 60609435759/23283064365386962890625)))))) /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) |
     skoXC1 *
     (1365984/78125 +
      skoXC1 *
      (5805432/244140625 +
       skoXC1 *
       (16448724/762939453125 +
        skoXC1 *
        (489349539/38146972656250000 +
         skoXC1 *
         (1188420309/238418579101562500000 +
          skoXC1 * 6734381751/5960464477539062500000000))))) <=
     -160704/25), inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (4115952/625 +
     skoXC1 *
     (1177704/78125 +
      skoXC1 *
      (12437937/390625000 +
       skoXC1 *
       (16920372/762939453125 +
        skoXC1 *
        (-1579903250199/4882812500000000000 +
         skoXC1 *
         (1222496877/238418579101562500000 +
          skoXC1 * 6927482303/5960464477539062500000000))))) <
     (370262016/25 +
      skoXC1 *
      (3147227136/78125 +
       skoXC1 *
       (13375715328/244140625 +
        skoXC1 *
        (37897860096/762939453125 +
         skoXC1 *
         (70466333616/2384185791015625 +
          skoXC1 *
          (85566262248/7450580596923828125 +
           skoXC1 * 60609435759/23283064365386962890625)))))) /
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000))))))),
    inference(resolve, [], [refute_0_61, refute_0_55])).

cnf(refute_0_63, plain,
    ((4115952/625 +
      skoXC1 *
      (1177704/78125 +
       skoXC1 *
       (12437937/390625000 +
        skoXC1 *
        (16920372/762939453125 +
         skoXC1 *
         (-1579903250199/4882812500000000000 +
          skoXC1 *
          (1222496877/238418579101562500000 +
           skoXC1 * 6927482303/5960464477539062500000000)))))) *
     (2304 +
      skoXC1 *
      (141552/625 +
       skoXC1 *
       (8696601/781250 +
        skoXC1 *
        (2849586261/7812500000 +
         skoXC1 *
         (19608003061941/2500000000000000 +
          skoXC1 *
          (2753523858555429/25000000000000000000 +
           skoXC1 * 902237984319995569/1000000000000000000000000)))))) <
     370262016/25 +
     skoXC1 *
     (3147227136/78125 +
      skoXC1 *
      (13375715328/244140625 +
       skoXC1 *
       (37897860096/762939453125 +
        skoXC1 *
        (70466333616/2384185791015625 +
         skoXC1 *
         (85566262248/7450580596923828125 +
          skoXC1 * 60609435759/23283064365386962890625))))) |
     2304 +
     skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_62])).

cnf(refute_0_64, plain,
    (226603008/625 <
     skoXC1 *
     (-580452251904/390625 +
      skoXC1 *
      (-18735462836568/244140625 +
       skoXC1 *
       (-3932294145237819/1525878906250 +
        skoXC1 *
        (-140402088616311031923/2441406250000000000 +
         skoXC1 *
         (-104413643519719049620119/122070312500000000000000 +
          skoXC1 *
          (-191803374358170253566497703/24414062500000000000000000000 +
           skoXC1 *
           (-261900051655502903039608941/15258789062500000000000000000000 +
            skoXC1 *
            (-34954910531894636890104415977/
             1220703125000000000000000000000000000 +
             skoXC1 *
             (47568699221988033551154817947/
              3051757812500000000000000000000000000000 +
              skoXC1 *
              (35566166094074988791911300466319/
               122070312500000000000000000000000000000000000 +
               skoXC1 *
               (-5667515493118095835118273961/
                1192092895507812500000000000000000000000000000 +
                skoXC1 *
                -6250237669471160793285915407/
                5960464477539062500000000000000000000000000000000))))))))))) |
     skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000))))) <=
     -2304), inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 *
     (-580452251904/390625 +
      skoXC1 *
      (-18735462836568/244140625 +
       skoXC1 *
       (-3932294145237819/1525878906250 +
        skoXC1 *
        (-140402088616311031923/2441406250000000000 +
         skoXC1 *
         (-104413643519719049620119/122070312500000000000000 +
          skoXC1 *
          (-191803374358170253566497703/24414062500000000000000000000 +
           skoXC1 *
           (-261900051655502903039608941/15258789062500000000000000000000 +
            skoXC1 *
            (-34954910531894636890104415977/
             1220703125000000000000000000000000000 +
             skoXC1 *
             (47568699221988033551154817947/
              3051757812500000000000000000000000000000 +
              skoXC1 *
              (35566166094074988791911300466319/
               122070312500000000000000000000000000000000000 +
               skoXC1 *
               (-5667515493118095835118273961/
                1192092895507812500000000000000000000000000000 +
                skoXC1 *
                -6250237669471160793285915407/
                5960464477539062500000000000000000000000000000000))))))))))) <=
     226603008/625 |
     skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000))))) <=
     -2304), inference(decision, [], [refute_0_39, refute_0_27])).

cnf(refute_0_66, plain,
    (skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000))))) <=
     -2304), inference(resolve, [], [refute_0_65, refute_0_64])).

cnf(refute_0_67, plain,
    (-2304 <
     skoXC1 *
     (141552/625 +
      skoXC1 *
      (8696601/781250 +
       skoXC1 *
       (2849586261/7812500000 +
        skoXC1 *
        (19608003061941/2500000000000000 +
         skoXC1 *
         (2753523858555429/25000000000000000000 +
          skoXC1 * 902237984319995569/1000000000000000000000000)))))),
    inference(decision, [], [refute_0_39, refute_0_27])).

cnf(refute_0_68, plain, ($false),
    inference(resolve, [], [refute_0_66, refute_0_67])).
% SZS output end CNFRefutation for Chua.tptp
