% SZS output start CNFRefutation for Chua-1-IL-U.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Chua', conjecture,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * C - 21/50 * exp(3/12500 * X) * S < 5))).

fof(subgoal_0, plain,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * C - 21/50 * exp(3/12500 * X) * S < 5)),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X, S, C] :
         ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
          151/100 - 23/1000 * exp(-19/1000 * X) -
          47/20 * exp(3/12500 * X) * C - 21/50 * exp(3/12500 * X) * S <
          5)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, S, X] :
       (X <= 289 & 0 <= X &
        5 <=
        151/100 - 23/1000 * exp(-19/1000 * X) -
        47/20 * exp(3/12500 * X) * C - 21/50 * exp(3/12500 * X) * S &
        S ^ 2 + C ^ 2 = 1)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * skoCC1 -
     21/50 * exp(3/12500 * skoXC1) * skoSC1 & skoXC1 <= 289 &
     skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * skoCC1 -
     21/50 * exp(3/12500 * skoXC1) * skoSC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 289),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     -1897100476416/25 + skoCC1 * -255483445248/5 +
     skoSC1 * -228304355328/25 +
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0 |
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     exp(skoXC1 * -19/1000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 <
     (skoCC1 * 255483445248/5 + skoSC1 * 228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * 255483445248/5 + skoSC1 * 228304355328/25 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_6, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) |
     0 <= skoCC1 * 47/20 + skoSC1 * 21/50 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_7, plain,
    (5 <=
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     47/20 * exp(3/12500 * skoXC1) * skoCC1 -
     21/50 * exp(3/12500 * skoXC1) * skoSC1),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (exp(skoXC1 * 3/12500) * (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (0 <= skoCC1 * 47/20 + skoSC1 * 21/50 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <= exp(skoXC1 * 3/12500)),
    inference(split, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (skoCC1 * -235/42 <= skoSC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <= exp(skoXC1 * 3/12500)),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (exp(X_000061) < X_000062 | ~ lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_12, plain,
    (3 < X_000061 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_13, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062)),
    inference(resolve, [], [refute_0_12, refute_0_11])).

cnf(refute_0_14, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     X_000062 <=
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061))))))))))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (3 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(subst, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (3 < skoXC1 * 3/12500 | skoCC1 * -235/42 <= skoSC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(resolve, [], [refute_0_10, refute_0_15])).

cnf(refute_0_17, plain,
    (12500 < skoXC1 | skoCC1 * -235/42 <= skoSC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (0 *
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <
     -201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000))))) |
     6624/125 +
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= 0 |
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_19, plain,
    (0 <
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) |
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <
     exp(skoXC1 * -19/1000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_20, plain,
    (-349/100 + exp(skoXC1 * -19/1000) * -23/1000 <
     (skoCC1 * 27072/5 + skoSC1 * 24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * 27072/5 + skoSC1 * 24192/25 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_21, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) |
     0 <= skoCC1 * 47/20 + skoSC1 * 21/50 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_22, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_23, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_24, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_23, refute_0_22])).

cnf(refute_0_25, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (0 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (0 < skoXC1 * 3/12500 | skoCC1 * -235/42 <= skoSC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(resolve, [], [refute_0_10, refute_0_26])).

cnf(refute_0_28, plain,
    (0 < skoXC1 | skoCC1 * -235/42 <= skoSC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_30, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_31, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_30, refute_0_29])).

cnf(refute_0_32, plain,
    (0 < -3490/23 | -3490/23 <= exp(skoXC1 * -19/1000)),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (skoCC1 * 47/20 + skoSC1 * 21/50 <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50)),
    inference(split, [], [refute_0_8])).

cnf(refute_0_34, plain,
    (exp(skoXC1 * 3/12500) <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) | skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) < exp(skoXC1 * 3/12500)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_36, plain,
    (0 <
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) | skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (0 * (skoCC1 * 47/20 + skoSC1 * 21/50) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     skoCC1 * 47/20 + skoSC1 * 21/50 <= 0 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_38, plain,
    (0 * (skoCC1 * 47/20 + skoSC1 * 21/50) <
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000 |
     skoCC1 * 47/20 + skoSC1 * 21/50 <= 0 | skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (exp(skoXC1 * -19/1000) < -3490/23 | skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain, (0 < -3490/23 | skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_32, refute_0_39])).

cnf(refute_0_41, plain, (skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (skoCC1 * 47/20 + skoSC1 * 21/50 < 0 |
     0 < skoCC1 * 47/20 + skoSC1 * 21/50 |
     0 <= -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(split, [], [refute_0_8])).

cnf(refute_0_43, plain,
    (skoSC1 < skoCC1 * -235/42 | skoCC1 * -235/42 < skoSC1 |
     exp(skoXC1 * -19/1000) <= -3490/23),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain, (0 < -3490/23 | -3490/23 < exp(skoXC1 * -19/1000)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_45, plain,
    (0 < -3490/23 | skoSC1 < skoCC1 * -235/42 | skoCC1 * -235/42 < skoSC1),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (skoSC1 < skoCC1 * -235/42 | skoCC1 * -235/42 < skoSC1),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_48, plain, (skoSC1 * skoSC1 = 1 + skoCC1 * (skoCC1 * -1)),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain, (skoXC1 <= 289),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_50, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_51, plain, (skoSC1 < skoCC1 * -235/42 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_41, refute_0_46, refute_0_48, refute_0_49,
               refute_0_50])).

cnf(refute_0_52, plain,
    (0 < skoXC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(resolve, [], [refute_0_28, refute_0_51])).

cnf(refute_0_53, plain,
    (0 < skoXC1 | 0 <= skoCC1 * 47/20 + skoSC1 * 21/50 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_52, refute_0_21])).

cnf(refute_0_54, plain,
    (0 < skoXC1 | skoCC1 * -235/42 <= skoSC1 |
     (skoCC1 * 27072/5 + skoSC1 * 24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 |
     (skoCC1 * 27072/5 + skoSC1 * 24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_54, refute_0_51])).

cnf(refute_0_56, plain,
    (0 < skoXC1 |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * 27072/5 + skoSC1 * 24192/25 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(resolve, [], [refute_0_55, refute_0_20])).

cnf(refute_0_57, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <=
     -201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000))))) |
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (-2304 <
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_41, refute_0_46, refute_0_48, refute_0_49,
               refute_0_50])).

cnf(refute_0_59, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <=
     -201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000)))))),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (0 < skoXC1 |
     6624/125 +
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= 0 |
     exp(skoXC1 * -19/1000) <=
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000))))))),
    inference(split, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125 |
     exp(skoXC1 * -19/1000) <=
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000))))))),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (-6624/125 <
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_41, refute_0_46, refute_0_48, refute_0_49,
               refute_0_50])).

cnf(refute_0_63, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) <=
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000))))))),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (0 <
     (-201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
      skoXC1 *
      (150768/78125 +
       skoXC1 *
       (-56538/244140625 +
        skoXC1 *
        (28269/1525878906250 +
         skoXC1 *
         (-593649/610351562500000000 +
          skoXC1 *
          (254421/7629394531250000000000 +
           skoXC1 * -254421/381469726562500000000000000)))))) /
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) | 0 < skoXC1),
    inference(resolve, [], [refute_0_63, refute_0_19])).

cnf(refute_0_65, plain,
    (0 *
     (6624/125 +
      skoXC1 *
      (-4968/390625 +
       skoXC1 *
       (1863/1220703125 +
        skoXC1 *
        (-1863/15258789062500 +
         skoXC1 *
         (39123/6103515625000000000 +
          skoXC1 *
          (-16767/76293945312500000000000 +
           skoXC1 * 16767/3814697265625000000000000000)))))) <
     -201024/25 + skoCC1 * -27072/5 + skoSC1 * -24192/25 +
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000))))) | 0 < skoXC1 |
     6624/125 +
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_18, refute_0_64])).

cnf(refute_0_66, plain,
    (201024/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 <
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000))))) | 0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125 |
     skoXC1 *
     (150768/78125 +
      skoXC1 *
      (-56538/244140625 +
       skoXC1 *
       (28269/1525878906250 +
        skoXC1 *
        (-593649/610351562500000000 +
         skoXC1 *
         (254421/7629394531250000000000 +
          skoXC1 * -254421/381469726562500000000000000))))) <=
     201024/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25),
    inference(decision, [],
              [refute_0_41, refute_0_46, refute_0_48, refute_0_49,
               refute_0_50])).

cnf(refute_0_68, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125),
    inference(resolve, [], [refute_0_67, refute_0_66])).

cnf(refute_0_69, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_68, refute_0_62])).

cnf(refute_0_70, plain, (skoSC1 < skoCC1 * -235/42 | 12500 < skoXC1),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_71, plain,
    (12500 < skoXC1 |
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_17, refute_0_70])).

cnf(refute_0_72, plain, (skoXC1 <= 12500),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_73, plain,
    ((-349/100 + exp(skoXC1 * -19/1000) * -23/1000) /
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_72, refute_0_71])).

cnf(refute_0_74, plain,
    (0 <= skoCC1 * 47/20 + skoSC1 * 21/50 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * 47/20 + skoSC1 * 21/50) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_73, refute_0_6])).

cnf(refute_0_75, plain,
    (skoCC1 * -235/42 <= skoSC1 |
     (skoCC1 * 255483445248/5 + skoSC1 * 228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain, (skoSC1 < skoCC1 * -235/42),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_77, plain,
    ((skoCC1 * 255483445248/5 + skoSC1 * 228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -349/100 + exp(skoXC1 * -19/1000) * -23/1000),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * 255483445248/5 + skoSC1 * 228304355328/25 <=
     (-349/100 + exp(skoXC1 * -19/1000) * -23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_77, refute_0_5])).

cnf(refute_0_79, plain,
    (exp(skoXC1 * -19/1000) *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1897100476416/25 + skoCC1 * -255483445248/5 +
     skoSC1 * -228304355328/25 +
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (-21743271936 <
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_81, plain,
    (exp(skoXC1 * -19/1000) *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1897100476416/25 + skoCC1 * -255483445248/5 +
     skoSC1 * -228304355328/25 +
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    (62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0 |
     exp(skoXC1 * -19/1000) <=
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(split, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125 |
     exp(skoXC1 * -19/1000) <=
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-62511906816/125 <
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_85, plain,
    (exp(skoXC1 * -19/1000) <=
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (0 <
     (-1897100476416/25 + skoCC1 * -255483445248/5 +
      skoSC1 * -228304355328/25 +
      skoXC1 *
      (1422825357312/78125 +
       skoXC1 *
       (-533559508992/244140625 +
        skoXC1 *
        (133389877248/762939453125 +
         skoXC1 *
         (-24619811328/2384185791015625 +
          skoXC1 *
          (3517115904/7450580596923828125 +
           skoXC1 *
           (-398423286/23283064365386962890625 +
            skoXC1 *
            (144256707/291038304567337036132812500 +
             skoXC1 *
             (-2658445029/232830643653869628906250000000000 +
              skoXC1 *
              (597634929/2910383045673370361328125000000000000 +
               skoXC1 *
               (-803715939/291038304567337036132812500000000000000000 +
                skoXC1 *
                (185472909/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -185472909/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_85, refute_0_4])).

cnf(refute_0_87, plain,
    (0 *
     (62511906816/125 +
      skoXC1 *
      (-46883930112/390625 +
       skoXC1 *
       (17581473792/1220703125 +
        skoXC1 *
        (-4395368448/3814697265625 +
         skoXC1 *
         (811254528/11920928955078125 +
          skoXC1 *
          (-115893504/37252902984619140625 +
           skoXC1 *
           (13128561/116415321826934814453125 +
            skoXC1 *
            (-9506889/2910383045673370361328125000 +
             skoXC1 *
             (175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (-39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (-12223143/
                 72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <
     -1897100476416/25 + skoCC1 * -255483445248/5 +
     skoSC1 * -228304355328/25 +
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     62511906816/125 +
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_86])).

cnf(refute_0_88, plain,
    (1897100476416/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 <
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125), inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125 |
     skoXC1 *
     (1422825357312/78125 +
      skoXC1 *
      (-533559508992/244140625 +
       skoXC1 *
       (133389877248/762939453125 +
        skoXC1 *
        (-24619811328/2384185791015625 +
         skoXC1 *
         (3517115904/7450580596923828125 +
          skoXC1 *
          (-398423286/23283064365386962890625 +
           skoXC1 *
           (144256707/291038304567337036132812500 +
            skoXC1 *
            (-2658445029/232830643653869628906250000000000 +
             skoXC1 *
             (597634929/2910383045673370361328125000000000000 +
              skoXC1 *
              (-803715939/291038304567337036132812500000000000000000 +
               skoXC1 *
               (185472909/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -185472909/
                1455191522836685180664062500000000000000000000000000))))))))))) <=
     1897100476416/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25),
    inference(decision, [],
              [refute_0_69, refute_0_41, refute_0_46, refute_0_48,
               refute_0_49, refute_0_50])).

cnf(refute_0_90, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125),
    inference(resolve, [], [refute_0_89, refute_0_88])).

cnf(refute_0_91, plain, ($false),
    inference(resolve, [], [refute_0_90, refute_0_84])).
% SZS output end CNFRefutation for Chua-1-IL-U.tptp
