% SZS output start CNFRefutation for Chua-1-VC1-L.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Chua', conjecture,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        0 <
        71/25 - 63/1000 * exp(-19/1000 * X) -
        177/100 * exp(3/12500 * X) * C +
        689/1000 * exp(3/12500 * X) * S))).

fof(subgoal_0, plain,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        0 <
        71/25 - 63/1000 * exp(-19/1000 * X) -
        177/100 * exp(3/12500 * X) * C + 689/1000 * exp(3/12500 * X) * S)),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X, S, C] :
         ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
          0 <
          71/25 - 63/1000 * exp(-19/1000 * X) -
          177/100 * exp(3/12500 * X) * C +
          689/1000 * exp(3/12500 * X) * S)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, S, X] :
       (X <= 289 & 0 <= X &
        71/25 - 63/1000 * exp(-19/1000 * X) -
        177/100 * exp(3/12500 * X) * C + 689/1000 * exp(3/12500 * X) * S <=
        0 & S ^ 2 + C ^ 2 = 1)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     71/25 - 63/1000 * exp(-19/1000 * skoXC1) -
     177/100 * exp(3/12500 * skoXC1) * skoCC1 +
     689/1000 * exp(3/12500 * skoXC1) * skoSC1 <= 0 & skoXC1 <= 289 &
     skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (71/25 - 63/1000 * exp(-19/1000 * skoXC1) -
     177/100 * exp(3/12500 * skoXC1) * skoCC1 +
     689/1000 * exp(3/12500 * skoXC1) * skoSC1 <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 289),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000))))))))))) <
     (-2054739197952/125 +
      skoXC1 *
      (62541124337664/390625 +
       skoXC1 *
       (-649974473293824/1220703125 +
        skoXC1 *
        (470035637600256/3814697265625 +
         skoXC1 *
         (-174091318281216/11920928955078125 +
          skoXC1 *
          (43239340202688/37252902984619140625 +
           skoXC1 *
           (-7947405382392/116415321826934814453125 +
            skoXC1 *
            (4527905132079/1455191522836685180664062500 +
             skoXC1 *
             (-65504483563869/582076609134674072265625000000000 +
              skoXC1 *
              (94691738073051/29103830456733703613281250000000000000 +
               skoXC1 *
               (-108891258070833/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (48888933223473/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-65659821996699/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (1210564670931/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -12086562663/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     -2054739197952/125 +
     skoXC1 *
     (62541124337664/390625 +
      skoXC1 *
      (-649974473293824/1220703125 +
       skoXC1 *
       (470035637600256/3814697265625 +
        skoXC1 *
        (-174091318281216/11920928955078125 +
         skoXC1 *
         (43239340202688/37252902984619140625 +
          skoXC1 *
          (-7947405382392/116415321826934814453125 +
           skoXC1 *
           (4527905132079/1455191522836685180664062500 +
            skoXC1 *
            (-65504483563869/582076609134674072265625000000000 +
             skoXC1 *
             (94691738073051/29103830456733703613281250000000000000 +
              skoXC1 *
              (-108891258070833/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (48888933223473/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (-65659821996699/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (1210564670931/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  -12086562663/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     0 <=
     -171228266496/125 +
     skoXC1 *
     (128421199872/390625 +
      skoXC1 *
      (-48157949952/1220703125 +
       skoXC1 *
       (12039487488/3814697265625 +
        skoXC1 *
        (-2222131968/11920928955078125 +
         skoXC1 *
         (317447424/37252902984619140625 +
          skoXC1 *
          (-35960841/116415321826934814453125 +
           skoXC1 *
           (26040609/2910383045673370361328125000 +
            skoXC1 *
            (-479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (-145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -33480783/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    (exp(X_000064) < X_000065 | ~ lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064 |
     lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_4, plain,
    (exp(X_000064) < X_000065 |
     ~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (exp(X_000064) < X_000065 |
     X_000065 <=
     (12 + X_000064 * (6 + X_000064)) / (12 + X_000064 * (-6 + X_000064)) |
     0 <= X_000064), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (exp(skoXC1 * -19/1000) <
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     0 <= skoXC1 * -19/1000 |
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-71/25 + exp(skoXC1 * -19/1000) * 63/1000 <
     (skoCC1 * -962139783168/25 + skoSC1 * 1872639295488/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * -962139783168/25 + skoSC1 * 1872639295488/125 <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     0 <= skoCC1 * -177/100 + skoSC1 * 689/1000 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_9, plain,
    (71/25 - 63/1000 * exp(-19/1000 * skoXC1) -
     177/100 * exp(3/12500 * skoXC1) * skoCC1 +
     689/1000 * exp(3/12500 * skoXC1) * skoSC1 <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    (exp(skoXC1 * 3/12500) * (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (0 <= skoCC1 * -177/100 + skoSC1 * 689/1000 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <= exp(skoXC1 * 3/12500)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (skoCC1 * 1770/689 <= skoSC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <= exp(skoXC1 * 3/12500)),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (exp(X_000061) < X_000062 | ~ lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (3 < X_000061 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_15, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     X_000062 <=
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061))))))))))))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (3 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (3 < skoXC1 * 3/12500 | skoCC1 * 1770/689 <= skoSC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(resolve, [], [refute_0_12, refute_0_17])).

cnf(refute_0_19, plain,
    (12500 < skoXC1 | skoCC1 * 1770/689 <= skoSC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000))))) <
     (-41803776/125 +
      skoXC1 *
      (31352832/390625 +
       skoXC1 *
       (-11757312/1220703125 +
        skoXC1 *
        (2939328/3814697265625 +
         skoXC1 *
         (-964467/23841857910156250 +
          skoXC1 *
          (413343/298023223876953125000 +
           skoXC1 * -413343/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     -41803776/125 +
     skoXC1 *
     (31352832/390625 +
      skoXC1 *
      (-11757312/1220703125 +
       skoXC1 *
       (2939328/3814697265625 +
        skoXC1 *
        (-964467/23841857910156250 +
         skoXC1 *
         (413343/298023223876953125000 +
          skoXC1 * -413343/14901161193847656250000000))))) <=
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_21, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) |
     0 <=
     -18144/125 +
     skoXC1 *
     (13608/390625 +
      skoXC1 *
      (-5103/1220703125 +
       skoXC1 *
       (5103/15258789062500 +
        skoXC1 *
        (-107163/6103515625000000000 +
         skoXC1 *
         (45927/76293945312500000000000 +
          skoXC1 * -45927/3814697265625000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000)))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_22, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_23, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_24, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_23, refute_0_22])).

cnf(refute_0_25, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (0 < skoXC1 * -19/1000 |
     exp(skoXC1 * -19/1000) <
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (-71/25 + exp(skoXC1 * -19/1000) * 63/1000 <
     (skoCC1 * -101952/25 + skoSC1 * 198432/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * -101952/25 + skoSC1 * 198432/125 <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_28, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     0 <= skoCC1 * -177/100 + skoSC1 * 689/1000 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_29, plain,
    (0 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_30, plain,
    (0 < skoXC1 * 3/12500 | skoCC1 * 1770/689 <= skoSC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(resolve, [], [refute_0_12, refute_0_29])).

cnf(refute_0_31, plain,
    (0 < skoXC1 | skoCC1 * 1770/689 <= skoSC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (0 < skoXC1 | 0 <= skoCC1 * -177/100 + skoSC1 * 689/1000 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000 |
     skoCC1 * 1770/689 <= skoSC1),
    inference(resolve, [], [refute_0_31, refute_0_28])).

cnf(refute_0_33, plain,
    (0 < skoXC1 | skoCC1 * 1770/689 <= skoSC1 |
     (skoCC1 * -101952/25 + skoSC1 * 198432/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2840/63 |
     2840/63 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_35, plain,
    (skoCC1 * -177/100 + skoSC1 * 689/1000 < 0 |
     0 < skoCC1 * -177/100 + skoSC1 * 689/1000 |
     0 <= -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(split, [], [refute_0_10])).

cnf(refute_0_36, plain,
    (skoSC1 < skoCC1 * 1770/689 | skoCC1 * 1770/689 < skoSC1 |
     2840/63 <= exp(skoXC1 * -19/1000)),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (exp(skoXC1 * -19/1000) < 2840/63 | 0 < skoXC1 * -19/1000 |
     2840/63 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_38, plain,
    (skoSC1 < skoCC1 * 1770/689 | 0 < skoXC1 * -19/1000 |
     skoCC1 * 1770/689 < skoSC1 |
     2840/63 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 < 0 | skoSC1 < skoCC1 * 1770/689 | skoCC1 * 1770/689 < skoSC1 |
     2840/63 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (2840/63 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2840/63),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (skoCC1 * -177/100 + skoSC1 * 689/1000 <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_42, plain,
    (exp(skoXC1 * 3/12500) <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_45, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (0 <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) < exp(skoXC1 * 3/12500)),
    inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (0 <
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_42, refute_0_46])).

cnf(refute_0_48, plain,
    (0 * (skoCC1 * -177/100 + skoSC1 * 689/1000) <
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000 |
     skoCC1 * -177/100 + skoSC1 * 689/1000 <= 0 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_49, plain,
    (0 * (skoCC1 * -177/100 + skoSC1 * 689/1000) <
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000 |
     skoCC1 * -177/100 + skoSC1 * 689/1000 <= 0 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    (2840/63 < exp(skoXC1 * -19/1000) | skoSC1 <= skoCC1 * 1770/689),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_52, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_53, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_52, refute_0_51])).

cnf(refute_0_54, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 * -19/1000 |
     2840/63 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     exp(skoXC1 * -19/1000) <= 2840/63),
    inference(subst, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (0 < skoXC1 * -19/1000 |
     2840/63 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_55, refute_0_50])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 |
     2840/63 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_59, plain, (skoSC1 * skoSC1 = 1 + skoCC1 * (skoCC1 * -1)),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain, (skoXC1 <= 289),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_61, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_62, plain, (0 <= skoXC1 | skoSC1 <= skoCC1 * 1770/689),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_63, plain,
    (2840/63 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_62, refute_0_57])).

cnf(refute_0_64, plain,
    (2840/63 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_40, refute_0_63])).

cnf(refute_0_65, plain,
    (710912/7 <
     skoXC1 *
     (-345344/175 +
      skoXC1 *
      (-410096/21875 +
       skoXC1 *
       (-973978/8203125 +
        skoXC1 *
        (-9252791/18750000000 +
         skoXC1 *
         (-175803029/131250000000000 +
          skoXC1 * -3340257551/1575000000000000000))))) |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 *
     (-345344/175 +
      skoXC1 *
      (-410096/21875 +
       skoXC1 *
       (-973978/8203125 +
        skoXC1 *
        (-9252791/18750000000 +
         skoXC1 *
         (-175803029/131250000000000 +
          skoXC1 * -3340257551/1575000000000000000))))) <= 710912/7 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_67, plain,
    (skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_66, refute_0_65])).

cnf(refute_0_68, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoSC1 <= skoCC1 * 1770/689),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_69, plain, (skoSC1 <= skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 | skoSC1 < skoCC1 * 1770/689 |
     skoSC1 <= skoCC1 * 1770/689),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 | skoSC1 < skoCC1 * 1770/689 |
     2840/63 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_70, refute_0_39])).

cnf(refute_0_72, plain, (skoSC1 < skoCC1 * 1770/689 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_73, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     2840/63 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_72, refute_0_71])).

cnf(refute_0_74, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     2840/63 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_73, refute_0_34])).

cnf(refute_0_75, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (345344/175 +
      skoXC1 *
      (410096/21875 +
       skoXC1 *
       (973978/8203125 +
        skoXC1 *
        (9252791/18750000000 +
         skoXC1 *
         (175803029/131250000000000 +
          skoXC1 * 3340257551/1575000000000000000))))) <= -710912/7),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     -2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoXC1 *
     (345344/175 +
      skoXC1 *
      (410096/21875 +
       skoXC1 *
       (973978/8203125 +
        skoXC1 *
        (9252791/18750000000 +
         skoXC1 *
         (175803029/131250000000000 +
          skoXC1 * 3340257551/1575000000000000000))))) <= -710912/7),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_77, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     skoXC1 *
     (345344/175 +
      skoXC1 *
      (410096/21875 +
       skoXC1 *
       (973978/8203125 +
        skoXC1 *
        (9252791/18750000000 +
         skoXC1 *
         (175803029/131250000000000 +
          skoXC1 * 3340257551/1575000000000000000))))) <= -710912/7),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (skoSC1 < skoCC1 * 1770/689 |
     -710912/7 <
     skoXC1 *
     (345344/175 +
      skoXC1 *
      (410096/21875 +
       skoXC1 *
       (973978/8203125 +
        skoXC1 *
        (9252791/18750000000 +
         skoXC1 *
         (175803029/131250000000000 +
          skoXC1 * 3340257551/1575000000000000000)))))),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_79, plain, (skoSC1 < skoCC1 * 1770/689),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain, (skoSC1 < skoCC1 * 1770/689 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_81, plain,
    (0 < skoXC1 |
     (skoCC1 * -101952/25 + skoSC1 * 198432/125) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(resolve, [], [refute_0_33, refute_0_80])).

cnf(refute_0_82, plain,
    (0 < skoXC1 |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * -101952/25 + skoSC1 * 198432/125 <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(resolve, [], [refute_0_81, refute_0_27])).

cnf(refute_0_83, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000))))) |
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-2304 <
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_85, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000)))))),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (0 < skoXC1 |
     0 <=
     -18144/125 +
     skoXC1 *
     (13608/390625 +
      skoXC1 *
      (-5103/1220703125 +
       skoXC1 *
       (5103/15258789062500 +
        skoXC1 *
        (-107163/6103515625000000000 +
         skoXC1 *
         (45927/76293945312500000000000 +
          skoXC1 * -45927/3814697265625000000000000000))))) |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-13608/390625 +
      skoXC1 *
      (5103/1220703125 +
       skoXC1 *
       (-5103/15258789062500 +
        skoXC1 *
        (107163/6103515625000000000 +
         skoXC1 *
         (-45927/76293945312500000000000 +
          skoXC1 * 45927/3814697265625000000000000000))))) <= -18144/125 |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (-18144/125 <
     skoXC1 *
     (-13608/390625 +
      skoXC1 *
      (5103/1220703125 +
       skoXC1 *
       (-5103/15258789062500 +
        skoXC1 *
        (107163/6103515625000000000 +
         skoXC1 *
         (-45927/76293945312500000000000 +
          skoXC1 * 45927/3814697265625000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_89, plain,
    (0 < skoXC1 |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (0 < skoXC1 * -19/1000 | 0 < skoXC1 |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_89, refute_0_26])).

cnf(refute_0_91, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain, (0 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_93, plain,
    (0 < skoXC1 |
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) /
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (0 < skoXC1 |
     0 <=
     -18144/125 +
     skoXC1 *
     (13608/390625 +
      skoXC1 *
      (-5103/1220703125 +
       skoXC1 *
       (5103/15258789062500 +
        skoXC1 *
        (-107163/6103515625000000000 +
         skoXC1 *
         (45927/76293945312500000000000 +
          skoXC1 * -45927/3814697265625000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-18144/125 +
      skoXC1 *
      (13608/390625 +
       skoXC1 *
       (-5103/1220703125 +
        skoXC1 *
        (5103/15258789062500 +
         skoXC1 *
         (-107163/6103515625000000000 +
          skoXC1 *
          (45927/76293945312500000000000 +
           skoXC1 * -45927/3814697265625000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000)))))),
    inference(resolve, [], [refute_0_93, refute_0_21])).

cnf(refute_0_95, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-13608/390625 +
      skoXC1 *
      (5103/1220703125 +
       skoXC1 *
       (-5103/15258789062500 +
        skoXC1 *
        (107163/6103515625000000000 +
         skoXC1 *
         (-45927/76293945312500000000000 +
          skoXC1 * 45927/3814697265625000000000000000))))) <= -18144/125 |
     (-41803776/125 +
      skoXC1 *
      (31352832/390625 +
       skoXC1 *
       (-11757312/1220703125 +
        skoXC1 *
        (2939328/3814697265625 +
         skoXC1 *
         (-964467/23841857910156250 +
          skoXC1 *
          (413343/298023223876953125000 +
           skoXC1 * -413343/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000)))))),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 < skoXC1 |
     (-41803776/125 +
      skoXC1 *
      (31352832/390625 +
       skoXC1 *
       (-11757312/1220703125 +
        skoXC1 *
        (2939328/3814697265625 +
         skoXC1 *
         (-964467/23841857910156250 +
          skoXC1 *
          (413343/298023223876953125000 +
           skoXC1 * -413343/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
     skoXC1 *
     (122688/78125 +
      skoXC1 *
      (-46008/244140625 +
       skoXC1 *
       (11502/762939453125 +
        skoXC1 *
        (-120771/152587890625000000 +
         skoXC1 *
         (51759/1907348632812500000000 +
          skoXC1 * -51759/95367431640625000000000000)))))),
    inference(resolve, [], [refute_0_95, refute_0_88])).

cnf(refute_0_97, plain,
    (0 < skoXC1 |
     -41803776/125 +
     skoXC1 *
     (31352832/390625 +
      skoXC1 *
      (-11757312/1220703125 +
       skoXC1 *
       (2939328/3814697265625 +
        skoXC1 *
        (-964467/23841857910156250 +
         skoXC1 *
         (413343/298023223876953125000 +
          skoXC1 * -413343/14901161193847656250000000))))) <=
     (-163584/25 + skoCC1 * 101952/25 + skoSC1 * -198432/125 +
      skoXC1 *
      (122688/78125 +
       skoXC1 *
       (-46008/244140625 +
        skoXC1 *
        (11502/762939453125 +
         skoXC1 *
         (-120771/152587890625000000 +
          skoXC1 *
          (51759/1907348632812500000000 +
           skoXC1 * -51759/95367431640625000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_20])).

cnf(refute_0_98, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (110509443072/390625 + skoCC1 * -557881344/3125 +
      skoSC1 * 1085819904/15625 +
      skoXC1 *
      (3238377262848/1220703125 + skoCC1 * -662484096/390625 +
       skoSC1 * 1289411136/1953125 +
       skoXC1 *
       (63283521708288/3814697265625 + skoCC1 * -524466576/48828125 +
        skoSC1 * 1020783816/244140625 +
        skoXC1 *
        (3222056123088561/47683715820312500 +
         skoCC1 * -4359628413/97656250000 +
         skoSC1 * 16970530941/976562500000 +
         skoXC1 *
         (211862343985343037/1192092895507812500000 +
          skoCC1 * -11833277121/97656250000000 +
          skoSC1 * 46062869697/976562500000000 +
          skoXC1 *
          (31375875162589707051/119209289550781250000000000 +
           skoCC1 * -74944088433/390625000000000000 +
           skoSC1 * 291731508081/3906250000000000000 +
           skoXC1 *
           (-815924416768113177/11920928955078125000000000000 +
            skoXC1 *
            (200900514788909037/23841857910156250000000000000000 +
             skoXC1 *
             (-16356400533753399/23841857910156250000000000000000000 +
              skoXC1 *
              (3474781667931951/95367431640625000000000000000000000000 +
               skoXC1 *
               (-60107231418129/47683715820312500000000000000000000000000 +
                skoXC1 *
                2435047754679/
                95367431640625000000000000000000000000000000))))))))))) <=
     -1842683904/125 + skoCC1 * 234897408/25 + skoSC1 * -457187328/125),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (-1842683904/125 + skoCC1 * 234897408/25 + skoSC1 * -457187328/125 <
     skoXC1 *
     (110509443072/390625 + skoCC1 * -557881344/3125 +
      skoSC1 * 1085819904/15625 +
      skoXC1 *
      (3238377262848/1220703125 + skoCC1 * -662484096/390625 +
       skoSC1 * 1289411136/1953125 +
       skoXC1 *
       (63283521708288/3814697265625 + skoCC1 * -524466576/48828125 +
        skoSC1 * 1020783816/244140625 +
        skoXC1 *
        (3222056123088561/47683715820312500 +
         skoCC1 * -4359628413/97656250000 +
         skoSC1 * 16970530941/976562500000 +
         skoXC1 *
         (211862343985343037/1192092895507812500000 +
          skoCC1 * -11833277121/97656250000000 +
          skoSC1 * 46062869697/976562500000000 +
          skoXC1 *
          (31375875162589707051/119209289550781250000000000 +
           skoCC1 * -74944088433/390625000000000000 +
           skoSC1 * 291731508081/3906250000000000000 +
           skoXC1 *
           (-815924416768113177/11920928955078125000000000000 +
            skoXC1 *
            (200900514788909037/23841857910156250000000000000000 +
             skoXC1 *
             (-16356400533753399/23841857910156250000000000000000000 +
              skoXC1 *
              (3474781667931951/95367431640625000000000000000000000000 +
               skoXC1 *
               (-60107231418129/47683715820312500000000000000000000000000 +
                skoXC1 *
                2435047754679/
                95367431640625000000000000000000000000000000))))))))))) |
     0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_100, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_102, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain, (skoSC1 < skoCC1 * 1770/689 | 12500 < skoXC1),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_104, plain,
    (12500 < skoXC1 |
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_19, refute_0_103])).

cnf(refute_0_105, plain, (skoXC1 <= 12500),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_106, plain,
    ((-71/25 + exp(skoXC1 * -19/1000) * 63/1000) /
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    (0 <= skoCC1 * -177/100 + skoSC1 * 689/1000 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * -177/100 + skoSC1 * 689/1000) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(resolve, [], [refute_0_106, refute_0_8])).

cnf(refute_0_108, plain,
    (skoCC1 * 1770/689 <= skoSC1 |
     (skoCC1 * -962139783168/25 + skoSC1 * 1872639295488/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    ((skoCC1 * -962139783168/25 + skoSC1 * 1872639295488/125) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -71/25 + exp(skoXC1 * -19/1000) * 63/1000),
    inference(resolve, [], [refute_0_108, refute_0_79])).

cnf(refute_0_110, plain,
    (21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * -962139783168/25 + skoSC1 * 1872639295488/125 <=
     (-71/25 + exp(skoXC1 * -19/1000) * 63/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_109, refute_0_7])).

cnf(refute_0_111, plain,
    (exp(skoXC1 * -19/1000) *
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000))))))))))) |
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-21743271936 <
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_113, plain,
    (exp(skoXC1 * -19/1000) *
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (0 <=
     -171228266496/125 +
     skoXC1 *
     (128421199872/390625 +
      skoXC1 *
      (-48157949952/1220703125 +
       skoXC1 *
       (12039487488/3814697265625 +
        skoXC1 *
        (-2222131968/11920928955078125 +
         skoXC1 *
         (317447424/37252902984619140625 +
          skoXC1 *
          (-35960841/116415321826934814453125 +
           skoXC1 *
           (26040609/2910383045673370361328125000 +
            skoXC1 *
            (-479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (-145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -33480783/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (-128421199872/390625 +
      skoXC1 *
      (48157949952/1220703125 +
       skoXC1 *
       (-12039487488/3814697265625 +
        skoXC1 *
        (2222131968/11920928955078125 +
         skoXC1 *
         (-317447424/37252902984619140625 +
          skoXC1 *
          (35960841/116415321826934814453125 +
           skoXC1 *
           (-26040609/2910383045673370361328125000 +
            skoXC1 *
            (479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (-107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                33480783/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -171228266496/125 |
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (-171228266496/125 <
     skoXC1 *
     (-128421199872/390625 +
      skoXC1 *
      (48157949952/1220703125 +
       skoXC1 *
       (-12039487488/3814697265625 +
        skoXC1 *
        (2222131968/11920928955078125 +
         skoXC1 *
         (-317447424/37252902984619140625 +
          skoXC1 *
          (35960841/116415321826934814453125 +
           skoXC1 *
           (-26040609/2910383045673370361328125000 +
            skoXC1 *
            (479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (-107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                33480783/
                14551915228366851806640625000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_117, plain,
    ((-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (0 <= skoXC1 * -19/1000 |
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(resolve, [], [refute_0_117, refute_0_6])).

cnf(refute_0_119, plain,
    ((-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) /
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000))),
    inference(resolve, [], [refute_0_119, refute_0_102])).

cnf(refute_0_121, plain,
    (0 <=
     -171228266496/125 +
     skoXC1 *
     (128421199872/390625 +
      skoXC1 *
      (-48157949952/1220703125 +
       skoXC1 *
       (12039487488/3814697265625 +
        skoXC1 *
        (-2222131968/11920928955078125 +
         skoXC1 *
         (317447424/37252902984619140625 +
          skoXC1 *
          (-35960841/116415321826934814453125 +
           skoXC1 *
           (26040609/2910383045673370361328125000 +
            skoXC1 *
            (-479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (-145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -33480783/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-171228266496/125 +
      skoXC1 *
      (128421199872/390625 +
       skoXC1 *
       (-48157949952/1220703125 +
        skoXC1 *
        (12039487488/3814697265625 +
         skoXC1 *
         (-2222131968/11920928955078125 +
          skoXC1 *
          (317447424/37252902984619140625 +
           skoXC1 *
           (-35960841/116415321826934814453125 +
            skoXC1 *
            (26040609/2910383045673370361328125000 +
             skoXC1 *
             (-479891223/2328306436538696289062500000000000 +
              skoXC1 *
              (107882523/29103830456733703613281250000000000000 +
               skoXC1 *
               (-145083393/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (33480783/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -33480783/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_120, refute_0_1])).

cnf(refute_0_122, plain,
    (skoXC1 *
     (-128421199872/390625 +
      skoXC1 *
      (48157949952/1220703125 +
       skoXC1 *
       (-12039487488/3814697265625 +
        skoXC1 *
        (2222131968/11920928955078125 +
         skoXC1 *
         (-317447424/37252902984619140625 +
          skoXC1 *
          (35960841/116415321826934814453125 +
           skoXC1 *
           (-26040609/2910383045673370361328125000 +
            skoXC1 *
            (479891223/2328306436538696289062500000000000 +
             skoXC1 *
             (-107882523/29103830456733703613281250000000000000 +
              skoXC1 *
              (145083393/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-33480783/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                33480783/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -171228266496/125 |
     (-2054739197952/125 +
      skoXC1 *
      (62541124337664/390625 +
       skoXC1 *
       (-649974473293824/1220703125 +
        skoXC1 *
        (470035637600256/3814697265625 +
         skoXC1 *
         (-174091318281216/11920928955078125 +
          skoXC1 *
          (43239340202688/37252902984619140625 +
           skoXC1 *
           (-7947405382392/116415321826934814453125 +
            skoXC1 *
            (4527905132079/1455191522836685180664062500 +
             skoXC1 *
             (-65504483563869/582076609134674072265625000000000 +
              skoXC1 *
              (94691738073051/29103830456733703613281250000000000000 +
               skoXC1 *
               (-108891258070833/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (48888933223473/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-65659821996699/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (1210564670931/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -12086562663/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    ((-2054739197952/125 +
      skoXC1 *
      (62541124337664/390625 +
       skoXC1 *
       (-649974473293824/1220703125 +
        skoXC1 *
        (470035637600256/3814697265625 +
         skoXC1 *
         (-174091318281216/11920928955078125 +
          skoXC1 *
          (43239340202688/37252902984619140625 +
           skoXC1 *
           (-7947405382392/116415321826934814453125 +
            skoXC1 *
            (4527905132079/1455191522836685180664062500 +
             skoXC1 *
             (-65504483563869/582076609134674072265625000000000 +
              skoXC1 *
              (94691738073051/29103830456733703613281250000000000000 +
               skoXC1 *
               (-108891258070833/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (48888933223473/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-65659821996699/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (1210564670931/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -12086562663/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -1543772307456/25 + skoCC1 * 962139783168/25 +
     skoSC1 * -1872639295488/125 +
     skoXC1 *
     (1157829230592/78125 +
      skoXC1 *
      (-434185961472/244140625 +
       skoXC1 *
       (108546490368/762939453125 +
        skoXC1 *
        (-20034459648/2384185791015625 +
         skoXC1 *
         (2862065664/7450580596923828125 +
          skoXC1 *
          (-324218376/23283064365386962890625 +
           skoXC1 *
           (29347353/72759576141834259033203125 +
            skoXC1 *
            (-540829791/58207660913467407226562500000000 +
             skoXC1 *
             (121581891/727595761418342590332031250000000000 +
              skoXC1 *
              (-163506681/72759576141834259033203125000000000000000 +
               skoXC1 *
               (37732311/1818989403545856475830078125000000000000000000 +
                skoXC1 *
                -37732311/
                363797880709171295166015625000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_122, refute_0_116])).

cnf(refute_0_124, plain,
    (-2054739197952/125 +
     skoXC1 *
     (62541124337664/390625 +
      skoXC1 *
      (-649974473293824/1220703125 +
       skoXC1 *
       (470035637600256/3814697265625 +
        skoXC1 *
        (-174091318281216/11920928955078125 +
         skoXC1 *
         (43239340202688/37252902984619140625 +
          skoXC1 *
          (-7947405382392/116415321826934814453125 +
           skoXC1 *
           (4527905132079/1455191522836685180664062500 +
            skoXC1 *
            (-65504483563869/582076609134674072265625000000000 +
             skoXC1 *
             (94691738073051/29103830456733703613281250000000000000 +
              skoXC1 *
              (-108891258070833/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (48888933223473/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (-65659821996699/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (1210564670931/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  -12086562663/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     (-1543772307456/25 + skoCC1 * 962139783168/25 +
      skoSC1 * -1872639295488/125 +
      skoXC1 *
      (1157829230592/78125 +
       skoXC1 *
       (-434185961472/244140625 +
        skoXC1 *
        (108546490368/762939453125 +
         skoXC1 *
         (-20034459648/2384185791015625 +
          skoXC1 *
          (2862065664/7450580596923828125 +
           skoXC1 *
           (-324218376/23283064365386962890625 +
            skoXC1 *
            (29347353/72759576141834259033203125 +
             skoXC1 *
             (-540829791/58207660913467407226562500000000 +
              skoXC1 *
              (121581891/727595761418342590332031250000000000 +
               skoXC1 *
               (-163506681/72759576141834259033203125000000000000000 +
                skoXC1 *
                (37732311/1818989403545856475830078125000000000000000000 +
                 skoXC1 *
                 -37732311/
                 363797880709171295166015625000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_0])).

cnf(refute_0_125, plain,
    (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12 |
     skoXC1 *
     (2742915793158144/390625 + skoCC1 * -13710491910144/3125 +
      skoSC1 * 26685109960704/15625 +
      skoXC1 *
      (24525695466602496/1220703125 + skoCC1 * -5427069714432/390625 +
       skoSC1 * 10562856026112/1953125 +
       skoXC1 *
       (-19172084982349824/3814697265625 +
        skoXC1 *
        (7287137903729664/11920928955078125 +
         skoXC1 *
         (-1834589900089152/37252902984619140625 +
          skoXC1 *
          (340120474356168/116415321826934814453125 +
           skoXC1 *
           (-194967071503641/1455191522836685180664062500 +
            skoXC1 *
            (2833866473475051/582076609134674072265625000000000 +
             skoXC1 *
             (-4112297096347629/29103830456733703613281250000000000000 +
              skoXC1 *
              (4744416176159607/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (-2136353374264167/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (2877037331022621/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (-53188862743989/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  532768008177/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     -90571599249408/125 + skoCC1 * 11545677398016/25 +
     skoSC1 * -22471671545856/125),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (-90571599249408/125 + skoCC1 * 11545677398016/25 +
     skoSC1 * -22471671545856/125 <
     skoXC1 *
     (2742915793158144/390625 + skoCC1 * -13710491910144/3125 +
      skoSC1 * 26685109960704/15625 +
      skoXC1 *
      (24525695466602496/1220703125 + skoCC1 * -5427069714432/390625 +
       skoSC1 * 10562856026112/1953125 +
       skoXC1 *
       (-19172084982349824/3814697265625 +
        skoXC1 *
        (7287137903729664/11920928955078125 +
         skoXC1 *
         (-1834589900089152/37252902984619140625 +
          skoXC1 *
          (340120474356168/116415321826934814453125 +
           skoXC1 *
           (-194967071503641/1455191522836685180664062500 +
            skoXC1 *
            (2833866473475051/582076609134674072265625000000000 +
             skoXC1 *
             (-4112297096347629/29103830456733703613281250000000000000 +
              skoXC1 *
              (4744416176159607/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (-2136353374264167/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (2877037331022621/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (-53188862743989/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  532768008177/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) |
     skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_127, plain, (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain, (-12 < skoXC1 * (57/500 + skoXC1 * 361/1000000)),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_129, plain, ($false),
    inference(resolve, [], [refute_0_127, refute_0_128])).
% SZS output end CNFRefutation for Chua-1-VC1-L.tptp
