% SZS output start CNFRefutation for Chua-2-IL-L.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

fof('Chua', conjecture,
    (! [X, S, C] :
       ((0 <= X & X <= 75 & S ^ 2 + C ^ 2 = 1) =>
        -2 <
        39/1000 * exp(277/10000 * X) + 54/625 * exp(-23/4000 * X) * C -
        2025130727/1000000000 * exp(-23/4000 * X) * S))).

fof(subgoal_0, plain,
    (! [X, S, C] :
       ((0 <= X & X <= 75 & S ^ 2 + C ^ 2 = 1) =>
        -2 <
        39/1000 * exp(277/10000 * X) + 54/625 * exp(-23/4000 * X) * C -
        2025130727/1000000000 * exp(-23/4000 * X) * S)),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X, S, C] :
         ((0 <= X & X <= 75 & S ^ 2 + C ^ 2 = 1) =>
          -2 <
          39/1000 * exp(277/10000 * X) + 54/625 * exp(-23/4000 * X) * C -
          2025130727/1000000000 * exp(-23/4000 * X) * S)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, S, X] :
       (X <= 75 & 0 <= X &
        39/1000 * exp(277/10000 * X) + 54/625 * exp(-23/4000 * X) * C -
        2025130727/1000000000 * exp(-23/4000 * X) * S <= -2 &
        S ^ 2 + C ^ 2 = 1)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     39/1000 * exp(277/10000 * skoXC1) +
     54/625 * exp(-23/4000 * skoXC1) * skoCC1 -
     2025130727/1000000000 * exp(-23/4000 * skoXC1) * skoSC1 <= -2 &
     skoXC1 <= 75 & skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (39/1000 * exp(277/10000 * skoXC1) +
     54/625 * exp(-23/4000 * skoXC1) * skoCC1 -
     2025130727/1000000000 * exp(-23/4000 * skoXC1) * skoSC1 <= -2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 75),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (2304 <
     (-1565952/13 +
      skoXC1 *
      (-6146184/8125 +
       skoXC1 *
       (-19161093/8125000 +
        skoXC1 *
        (-158331287/32500000000 +
         skoXC1 *
         (-28383019429/4160000000000000 +
          skoXC1 *
          (-107942229407/16640000000000000000 +
           skoXC1 *
           (-3177989410651/798720000000000000000000 +
            skoXC1 * -41005941253/40960000000000000000000000))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) |
     0 <= skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 |
     2304 * (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     -1565952/13 +
     skoXC1 *
     (-6146184/8125 +
      skoXC1 *
      (-19161093/8125000 +
       skoXC1 *
       (-158331287/32500000000 +
        skoXC1 *
        (-28383019429/4160000000000000 +
         skoXC1 *
         (-107942229407/16640000000000000000 +
          skoXC1 *
          (-3177989410651/798720000000000000000000 +
           skoXC1 * -41005941253/40960000000000000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) |
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) *
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (39/1000 * exp(277/10000 * skoXC1) +
     54/625 * exp(-23/4000 * skoXC1) * skoCC1 -
     2025130727/1000000000 * exp(-23/4000 * skoXC1) * skoSC1 <= -2),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (exp(skoXC1 * 277/10000) <=
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (X_000048 < exp(X_000047) | ~ lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_5, plain,
    (X_000047 < -1 | ~ lgen(1, X_000048, 1 + X_000047) |
     lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_6, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) |
     ~ lgen(1, X_000048, 1 + X_000047)),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) | 1 + X_000047 <= X_000048),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (skoXC1 * 277/10000 < -1 |
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     exp(skoXC1 * 277/10000) |
     1 + skoXC1 * 277/10000 <=
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (skoXC1 * 277/10000 < -1 |
     1 + skoXC1 * 277/10000 <=
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(resolve, [], [refute_0_3, refute_0_8])).

cnf(refute_0_10, plain,
    (skoXC1 < -10000/277 |
     exp(skoXC1 * -23/4000) *
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     -2039/39 + skoXC1 * -277/10000),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_12, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_13, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_12, refute_0_11])).

cnf(refute_0_14, plain,
    (0 <
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) |
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <
     exp(skoXC1 * -23/4000)), inference(subst, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (X_000017 < exp(X_000016) | ~ lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_16, plain,
    (~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3) |
     lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_17, plain,
    (X_000017 < exp(X_000016) |
     ~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3)),
    inference(resolve, [], [refute_0_16, refute_0_15])).

cnf(refute_0_18, plain,
    (X_000017 < exp(X_000016) |
     1 +
     X_000016 *
     (1 +
      X_000016 *
      (1/2 +
       X_000016 *
       (1/6 +
        X_000016 *
        (1/24 +
         X_000016 *
         (1/120 +
          X_000016 *
          (121/87480 +
           X_000016 *
           (17/87480 +
            X_000016 *
            (49/2099520 +
             X_000016 *
             (409/170061120 +
              X_000016 *
              (361/1700611200 +
               X_000016 *
               (1/62985600 +
                X_000016 *
                (181/183666009600 +
                 X_000016 *
                 (1/20407334400 +
                  X_000016 *
                  (1/550998028800 +
                   X_000016 * 1/24794911296000)))))))))))))) <= X_000017),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (-2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     exp(skoXC1 * 277/10000) |
     1 +
     skoXC1 * 277/10000 *
     (1 +
      skoXC1 * 277/10000 *
      (1/2 +
       skoXC1 * 277/10000 *
       (1/6 +
        skoXC1 * 277/10000 *
        (1/24 +
         skoXC1 * 277/10000 *
         (1/120 +
          skoXC1 * 277/10000 *
          (121/87480 +
           skoXC1 * 277/10000 *
           (17/87480 +
            skoXC1 * 277/10000 *
            (49/2099520 +
             skoXC1 * 277/10000 *
             (409/170061120 +
              skoXC1 * 277/10000 *
              (361/1700611200 +
               skoXC1 * 277/10000 *
               (1/62985600 +
                skoXC1 * 277/10000 *
                (181/183666009600 +
                 skoXC1 * 277/10000 *
                 (1/20407334400 +
                  skoXC1 * 277/10000 *
                  (1/550998028800 +
                   skoXC1 * 277/10000 * 1/24794911296000)))))))))))))) <=
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(subst, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (1 +
     skoXC1 * 277/10000 *
     (1 +
      skoXC1 * 277/10000 *
      (1/2 +
       skoXC1 * 277/10000 *
       (1/6 +
        skoXC1 * 277/10000 *
        (1/24 +
         skoXC1 * 277/10000 *
         (1/120 +
          skoXC1 * 277/10000 *
          (121/87480 +
           skoXC1 * 277/10000 *
           (17/87480 +
            skoXC1 * 277/10000 *
            (49/2099520 +
             skoXC1 * 277/10000 *
             (409/170061120 +
              skoXC1 * 277/10000 *
              (361/1700611200 +
               skoXC1 * 277/10000 *
               (1/62985600 +
                skoXC1 * 277/10000 *
                (181/183666009600 +
                 skoXC1 * 277/10000 *
                 (1/20407334400 +
                  skoXC1 * 277/10000 *
                  (1/550998028800 +
                   skoXC1 * 277/10000 * 1/24794911296000)))))))))))))) <=
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(resolve, [], [refute_0_3, refute_0_19])).

cnf(refute_0_21, plain,
    (exp(skoXC1 * -23/4000) *
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     -2039/39 +
     skoXC1 *
     (-277/10000 +
      skoXC1 *
      (-76729/200000000 +
       skoXC1 *
       (-21253933/6000000000000 +
        skoXC1 *
        (-5887339441/240000000000000000 +
         skoXC1 *
         (-1630793025157/12000000000000000000000 +
          skoXC1 *
          (-54659289824187169/87480000000000000000000000000 +
           skoXC1 *
           (-2127195006463614701/874800000000000000000000000000000 +
            skoXC1 *
            (-1698377518984155431569/
             209952000000000000000000000000000000000 +
             skoXC1 *
             (-3926822127719835128749933/
              170061120000000000000000000000000000000000000 +
              skoXC1 *
              (-960074406615159788189748289/
               17006112000000000000000000000000000000000000000000 +
               skoXC1 *
               (-736677591779499338860277773/
                6298560000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-36934804419048758352437746704901/
                 183666009600000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-56524534939649204771410253244517/
                  204073344000000000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-15657296178282829721680640148731209/
                   55099802880000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   -4337071041384343832905537321198544893/
                   24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 <= 0 |
     exp(skoXC1 * -23/4000) <=
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000)),
    inference(split, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoCC1 * 86400000/2025130727 <= skoSC1 |
     exp(skoXC1 * -23/4000) <=
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000)),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (0 <
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) |
     skoCC1 * 86400000/2025130727 <= skoSC1),
    inference(resolve, [], [refute_0_23, refute_0_14])).

cnf(refute_0_25, plain,
    (0 * (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <
     -2039/39 +
     skoXC1 *
     (-277/10000 +
      skoXC1 *
      (-76729/200000000 +
       skoXC1 *
       (-21253933/6000000000000 +
        skoXC1 *
        (-5887339441/240000000000000000 +
         skoXC1 *
         (-1630793025157/12000000000000000000000 +
          skoXC1 *
          (-54659289824187169/87480000000000000000000000000 +
           skoXC1 *
           (-2127195006463614701/874800000000000000000000000000000 +
            skoXC1 *
            (-1698377518984155431569/
             209952000000000000000000000000000000000 +
             skoXC1 *
             (-3926822127719835128749933/
              170061120000000000000000000000000000000000000 +
              skoXC1 *
              (-960074406615159788189748289/
               17006112000000000000000000000000000000000000000000 +
               skoXC1 *
               (-736677591779499338860277773/
                6298560000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-36934804419048758352437746704901/
                 183666009600000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-56524534939649204771410253244517/
                  204073344000000000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-15657296178282829721680640148731209/
                   55099802880000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   -4337071041384343832905537321198544893/
                   24794911296000000000000000000000000000000000000000000000000000000000000000)))))))))))))) |
     skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 <= 0 |
     (-2039/39 +
      skoXC1 *
      (-277/10000 +
       skoXC1 *
       (-76729/200000000 +
        skoXC1 *
        (-21253933/6000000000000 +
         skoXC1 *
         (-5887339441/240000000000000000 +
          skoXC1 *
          (-1630793025157/12000000000000000000000 +
           skoXC1 *
           (-54659289824187169/87480000000000000000000000000 +
            skoXC1 *
            (-2127195006463614701/874800000000000000000000000000000 +
             skoXC1 *
             (-1698377518984155431569/
              209952000000000000000000000000000000000 +
              skoXC1 *
              (-3926822127719835128749933/
               170061120000000000000000000000000000000000000 +
               skoXC1 *
               (-960074406615159788189748289/
                17006112000000000000000000000000000000000000000000 +
                skoXC1 *
                (-736677591779499338860277773/
                 6298560000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-36934804419048758352437746704901/
                  183666009600000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-56524534939649204771410253244517/
                   204073344000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (-15657296178282829721680640148731209/
                    55099802880000000000000000000000000000000000000000000000000000000000 +
                    skoXC1 *
                    -4337071041384343832905537321198544893/
                    24794911296000000000000000000000000000000000000000000000000000000000000000))))))))))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_26, plain,
    (0 * (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <
     -2039/39 +
     skoXC1 *
     (-277/10000 +
      skoXC1 *
      (-76729/200000000 +
       skoXC1 *
       (-21253933/6000000000000 +
        skoXC1 *
        (-5887339441/240000000000000000 +
         skoXC1 *
         (-1630793025157/12000000000000000000000 +
          skoXC1 *
          (-54659289824187169/87480000000000000000000000000 +
           skoXC1 *
           (-2127195006463614701/874800000000000000000000000000000 +
            skoXC1 *
            (-1698377518984155431569/
             209952000000000000000000000000000000000 +
             skoXC1 *
             (-3926822127719835128749933/
              170061120000000000000000000000000000000000000 +
              skoXC1 *
              (-960074406615159788189748289/
               17006112000000000000000000000000000000000000000000 +
               skoXC1 *
               (-736677591779499338860277773/
                6298560000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-36934804419048758352437746704901/
                 183666009600000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-56524534939649204771410253244517/
                  204073344000000000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-15657296178282829721680640148731209/
                   55099802880000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   -4337071041384343832905537321198544893/
                   24794911296000000000000000000000000000000000000000000000000000000000000000)))))))))))))) |
     skoCC1 * 86400000/2025130727 <= skoSC1 |
     skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 <= 0),
    inference(resolve, [], [refute_0_25, refute_0_24])).

cnf(refute_0_27, plain,
    (2039/39 <
     skoXC1 *
     (-277/10000 +
      skoXC1 *
      (-76729/200000000 +
       skoXC1 *
       (-21253933/6000000000000 +
        skoXC1 *
        (-5887339441/240000000000000000 +
         skoXC1 *
         (-1630793025157/12000000000000000000000 +
          skoXC1 *
          (-54659289824187169/87480000000000000000000000000 +
           skoXC1 *
           (-2127195006463614701/874800000000000000000000000000000 +
            skoXC1 *
            (-1698377518984155431569/
             209952000000000000000000000000000000000 +
             skoXC1 *
             (-3926822127719835128749933/
              170061120000000000000000000000000000000000000 +
              skoXC1 *
              (-960074406615159788189748289/
               17006112000000000000000000000000000000000000000000 +
               skoXC1 *
               (-736677591779499338860277773/
                6298560000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-36934804419048758352437746704901/
                 183666009600000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-56524534939649204771410253244517/
                  204073344000000000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-15657296178282829721680640148731209/
                   55099802880000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   -4337071041384343832905537321198544893/
                   24794911296000000000000000000000000000000000000000000000000000000000000000)))))))))))))) |
     skoCC1 * 86400000/2025130727 <= skoSC1),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (0 <
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) |
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     exp(skoXC1 * 277/10000)), inference(subst, [], [refute_0_13])).

cnf(refute_0_29, plain,
    (0 <
     -2000/39 +
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(resolve, [], [refute_0_3, refute_0_28])).

cnf(refute_0_30, plain,
    (2000/39 <
     exp(skoXC1 * -23/4000) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (2000/39 < 0 | skoCC1 * -144/65 + skoSC1 * 2025130727/39000000 < 0 |
     0 < skoCC1 * -144/65 + skoSC1 * 2025130727/39000000),
    inference(split, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (skoSC1 < skoCC1 * 86400000/2025130727 |
     skoCC1 * 86400000/2025130727 < skoSC1),
    inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_34, plain, (skoSC1 * skoSC1 = 1 + skoCC1 * (skoCC1 * -1)),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain, (skoXC1 <= 75),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_36, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_37, plain,
    (skoCC1 * 86400000/2025130727 <= skoSC1 |
     skoXC1 *
     (-277/10000 +
      skoXC1 *
      (-76729/200000000 +
       skoXC1 *
       (-21253933/6000000000000 +
        skoXC1 *
        (-5887339441/240000000000000000 +
         skoXC1 *
         (-1630793025157/12000000000000000000000 +
          skoXC1 *
          (-54659289824187169/87480000000000000000000000000 +
           skoXC1 *
           (-2127195006463614701/874800000000000000000000000000000 +
            skoXC1 *
            (-1698377518984155431569/
             209952000000000000000000000000000000000 +
             skoXC1 *
             (-3926822127719835128749933/
              170061120000000000000000000000000000000000000 +
              skoXC1 *
              (-960074406615159788189748289/
               17006112000000000000000000000000000000000000000000 +
               skoXC1 *
               (-736677591779499338860277773/
                6298560000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-36934804419048758352437746704901/
                 183666009600000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (-56524534939649204771410253244517/
                  204073344000000000000000000000000000000000000000000000000000000 +
                  skoXC1 *
                  (-15657296178282829721680640148731209/
                   55099802880000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   -4337071041384343832905537321198544893/
                   24794911296000000000000000000000000000000000000000000000000000000000000000)))))))))))))) <=
     2039/39),
    inference(decision, [],
              [refute_0_32, refute_0_34, refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (skoCC1 * 86400000/2025130727 <= skoSC1),
    inference(resolve, [], [refute_0_37, refute_0_27])).

cnf(refute_0_39, plain, (-10000/277 <= skoXC1),
    inference(decision, [],
              [refute_0_38, refute_0_32, refute_0_34, refute_0_35,
               refute_0_36])).

cnf(refute_0_40, plain,
    (exp(skoXC1 * -23/4000) *
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     -2039/39 + skoXC1 * -277/10000),
    inference(resolve, [], [refute_0_39, refute_0_10])).

cnf(refute_0_41, plain,
    (0 <= skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 |
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     exp(skoXC1 * -23/4000)), inference(split, [], [refute_0_40])).

cnf(refute_0_42, plain,
    ((-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     exp(skoXC1 * -23/4000) | skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain, (skoCC1 * 86400000/2025130727 < skoSC1),
    inference(resolve, [], [refute_0_38, refute_0_32])).

cnf(refute_0_44, plain,
    ((-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     exp(skoXC1 * -23/4000)),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_46, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_47, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_46, refute_0_45])).

cnf(refute_0_48, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (0 < skoXC1 * -23/4000 |
     exp(skoXC1 * -23/4000) <
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) |
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     2304 /
     (2304 +
      skoXC1 * -23/4000 *
      (-2304 +
       skoXC1 * -23/4000 *
       (1152 +
        skoXC1 * -23/4000 *
        (-384 +
         skoXC1 * -23/4000 *
         (84 + skoXC1 * -23/4000 * (-12 + skoXC1 * -23/4000))))))),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 < skoXC1 * -23/4000 |
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     2304 /
     (2304 +
      skoXC1 * -23/4000 *
      (-2304 +
       skoXC1 * -23/4000 *
       (1152 +
        skoXC1 * -23/4000 *
        (-384 +
         skoXC1 * -23/4000 *
         (84 + skoXC1 * -23/4000 * (-12 + skoXC1 * -23/4000))))))),
    inference(resolve, [], [refute_0_44, refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 < 0 |
     (-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000))))))),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    ((-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000))))))),
    inference(resolve, [], [refute_0_36, refute_0_51])).

cnf(refute_0_53, plain,
    ((-2039/39 + skoXC1 * -277/10000) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) *
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_52, refute_0_1])).

cnf(refute_0_54, plain,
    (skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= -2304 |
     (-1565952/13 +
      skoXC1 *
      (-6146184/8125 +
       skoXC1 *
       (-19161093/8125000 +
        skoXC1 *
        (-158331287/32500000000 +
         skoXC1 *
         (-28383019429/4160000000000000 +
          skoXC1 *
          (-107942229407/16640000000000000000 +
           skoXC1 *
           (-3177989410651/798720000000000000000000 +
            skoXC1 * -41005941253/40960000000000000000000000))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <= 2304),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (2000/39 *
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <
     skoCC1 * -331776/65 + skoSC1 * 24301568724/203125 |
     2304 +
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= 0 |
     (skoCC1 * -331776/65 + skoSC1 * 24301568724/203125) /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <= 2000/39),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_56, plain,
    (2000/39 <
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) |
     skoCC1 * -144/65 + skoSC1 * 2025130727/39000000 <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <=
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_57, plain,
    (2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     exp(skoXC1 * -23/4000) |
     skoCC1 * -144/65 + skoSC1 * 2025130727/39000000 <= 0),
    inference(split, [], [refute_0_30])).

cnf(refute_0_58, plain,
    (2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     exp(skoXC1 * -23/4000) | skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_60, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_61, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (0 < skoXC1 * -23/4000 |
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     2304 /
     (2304 +
      skoXC1 * -23/4000 *
      (-2304 +
       skoXC1 * -23/4000 *
       (1152 +
        skoXC1 * -23/4000 *
        (-384 +
         skoXC1 * -23/4000 *
         (84 + skoXC1 * -23/4000 * (-12 + skoXC1 * -23/4000)))))) |
     exp(skoXC1 * -23/4000) <=
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000)),
    inference(subst, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (0 < skoXC1 * -23/4000 |
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     2304 /
     (2304 +
      skoXC1 * -23/4000 *
      (-2304 +
       skoXC1 * -23/4000 *
       (1152 +
        skoXC1 * -23/4000 *
        (-384 +
         skoXC1 * -23/4000 *
         (84 + skoXC1 * -23/4000 * (-12 + skoXC1 * -23/4000)))))) |
     skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(resolve, [], [refute_0_63, refute_0_58])).

cnf(refute_0_65, plain,
    (0 < skoXC1 * -23/4000 |
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     2304 /
     (2304 +
      skoXC1 * -23/4000 *
      (-2304 +
       skoXC1 * -23/4000 *
       (1152 +
        skoXC1 * -23/4000 *
        (-384 +
         skoXC1 * -23/4000 *
         (84 + skoXC1 * -23/4000 * (-12 + skoXC1 * -23/4000))))))),
    inference(resolve, [], [refute_0_64, refute_0_43])).

cnf(refute_0_66, plain,
    (skoXC1 < 0 |
     2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000))))))),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (2000/39 / (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) <
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000))))))),
    inference(resolve, [], [refute_0_36, refute_0_66])).

cnf(refute_0_68, plain,
    (2000/39 <
     2304 /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) *
     (skoCC1 * -144/65 + skoSC1 * 2025130727/39000000) |
     skoCC1 * -144/65 + skoSC1 * 2025130727/39000000 <= 0),
    inference(resolve, [], [refute_0_56, refute_0_67])).

cnf(refute_0_69, plain,
    (2000/39 <
     (skoCC1 * -331776/65 + skoSC1 * 24301568724/203125) /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) |
     skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (2000/39 <
     (skoCC1 * -331776/65 + skoSC1 * 24301568724/203125) /
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000))))))),
    inference(resolve, [], [refute_0_69, refute_0_43])).

cnf(refute_0_71, plain,
    (2000/39 *
     (2304 +
      skoXC1 *
      (1656/125 +
       skoXC1 *
       (4761/125000 +
        skoXC1 *
        (36501/500000000 +
         skoXC1 *
         (5876661/64000000000000 +
          skoXC1 *
          (19309029/256000000000000000 +
           skoXC1 * 148035889/4096000000000000000000)))))) <
     skoCC1 * -331776/65 + skoSC1 * 24301568724/203125 |
     2304 +
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_55, refute_0_70])).

cnf(refute_0_72, plain,
    (1536000/13 + skoCC1 * 331776/65 + skoSC1 * -24301568724/203125 <
     skoXC1 *
     (-8832/13 +
      skoXC1 *
      (-3174/1625 +
       skoXC1 *
       (-12167/3250000 +
        skoXC1 *
        (-1958887/416000000000 +
         skoXC1 *
         (-6436343/1664000000000000 +
          skoXC1 * -148035889/79872000000000000000))))) |
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (1536000/13 + skoCC1 * 331776/65 + skoSC1 * -24301568724/203125 <
     skoXC1 *
     (-8832/13 +
      skoXC1 *
      (-3174/1625 +
       skoXC1 *
       (-12167/3250000 +
        skoXC1 *
        (-1958887/416000000000 +
         skoXC1 *
         (-6436343/1664000000000000 +
          skoXC1 * -148035889/79872000000000000000))))) |
     -2304 <
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000)))))),
    inference(decision, [],
              [refute_0_38, refute_0_32, refute_0_34, refute_0_35,
               refute_0_36])).

cnf(refute_0_74, plain,
    (1536000/13 + skoCC1 * 331776/65 + skoSC1 * -24301568724/203125 <
     skoXC1 *
     (-8832/13 +
      skoXC1 *
      (-3174/1625 +
       skoXC1 *
       (-12167/3250000 +
        skoXC1 *
        (-1958887/416000000000 +
         skoXC1 *
         (-6436343/1664000000000000 +
          skoXC1 * -148035889/79872000000000000000)))))),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain,
    (-2304 <
     skoXC1 *
     (1656/125 +
      skoXC1 *
      (4761/125000 +
       skoXC1 *
       (36501/500000000 +
        skoXC1 *
        (5876661/64000000000000 +
         skoXC1 *
         (19309029/256000000000000000 +
          skoXC1 * 148035889/4096000000000000000000)))))),
    inference(decision, [],
              [refute_0_74, refute_0_38, refute_0_32, refute_0_34,
               refute_0_35, refute_0_36])).

cnf(refute_0_76, plain,
    ((-1565952/13 +
      skoXC1 *
      (-6146184/8125 +
       skoXC1 *
       (-19161093/8125000 +
        skoXC1 *
        (-158331287/32500000000 +
         skoXC1 *
         (-28383019429/4160000000000000 +
          skoXC1 *
          (-107942229407/16640000000000000000 +
           skoXC1 *
           (-3177989410651/798720000000000000000000 +
            skoXC1 * -41005941253/40960000000000000000000000))))))) /
     (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <= 2304),
    inference(resolve, [], [refute_0_54, refute_0_75])).

cnf(refute_0_77, plain,
    (0 <= skoCC1 * 144/65 + skoSC1 * -2025130727/39000000 |
     2304 * (skoCC1 * 144/65 + skoSC1 * -2025130727/39000000) <=
     -1565952/13 +
     skoXC1 *
     (-6146184/8125 +
      skoXC1 *
      (-19161093/8125000 +
       skoXC1 *
       (-158331287/32500000000 +
        skoXC1 *
        (-28383019429/4160000000000000 +
         skoXC1 *
         (-107942229407/16640000000000000000 +
          skoXC1 *
          (-3177989410651/798720000000000000000000 +
           skoXC1 * -41005941253/40960000000000000000000000))))))),
    inference(resolve, [], [refute_0_76, refute_0_0])).

cnf(refute_0_78, plain,
    (skoXC1 *
     (6146184/8125 +
      skoXC1 *
      (19161093/8125000 +
       skoXC1 *
       (158331287/32500000000 +
        skoXC1 *
        (28383019429/4160000000000000 +
         skoXC1 *
         (107942229407/16640000000000000000 +
          skoXC1 *
          (3177989410651/798720000000000000000000 +
           skoXC1 * 41005941253/40960000000000000000000000)))))) <=
     -1565952/13 + skoCC1 * -331776/65 + skoSC1 * 24301568724/203125 |
     skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (-1565952/13 + skoCC1 * -331776/65 + skoSC1 * 24301568724/203125 <
     skoXC1 *
     (6146184/8125 +
      skoXC1 *
      (19161093/8125000 +
       skoXC1 *
       (158331287/32500000000 +
        skoXC1 *
        (28383019429/4160000000000000 +
         skoXC1 *
         (107942229407/16640000000000000000 +
          skoXC1 *
          (3177989410651/798720000000000000000000 +
           skoXC1 * 41005941253/40960000000000000000000000)))))) |
     skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(decision, [],
              [refute_0_74, refute_0_38, refute_0_32, refute_0_34,
               refute_0_35, refute_0_36])).

cnf(refute_0_80, plain, (skoSC1 <= skoCC1 * 86400000/2025130727),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain, ($false),
    inference(resolve, [], [refute_0_80, refute_0_43])).
% SZS output end CNFRefutation for Chua-2-IL-L.tptp
