% SZS output start CNFRefutation for Chua-1-IL-L.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Chua', conjecture,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        -2 <
        151/100 - 23/1000 * exp(-19/1000 * X) -
        (47/20 * C + 21/50 * S) * exp(3/12500 * X)))).

fof(subgoal_0, plain,
    (! [X, S, C] :
       ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
        -2 <
        151/100 - 23/1000 * exp(-19/1000 * X) -
        (47/20 * C + 21/50 * S) * exp(3/12500 * X))),
    inference(strip, [], ['Chua'])).

fof(negate_0_0, plain,
    (~ ! [X, S, C] :
         ((0 <= X & X <= 289 & S ^ 2 + C ^ 2 = 1) =>
          -2 <
          151/100 - 23/1000 * exp(-19/1000 * X) -
          (47/20 * C + 21/50 * S) * exp(3/12500 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, S, X] :
       (X <= 289 & 0 <= X &
        151/100 - 23/1000 * exp(-19/1000 * X) -
        (47/20 * C + 21/50 * S) * exp(3/12500 * X) <= -2 &
        S ^ 2 + C ^ 2 = 1)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     (47/20 * skoCC1 + 21/50 * skoSC1) * exp(3/12500 * skoXC1) <= -2 &
     skoXC1 <= 289 & skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     (47/20 * skoCC1 + 21/50 * skoSC1) * exp(3/12500 * skoXC1) <= -2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 289),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000))))))))))) <
     (-750142881792/125 +
      skoXC1 *
      (22832473964544/390625 +
       skoXC1 *
       (-237292268027904/1220703125 +
        skoXC1 *
        (171600312139776/3814697265625 +
         skoXC1 *
         (-63557147943936/11920928955078125 +
          skoXC1 *
          (15785790867648/37252902984619140625 +
           skoXC1 *
           (-2901433711032/116415321826934814453125 +
            skoXC1 *
            (1653044730759/1455191522836685180664062500 +
             skoXC1 *
             (-23914335269349/582076609134674072265625000000000 +
              skoXC1 *
              (34569999613971/29103830456733703613281250000000000000 +
               skoXC1 *
               (-39753951359193/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (17848340700633/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-23971046125779/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (441952181451/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -4412554623/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     -750142881792/125 +
     skoXC1 *
     (22832473964544/390625 +
      skoXC1 *
      (-237292268027904/1220703125 +
       skoXC1 *
       (171600312139776/3814697265625 +
        skoXC1 *
        (-63557147943936/11920928955078125 +
         skoXC1 *
         (15785790867648/37252902984619140625 +
          skoXC1 *
          (-2901433711032/116415321826934814453125 +
           skoXC1 *
           (1653044730759/1455191522836685180664062500 +
            skoXC1 *
            (-23914335269349/582076609134674072265625000000000 +
             skoXC1 *
             (34569999613971/29103830456733703613281250000000000000 +
              skoXC1 *
              (-39753951359193/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (17848340700633/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (-23971046125779/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (441952181451/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  -4412554623/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     0 <=
     -62511906816/125 +
     skoXC1 *
     (46883930112/390625 +
      skoXC1 *
      (-17581473792/1220703125 +
       skoXC1 *
       (4395368448/3814697265625 +
        skoXC1 *
        (-811254528/11920928955078125 +
         skoXC1 *
         (115893504/37252902984619140625 +
          skoXC1 *
          (-13128561/116415321826934814453125 +
           skoXC1 *
           (9506889/2910383045673370361328125000 +
            skoXC1 *
            (-175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (-52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    (exp(X_000064) < X_000065 | ~ lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064 |
     lgen(1, exp(X_000064), X_000065)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_4, plain,
    (exp(X_000064) < X_000065 |
     ~ lgen(1,
            (X_000064 ^ 2 + 6 * X_000064 + 12) /
            (X_000064 ^ 2 - 6 * X_000064 + 12), X_000065) | 0 <= X_000064),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (exp(X_000064) < X_000065 |
     X_000065 <=
     (12 + X_000064 * (6 + X_000064)) / (12 + X_000064 * (-6 + X_000064)) |
     0 <= X_000064), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (exp(skoXC1 * -19/1000) <
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) |
     0 <= skoXC1 * -19/1000 |
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-351/100 + exp(skoXC1 * -19/1000) * 23/1000 <
     (skoCC1 * -255483445248/5 + skoSC1 * -228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * -255483445248/5 + skoSC1 * -228304355328/25 <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_8, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) |
     0 <= skoCC1 * -47/20 + skoSC1 * -21/50 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_9, plain,
    (151/100 - 23/1000 * exp(-19/1000 * skoXC1) -
     (47/20 * skoCC1 + 21/50 * skoSC1) * exp(3/12500 * skoXC1) <= -2),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    (exp(skoXC1 * 3/12500) * (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (0 <= skoCC1 * -47/20 + skoSC1 * -21/50 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <= exp(skoXC1 * 3/12500)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_12, plain,
    ((-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <= exp(skoXC1 * 3/12500) |
     skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (exp(X_000061) < X_000062 | ~ lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (3 < X_000061 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_15, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     X_000062 <=
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061))))))))))))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (3 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500))))))))))))),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (3 < skoXC1 * 3/12500 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 3/12500 *
      (-21743271936 +
       skoXC1 * 3/12500 *
       (10871635968 +
        skoXC1 * 3/12500 *
        (-3623878656 +
         skoXC1 * 3/12500 *
         (891813888 +
          skoXC1 * 3/12500 *
          (-169869312 +
           skoXC1 * 3/12500 *
           (25657344 +
            skoXC1 * 3/12500 *
            (-3096576 +
             skoXC1 * 3/12500 *
             (297216 +
              skoXC1 * 3/12500 *
              (-22272 +
               skoXC1 * 3/12500 *
               (1248 +
                skoXC1 * 3/12500 * (-48 + skoXC1 * 3/12500)))))))))))) |
     skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_12, refute_0_17])).

cnf(refute_0_19, plain,
    (12500 < skoXC1 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) |
     skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000))))) <
     (-15261696/125 +
      skoXC1 *
      (11446272/390625 +
       skoXC1 *
       (-4292352/1220703125 +
        skoXC1 *
        (1073088/3814697265625 +
         skoXC1 *
         (-352107/23841857910156250 +
          skoXC1 *
          (150903/298023223876953125000 +
           skoXC1 * -150903/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     -15261696/125 +
     skoXC1 *
     (11446272/390625 +
      skoXC1 *
      (-4292352/1220703125 +
       skoXC1 *
       (1073088/3814697265625 +
        skoXC1 *
        (-352107/23841857910156250 +
         skoXC1 *
         (150903/298023223876953125000 +
          skoXC1 * -150903/14901161193847656250000000))))) <=
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_21, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) |
     0 <=
     -6624/125 +
     skoXC1 *
     (4968/390625 +
      skoXC1 *
      (-1863/1220703125 +
       skoXC1 *
       (1863/15258789062500 +
        skoXC1 *
        (-39123/6103515625000000000 +
         skoXC1 *
         (16767/76293945312500000000000 +
          skoXC1 * -16767/3814697265625000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000)))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_22, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_23, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_24, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_23, refute_0_22])).

cnf(refute_0_25, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (0 < skoXC1 * -19/1000 |
     exp(skoXC1 * -19/1000) <
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (-351/100 + exp(skoXC1 * -19/1000) * 23/1000 <
     (skoCC1 * -27072/5 + skoSC1 * -24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * -27072/5 + skoSC1 * -24192/25 <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_28, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) |
     0 <= skoCC1 * -47/20 + skoSC1 * -21/50 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_29, plain,
    (0 < skoXC1 * 3/12500 |
     exp(skoXC1 * 3/12500) <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_30, plain,
    (0 < skoXC1 * 3/12500 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     2304 /
     (2304 +
      skoXC1 * 3/12500 *
      (-2304 +
       skoXC1 * 3/12500 *
       (1152 +
        skoXC1 * 3/12500 *
        (-384 +
         skoXC1 * 3/12500 *
         (84 + skoXC1 * 3/12500 * (-12 + skoXC1 * 3/12500)))))) |
     skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_12, refute_0_29])).

cnf(refute_0_31, plain,
    (0 < skoXC1 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) |
     skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (0 < skoXC1 | 0 <= skoCC1 * -47/20 + skoSC1 * -21/50 |
     2304 /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) *
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000 |
     skoSC1 <= skoCC1 * -235/42),
    inference(resolve, [], [refute_0_31, refute_0_28])).

cnf(refute_0_33, plain,
    (0 < skoXC1 |
     (skoCC1 * -27072/5 + skoSC1 * -24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000 |
     skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 3510/23 |
     3510/23 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_35, plain,
    (skoCC1 * -47/20 + skoSC1 * -21/50 < 0 |
     0 < skoCC1 * -47/20 + skoSC1 * -21/50 |
     0 <= -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(split, [], [refute_0_10])).

cnf(refute_0_36, plain,
    (skoSC1 < skoCC1 * -235/42 | skoCC1 * -235/42 < skoSC1 |
     3510/23 <= exp(skoXC1 * -19/1000)),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (exp(skoXC1 * -19/1000) < 3510/23 | 0 < skoXC1 * -19/1000 |
     3510/23 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_38, plain,
    (skoSC1 < skoCC1 * -235/42 | 0 < skoXC1 * -19/1000 |
     skoCC1 * -235/42 < skoSC1 |
     3510/23 <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 < 0 | skoSC1 < skoCC1 * -235/42 | skoCC1 * -235/42 < skoSC1 |
     3510/23 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (3510/23 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 3510/23),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (skoCC1 * -47/20 + skoSC1 * -21/50 <= 0 |
     exp(skoXC1 * 3/12500) <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_42, plain,
    (skoCC1 * -235/42 <= skoSC1 |
     exp(skoXC1 * 3/12500) <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_45, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (0 <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) < exp(skoXC1 * 3/12500)),
    inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (0 <
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) | skoCC1 * -235/42 <= skoSC1),
    inference(resolve, [], [refute_0_42, refute_0_46])).

cnf(refute_0_48, plain,
    (0 * (skoCC1 * -47/20 + skoSC1 * -21/50) <
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000 |
     skoCC1 * -47/20 + skoSC1 * -21/50 <= 0 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_49, plain,
    (0 * (skoCC1 * -47/20 + skoSC1 * -21/50) <
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000 |
     skoCC1 * -235/42 <= skoSC1 | skoCC1 * -47/20 + skoSC1 * -21/50 <= 0),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    (3510/23 < exp(skoXC1 * -19/1000) | skoCC1 * -235/42 <= skoSC1),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_52, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_53, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_52, refute_0_51])).

cnf(refute_0_54, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 * -19/1000 |
     3510/23 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     exp(skoXC1 * -19/1000) <= 3510/23),
    inference(subst, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (0 < skoXC1 * -19/1000 |
     3510/23 <
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000)))))) |
     skoCC1 * -235/42 <= skoSC1),
    inference(resolve, [], [refute_0_55, refute_0_50])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 |
     3510/23 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoCC1 * -235/42 <= skoSC1),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain, (skoSC1 ^ 2 + skoCC1 ^ 2 = 1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_59, plain, (skoSC1 * skoSC1 = 1 + skoCC1 * (skoCC1 * -1)),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain, (skoXC1 <= 289),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_61, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_62, plain, (0 <= skoXC1 | skoCC1 * -235/42 <= skoSC1),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_63, plain,
    (3510/23 <
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     skoCC1 * -235/42 <= skoSC1),
    inference(resolve, [], [refute_0_62, refute_0_57])).

cnf(refute_0_64, plain,
    (3510/23 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) < 2304 |
     skoCC1 * -235/42 <= skoSC1 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_40, refute_0_63])).

cnf(refute_0_65, plain,
    (8034048/23 <
     skoXC1 *
     (-3841344/575 +
      skoXC1 *
      (-4561596/71875 +
       skoXC1 *
       (-7222527/17968750 +
        skoXC1 *
        (-960596091/575000000000 +
         skoXC1 *
         (-2607332247/575000000000000 +
          skoXC1 * -16513104231/2300000000000000000))))) |
     skoCC1 * -235/42 <= skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoCC1 * -235/42 <= skoSC1 |
     skoXC1 *
     (-3841344/575 +
      skoXC1 *
      (-4561596/71875 +
       skoXC1 *
       (-7222527/17968750 +
        skoXC1 *
        (-960596091/575000000000 +
         skoXC1 *
         (-2607332247/575000000000000 +
          skoXC1 * -16513104231/2300000000000000000))))) <= 8034048/23 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_67, plain,
    (skoCC1 * -235/42 <= skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_66, refute_0_65])).

cnf(refute_0_68, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoCC1 * -235/42 <= skoSC1),
    inference(decision, [], [refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_69, plain, (skoCC1 * -235/42 <= skoSC1),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 | skoCC1 * -235/42 < skoSC1 | skoCC1 * -235/42 <= skoSC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 | skoCC1 * -235/42 < skoSC1 |
     3510/23 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_70, refute_0_39])).

cnf(refute_0_72, plain, (skoCC1 * -235/42 < skoSC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_73, plain,
    (skoCC1 * -235/42 < skoSC1 |
     3510/23 <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_72, refute_0_71])).

cnf(refute_0_74, plain,
    (skoCC1 * -235/42 < skoSC1 |
     3510/23 *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_73, refute_0_34])).

cnf(refute_0_75, plain,
    (skoCC1 * -235/42 < skoSC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (3841344/575 +
      skoXC1 *
      (4561596/71875 +
       skoXC1 *
       (7222527/17968750 +
        skoXC1 *
        (960596091/575000000000 +
         skoXC1 *
         (2607332247/575000000000000 +
          skoXC1 * 16513104231/2300000000000000000))))) <= -8034048/23),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) |
     skoCC1 * -235/42 < skoSC1 |
     skoXC1 *
     (3841344/575 +
      skoXC1 *
      (4561596/71875 +
       skoXC1 *
       (7222527/17968750 +
        skoXC1 *
        (960596091/575000000000 +
         skoXC1 *
         (2607332247/575000000000000 +
          skoXC1 * 16513104231/2300000000000000000))))) <= -8034048/23),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_77, plain,
    (skoCC1 * -235/42 < skoSC1 |
     skoXC1 *
     (3841344/575 +
      skoXC1 *
      (4561596/71875 +
       skoXC1 *
       (7222527/17968750 +
        skoXC1 *
        (960596091/575000000000 +
         skoXC1 *
         (2607332247/575000000000000 +
          skoXC1 * 16513104231/2300000000000000000))))) <= -8034048/23),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (-8034048/23 <
     skoXC1 *
     (3841344/575 +
      skoXC1 *
      (4561596/71875 +
       skoXC1 *
       (7222527/17968750 +
        skoXC1 *
        (960596091/575000000000 +
         skoXC1 *
         (2607332247/575000000000000 +
          skoXC1 * 16513104231/2300000000000000000))))) |
     skoCC1 * -235/42 < skoSC1),
    inference(decision, [],
              [refute_0_69, refute_0_59, refute_0_60, refute_0_61])).

cnf(refute_0_79, plain, (skoCC1 * -235/42 < skoSC1),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain, (skoCC1 * -235/42 < skoSC1 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_81, plain,
    (0 < skoXC1 |
     (skoCC1 * -27072/5 + skoSC1 * -24192/25) /
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000)))))) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(resolve, [], [refute_0_33, refute_0_80])).

cnf(refute_0_82, plain,
    (0 < skoXC1 |
     2304 +
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= 0 |
     skoCC1 * -27072/5 + skoSC1 * -24192/25 <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) *
     (2304 +
      skoXC1 *
      (-1728/3125 +
       skoXC1 *
       (648/9765625 +
        skoXC1 *
        (-162/30517578125 +
         skoXC1 *
         (1701/6103515625000000 +
          skoXC1 *
          (-729/76293945312500000000 +
           skoXC1 * 729/3814697265625000000000000))))))),
    inference(resolve, [], [refute_0_81, refute_0_27])).

cnf(refute_0_83, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000))))) |
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-2304 <
     skoXC1 *
     (-1728/3125 +
      skoXC1 *
      (648/9765625 +
       skoXC1 *
       (-162/30517578125 +
        skoXC1 *
        (1701/6103515625000000 +
         skoXC1 *
         (-729/76293945312500000000 +
          skoXC1 * 729/3814697265625000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_85, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -19/1000) *
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000)))))),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (0 < skoXC1 |
     0 <=
     -6624/125 +
     skoXC1 *
     (4968/390625 +
      skoXC1 *
      (-1863/1220703125 +
       skoXC1 *
       (1863/15258789062500 +
        skoXC1 *
        (-39123/6103515625000000000 +
         skoXC1 *
         (16767/76293945312500000000000 +
          skoXC1 * -16767/3814697265625000000000000000))))) |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125 |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (-6624/125 <
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_89, plain,
    (0 < skoXC1 |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (0 < skoXC1 * -19/1000 | 0 < skoXC1 |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -19/1000 *
      (-2304 +
       skoXC1 * -19/1000 *
       (1152 +
        skoXC1 * -19/1000 *
        (-384 +
         skoXC1 * -19/1000 *
         (84 + skoXC1 * -19/1000 * (-12 + skoXC1 * -19/1000))))))),
    inference(resolve, [], [refute_0_89, refute_0_26])).

cnf(refute_0_91, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain, (0 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_93, plain,
    (0 < skoXC1 |
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) /
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000))))))),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (0 < skoXC1 |
     0 <=
     -6624/125 +
     skoXC1 *
     (4968/390625 +
      skoXC1 *
      (-1863/1220703125 +
       skoXC1 *
       (1863/15258789062500 +
        skoXC1 *
        (-39123/6103515625000000000 +
         skoXC1 *
         (16767/76293945312500000000000 +
          skoXC1 * -16767/3814697265625000000000000000))))) |
     2304 /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) *
     (-6624/125 +
      skoXC1 *
      (4968/390625 +
       skoXC1 *
       (-1863/1220703125 +
        skoXC1 *
        (1863/15258789062500 +
         skoXC1 *
         (-39123/6103515625000000000 +
          skoXC1 *
          (16767/76293945312500000000000 +
           skoXC1 * -16767/3814697265625000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000)))))),
    inference(resolve, [], [refute_0_93, refute_0_21])).

cnf(refute_0_95, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-4968/390625 +
      skoXC1 *
      (1863/1220703125 +
       skoXC1 *
       (-1863/15258789062500 +
        skoXC1 *
        (39123/6103515625000000000 +
         skoXC1 *
         (-16767/76293945312500000000000 +
          skoXC1 * 16767/3814697265625000000000000000))))) <= -6624/125 |
     (-15261696/125 +
      skoXC1 *
      (11446272/390625 +
       skoXC1 *
       (-4292352/1220703125 +
        skoXC1 *
        (1073088/3814697265625 +
         skoXC1 *
         (-352107/23841857910156250 +
          skoXC1 *
          (150903/298023223876953125000 +
           skoXC1 * -150903/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000)))))),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 < skoXC1 |
     (-15261696/125 +
      skoXC1 *
      (11446272/390625 +
       skoXC1 *
       (-4292352/1220703125 +
        skoXC1 *
        (1073088/3814697265625 +
         skoXC1 *
         (-352107/23841857910156250 +
          skoXC1 *
          (150903/298023223876953125000 +
           skoXC1 * -150903/14901161193847656250000000)))))) /
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) <=
     -202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
     skoXC1 *
     (151632/78125 +
      skoXC1 *
      (-56862/244140625 +
       skoXC1 *
       (28431/1525878906250 +
        skoXC1 *
        (-597051/610351562500000000 +
         skoXC1 *
         (255879/7629394531250000000000 +
          skoXC1 * -255879/381469726562500000000000000)))))),
    inference(resolve, [], [refute_0_95, refute_0_88])).

cnf(refute_0_97, plain,
    (0 < skoXC1 |
     -15261696/125 +
     skoXC1 *
     (11446272/390625 +
      skoXC1 *
      (-4292352/1220703125 +
       skoXC1 *
       (1073088/3814697265625 +
        skoXC1 *
        (-352107/23841857910156250 +
         skoXC1 *
         (150903/298023223876953125000 +
          skoXC1 * -150903/14901161193847656250000000))))) <=
     (-202176/25 + skoCC1 * 27072/5 + skoSC1 * 24192/25 +
      skoXC1 *
      (151632/78125 +
       skoXC1 *
       (-56862/244140625 +
        skoXC1 *
        (28431/1525878906250 +
         skoXC1 *
         (-597051/610351562500000000 +
          skoXC1 *
          (255879/7629394531250000000000 +
           skoXC1 * -255879/381469726562500000000000000)))))) *
     (2304 +
      skoXC1 *
      (5472/125 +
       skoXC1 *
       (6498/15625 +
        skoXC1 *
        (20577/7812500 +
         skoXC1 *
         (2736741/250000000000 +
          skoXC1 *
          (7428297/250000000000000 +
           skoXC1 * 47045881/1000000000000000000)))))) |
     2304 +
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_20])).

cnf(refute_0_98, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304 |
     skoXC1 *
     (136553029632/390625 + skoCC1 * -148137984/625 +
      skoSC1 * -132378624/3125 +
      skoXC1 *
      (4002370869888/1220703125 + skoCC1 * -175913856/78125 +
       skoSC1 * -157199616/390625 +
       skoXC1 *
       (78213082368528/3814697265625 + skoCC1 * -139265136/9765625 +
        skoSC1 * -124449696/48828125 +
        skoXC1 *
        (15928756333537239/190734863281250000 +
         skoCC1 * -1157641443/19531250000 +
         skoSC1 * -517244049/48828125000 +
         skoXC1 *
         (1047375813217556313/4768371582031250000000 +
          skoCC1 * -3142169631/19531250000000 +
          skoSC1 * -1403948133/48828125000000 +
          skoXC1 *
          (155111720874222605499/476837158203125000000000000 +
           skoCC1 * -19900407663/78125000000000000 +
           skoSC1 * -8891671509/195312500000000000 +
           skoXC1 *
           (-4033654511064897537/47683715820312500000000000000 +
            skoXC1 *
            (993184235083198197/95367431640625000000000000000000 +
             skoXC1 *
             (-80860515314752719/95367431640625000000000000000000000 +
              skoXC1 *
              (17178145992170631/381469726562500000000000000000000000000 +
               skoXC1 *
               (-297149834193849/
                190734863281250000000000000000000000000000 +
                skoXC1 *
                12038052984399/
                381469726562500000000000000000000000000000000))))))))))) <=
     -2313805824/125 + skoCC1 * 62373888/5 + skoSC1 * 55738368/25),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (-2313805824/125 + skoCC1 * 62373888/5 + skoSC1 * 55738368/25 <
     skoXC1 *
     (136553029632/390625 + skoCC1 * -148137984/625 +
      skoSC1 * -132378624/3125 +
      skoXC1 *
      (4002370869888/1220703125 + skoCC1 * -175913856/78125 +
       skoSC1 * -157199616/390625 +
       skoXC1 *
       (78213082368528/3814697265625 + skoCC1 * -139265136/9765625 +
        skoSC1 * -124449696/48828125 +
        skoXC1 *
        (15928756333537239/190734863281250000 +
         skoCC1 * -1157641443/19531250000 +
         skoSC1 * -517244049/48828125000 +
         skoXC1 *
         (1047375813217556313/4768371582031250000000 +
          skoCC1 * -3142169631/19531250000000 +
          skoSC1 * -1403948133/48828125000000 +
          skoXC1 *
          (155111720874222605499/476837158203125000000000000 +
           skoCC1 * -19900407663/78125000000000000 +
           skoSC1 * -8891671509/195312500000000000 +
           skoXC1 *
           (-4033654511064897537/47683715820312500000000000000 +
            skoXC1 *
            (993184235083198197/95367431640625000000000000000000 +
             skoXC1 *
             (-80860515314752719/95367431640625000000000000000000000 +
              skoXC1 *
              (17178145992170631/381469726562500000000000000000000000000 +
               skoXC1 *
               (-297149834193849/
                190734863281250000000000000000000000000000 +
                skoXC1 *
                12038052984399/
                381469726562500000000000000000000000000000000))))))))))) |
     0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_100, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (-2304 <
     skoXC1 *
     (5472/125 +
      skoXC1 *
      (6498/15625 +
       skoXC1 *
       (20577/7812500 +
        skoXC1 *
        (2736741/250000000000 +
         skoXC1 *
         (7428297/250000000000000 +
          skoXC1 * 47045881/1000000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_79, refute_0_69, refute_0_59, refute_0_60,
               refute_0_61])).

cnf(refute_0_102, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain, (skoCC1 * -235/42 < skoSC1 | 12500 < skoXC1),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_104, plain,
    (12500 < skoXC1 |
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_19, refute_0_103])).

cnf(refute_0_105, plain, (skoXC1 <= 12500),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_106, plain,
    ((-351/100 + exp(skoXC1 * -19/1000) * 23/1000) /
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    (0 <= skoCC1 * -47/20 + skoSC1 * -21/50 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) *
     (skoCC1 * -47/20 + skoSC1 * -21/50) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(resolve, [], [refute_0_106, refute_0_8])).

cnf(refute_0_108, plain,
    ((skoCC1 * -255483445248/5 + skoSC1 * -228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000 |
     skoSC1 <= skoCC1 * -235/42),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    ((skoCC1 * -255483445248/5 + skoSC1 * -228304355328/25) /
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000)))))))))))) <=
     -351/100 + exp(skoXC1 * -19/1000) * 23/1000),
    inference(resolve, [], [refute_0_108, refute_0_79])).

cnf(refute_0_110, plain,
    (21743271936 +
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     0 |
     skoCC1 * -255483445248/5 + skoSC1 * -228304355328/25 <=
     (-351/100 + exp(skoXC1 * -19/1000) * 23/1000) *
     (21743271936 +
      skoXC1 *
      (-16307453952/3125 +
       skoXC1 *
       (6115295232/9765625 +
        skoXC1 *
        (-1528823808/30517578125 +
         skoXC1 *
         (282175488/95367431640625 +
          skoXC1 *
          (-40310784/298023223876953125 +
           skoXC1 *
           (4566456/931322574615478515625 +
            skoXC1 *
            (-413343/2910383045673370361328125 +
             skoXC1 *
             (7617321/2328306436538696289062500000000 +
              skoXC1 *
              (-1712421/29103830456733703613281250000000000 +
               skoXC1 *
               (2302911/2910383045673370361328125000000000000000 +
                skoXC1 *
                (-531441/72759576141834259033203125000000000000000000 +
                 skoXC1 *
                 531441/
                 14551915228366851806640625000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_109, refute_0_7])).

cnf(refute_0_111, plain,
    (exp(skoXC1 * -19/1000) *
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000))))))))))) |
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-21743271936 <
     skoXC1 *
     (-16307453952/3125 +
      skoXC1 *
      (6115295232/9765625 +
       skoXC1 *
       (-1528823808/30517578125 +
        skoXC1 *
        (282175488/95367431640625 +
         skoXC1 *
         (-40310784/298023223876953125 +
          skoXC1 *
          (4566456/931322574615478515625 +
           skoXC1 *
           (-413343/2910383045673370361328125 +
            skoXC1 *
            (7617321/2328306436538696289062500000000 +
             skoXC1 *
             (-1712421/29103830456733703613281250000000000 +
              skoXC1 *
              (2302911/2910383045673370361328125000000000000000 +
               skoXC1 *
               (-531441/72759576141834259033203125000000000000000000 +
                skoXC1 *
                531441/
                14551915228366851806640625000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_113, plain,
    (exp(skoXC1 * -19/1000) *
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (0 <=
     -62511906816/125 +
     skoXC1 *
     (46883930112/390625 +
      skoXC1 *
      (-17581473792/1220703125 +
       skoXC1 *
       (4395368448/3814697265625 +
        skoXC1 *
        (-811254528/11920928955078125 +
         skoXC1 *
         (115893504/37252902984619140625 +
          skoXC1 *
          (-13128561/116415321826934814453125 +
           skoXC1 *
           (9506889/2910383045673370361328125000 +
            skoXC1 *
            (-175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (-52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(split, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125 |
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (-62511906816/125 <
     skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_117, plain,
    ((-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     exp(skoXC1 * -19/1000)),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (0 <= skoXC1 * -19/1000 |
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * -19/1000 * (6 + skoXC1 * -19/1000)) /
     (12 + skoXC1 * -19/1000 * (-6 + skoXC1 * -19/1000))),
    inference(resolve, [], [refute_0_117, refute_0_6])).

cnf(refute_0_119, plain,
    ((-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) /
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000))),
    inference(resolve, [], [refute_0_119, refute_0_102])).

cnf(refute_0_121, plain,
    (0 <=
     -62511906816/125 +
     skoXC1 *
     (46883930112/390625 +
      skoXC1 *
      (-17581473792/1220703125 +
       skoXC1 *
       (4395368448/3814697265625 +
        skoXC1 *
        (-811254528/11920928955078125 +
         skoXC1 *
         (115893504/37252902984619140625 +
          skoXC1 *
          (-13128561/116415321826934814453125 +
           skoXC1 *
           (9506889/2910383045673370361328125000 +
            skoXC1 *
            (-175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (-52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                -12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) |
     (12 + skoXC1 * (-57/500 + skoXC1 * 361/1000000)) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) *
     (-62511906816/125 +
      skoXC1 *
      (46883930112/390625 +
       skoXC1 *
       (-17581473792/1220703125 +
        skoXC1 *
        (4395368448/3814697265625 +
         skoXC1 *
         (-811254528/11920928955078125 +
          skoXC1 *
          (115893504/37252902984619140625 +
           skoXC1 *
           (-13128561/116415321826934814453125 +
            skoXC1 *
            (9506889/2910383045673370361328125000 +
             skoXC1 *
             (-175198383/2328306436538696289062500000000000 +
              skoXC1 *
              (39385683/29103830456733703613281250000000000000 +
               skoXC1 *
               (-52966953/2910383045673370361328125000000000000000000 +
                skoXC1 *
                (12223143/72759576141834259033203125000000000000000000000 +
                 skoXC1 *
                 -12223143/
                 14551915228366851806640625000000000000000000000000000)))))))))))) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_120, refute_0_1])).

cnf(refute_0_122, plain,
    (skoXC1 *
     (-46883930112/390625 +
      skoXC1 *
      (17581473792/1220703125 +
       skoXC1 *
       (-4395368448/3814697265625 +
        skoXC1 *
        (811254528/11920928955078125 +
         skoXC1 *
         (-115893504/37252902984619140625 +
          skoXC1 *
          (13128561/116415321826934814453125 +
           skoXC1 *
           (-9506889/2910383045673370361328125000 +
            skoXC1 *
            (175198383/2328306436538696289062500000000000 +
             skoXC1 *
             (-39385683/29103830456733703613281250000000000000 +
              skoXC1 *
              (52966953/2910383045673370361328125000000000000000000 +
               skoXC1 *
               (-12223143/72759576141834259033203125000000000000000000000 +
                skoXC1 *
                12223143/
                14551915228366851806640625000000000000000000000000000))))))))))) <=
     -62511906816/125 |
     (-750142881792/125 +
      skoXC1 *
      (22832473964544/390625 +
       skoXC1 *
       (-237292268027904/1220703125 +
        skoXC1 *
        (171600312139776/3814697265625 +
         skoXC1 *
         (-63557147943936/11920928955078125 +
          skoXC1 *
          (15785790867648/37252902984619140625 +
           skoXC1 *
           (-2901433711032/116415321826934814453125 +
            skoXC1 *
            (1653044730759/1455191522836685180664062500 +
             skoXC1 *
             (-23914335269349/582076609134674072265625000000000 +
              skoXC1 *
              (34569999613971/29103830456733703613281250000000000000 +
               skoXC1 *
               (-39753951359193/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (17848340700633/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-23971046125779/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (441952181451/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -4412554623/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    ((-750142881792/125 +
      skoXC1 *
      (22832473964544/390625 +
       skoXC1 *
       (-237292268027904/1220703125 +
        skoXC1 *
        (171600312139776/3814697265625 +
         skoXC1 *
         (-63557147943936/11920928955078125 +
          skoXC1 *
          (15785790867648/37252902984619140625 +
           skoXC1 *
           (-2901433711032/116415321826934814453125 +
            skoXC1 *
            (1653044730759/1455191522836685180664062500 +
             skoXC1 *
             (-23914335269349/582076609134674072265625000000000 +
              skoXC1 *
              (34569999613971/29103830456733703613281250000000000000 +
               skoXC1 *
               (-39753951359193/
                1455191522836685180664062500000000000000000 +
                skoXC1 *
                (17848340700633/
                 36379788070917129516601562500000000000000000000 +
                 skoXC1 *
                 (-23971046125779/
                  3637978807091712951660156250000000000000000000000000 +
                  skoXC1 *
                  (441952181451/
                   7275957614183425903320312500000000000000000000000000000 +
                   skoXC1 *
                   -4412554623/
                   14551915228366851806640625000000000000000000000000000000000)))))))))))))) /
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) <=
     -1907972112384/25 + skoCC1 * 255483445248/5 +
     skoSC1 * 228304355328/25 +
     skoXC1 *
     (1430979084288/78125 +
      skoXC1 *
      (-536617156608/244140625 +
       skoXC1 *
       (134154289152/762939453125 +
        skoXC1 *
        (-24760899072/2384185791015625 +
         skoXC1 *
         (3537271296/7450580596923828125 +
          skoXC1 *
          (-400706514/23283064365386962890625 +
           skoXC1 *
           (145083393/291038304567337036132812500 +
            skoXC1 *
            (-2673679671/232830643653869628906250000000000 +
             skoXC1 *
             (601059771/2910383045673370361328125000000000000 +
              skoXC1 *
              (-808321761/291038304567337036132812500000000000000000 +
               skoXC1 *
               (186535791/7275957614183425903320312500000000000000000000 +
                skoXC1 *
                -186535791/
                1455191522836685180664062500000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_122, refute_0_116])).

cnf(refute_0_124, plain,
    (-750142881792/125 +
     skoXC1 *
     (22832473964544/390625 +
      skoXC1 *
      (-237292268027904/1220703125 +
       skoXC1 *
       (171600312139776/3814697265625 +
        skoXC1 *
        (-63557147943936/11920928955078125 +
         skoXC1 *
         (15785790867648/37252902984619140625 +
          skoXC1 *
          (-2901433711032/116415321826934814453125 +
           skoXC1 *
           (1653044730759/1455191522836685180664062500 +
            skoXC1 *
            (-23914335269349/582076609134674072265625000000000 +
             skoXC1 *
             (34569999613971/29103830456733703613281250000000000000 +
              skoXC1 *
              (-39753951359193/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (17848340700633/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (-23971046125779/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (441952181451/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  -4412554623/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     (-1907972112384/25 + skoCC1 * 255483445248/5 +
      skoSC1 * 228304355328/25 +
      skoXC1 *
      (1430979084288/78125 +
       skoXC1 *
       (-536617156608/244140625 +
        skoXC1 *
        (134154289152/762939453125 +
         skoXC1 *
         (-24760899072/2384185791015625 +
          skoXC1 *
          (3537271296/7450580596923828125 +
           skoXC1 *
           (-400706514/23283064365386962890625 +
            skoXC1 *
            (145083393/291038304567337036132812500 +
             skoXC1 *
             (-2673679671/232830643653869628906250000000000 +
              skoXC1 *
              (601059771/2910383045673370361328125000000000000 +
               skoXC1 *
               (-808321761/291038304567337036132812500000000000000000 +
                skoXC1 *
                (186535791/7275957614183425903320312500000000000000000000 +
                 skoXC1 *
                 -186535791/
                 1455191522836685180664062500000000000000000000000000)))))))))))) *
     (12 + skoXC1 * (57/500 + skoXC1 * 361/1000000)) |
     12 + skoXC1 * (57/500 + skoXC1 * 361/1000000) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_0])).

cnf(refute_0_125, plain,
    (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12 |
     skoXC1 *
     (3335549054091264/390625 + skoCC1 * -3640639094784/625 +
      skoSC1 * -3253337063424/3125 +
      skoXC1 *
      (30877708256280576/1220703125 + skoCC1 * -1441086308352/78125 +
       skoSC1 * -1287779254272/390625 +
       skoXC1 *
       (-24104400876601344/3814697265625 +
        skoXC1 *
        (9157891644048384/11920928955078125 +
         skoXC1 *
         (-2305052460338112/37252902984619140625 +
          skoXC1 *
          (427281051458808/116415321826934814453125 +
           skoXC1 *
           (-244905887660571/1455191522836685180664062500 +
            skoXC1 *
            (3559463150366781/582076609134674072265625000000000 +
             skoXC1 *
             (-5164912679546799/29103830456733703613281250000000000000 +
              skoXC1 *
              (5958523898693217/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (-2682926482990077/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (3612953956373451/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (-66791142968679/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  668981650887/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) <=
     -113728183861248/125 + skoCC1 * 3065801342976/5 +
     skoSC1 * 2739652263936/25),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (-113728183861248/125 + skoCC1 * 3065801342976/5 +
     skoSC1 * 2739652263936/25 <
     skoXC1 *
     (3335549054091264/390625 + skoCC1 * -3640639094784/625 +
      skoSC1 * -3253337063424/3125 +
      skoXC1 *
      (30877708256280576/1220703125 + skoCC1 * -1441086308352/78125 +
       skoSC1 * -1287779254272/390625 +
       skoXC1 *
       (-24104400876601344/3814697265625 +
        skoXC1 *
        (9157891644048384/11920928955078125 +
         skoXC1 *
         (-2305052460338112/37252902984619140625 +
          skoXC1 *
          (427281051458808/116415321826934814453125 +
           skoXC1 *
           (-244905887660571/1455191522836685180664062500 +
            skoXC1 *
            (3559463150366781/582076609134674072265625000000000 +
             skoXC1 *
             (-5164912679546799/29103830456733703613281250000000000000 +
              skoXC1 *
              (5958523898693217/
               1455191522836685180664062500000000000000000 +
               skoXC1 *
               (-2682926482990077/
                36379788070917129516601562500000000000000000000 +
                skoXC1 *
                (3612953956373451/
                 3637978807091712951660156250000000000000000000000000 +
                 skoXC1 *
                 (-66791142968679/
                  7275957614183425903320312500000000000000000000000000000 +
                  skoXC1 *
                  668981650887/
                  14551915228366851806640625000000000000000000000000000000000))))))))))))) |
     skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_127, plain, (skoXC1 * (57/500 + skoXC1 * 361/1000000) <= -12),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain, (-12 < skoXC1 * (57/500 + skoXC1 * 361/1000000)),
    inference(decision, [],
              [refute_0_102, refute_0_79, refute_0_69, refute_0_59,
               refute_0_60, refute_0_61])).

cnf(refute_0_129, plain, ($false),
    inference(resolve, [], [refute_0_127, refute_0_128])).
% SZS output end CNFRefutation for Chua-1-IL-L.tptp
