% SZS output start CNFRefutation for Colpitts-BJT-2-VC2+VC1.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sin_lower_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, Y, X) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

fof('Colpitts', conjecture,
    (! [X] :
       ((0 <= X & X <= 73/50000000) =>
        0 <
        (3221/50 * sin(1710000 * X) - 527/25 * cos(1710000 * X)) *
        exp(905000 * X) + 606/25 * exp(-1860000 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 73/50000000) =>
        0 <
        (3221/50 * sin(1710000 * X) - 527/25 * cos(1710000 * X)) *
        exp(905000 * X) + 606/25 * exp(-1860000 * X))),
    inference(strip, [], ['Colpitts'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 73/50000000) =>
          0 <
          (3221/50 * sin(1710000 * X) - 527/25 * cos(1710000 * X)) *
          exp(905000 * X) + 606/25 * exp(-1860000 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 73/50000000 & 0 <= X &
        (3221/50 * sin(1710000 * X) - 527/25 * cos(1710000 * X)) *
        exp(905000 * X) + 606/25 * exp(-1860000 * X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     (3221/50 * sin(1710000 * skoXC1) - 527/25 * cos(1710000 * skoXC1)) *
     exp(905000 * skoXC1) + 606/25 * exp(-1860000 * skoXC1) <= 0 &
     skoXC1 <= 73/50000000), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    ((3221/50 * sin(1710000 * skoXC1) - 527/25 * cos(1710000 * skoXC1)) *
     exp(905000 * skoXC1) + 606/25 * exp(-1860000 * skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 73/50000000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <
     1 + skoXC1 * -1860000 |
     (1 + skoXC1 * -1860000) *
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <=
     11458704310272/25 +
     skoXC1 *
     (-2395199698580275200 +
      skoXC1 *
      (-670127945473327104000000 +
       skoXC1 *
       (1167300573103097118720000000000 +
        skoXC1 * 163293427113212982067200000000000000))) |
     13176422793216/25 +
     skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_1, plain,
    (X_000106 < exp(X_000105) | ~ lgen(1, X_000106, exp(X_000105))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain,
    (X_000105 < -1 | ~ lgen(1, X_000106, 1 + X_000105) |
     lgen(1, X_000106, exp(X_000105))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_3, plain,
    (X_000105 < -1 | X_000106 < exp(X_000105) |
     ~ lgen(1, X_000106, 1 + X_000105)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000105 < -1 | X_000106 < exp(X_000105) | 1 + X_000105 <= X_000106),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoXC1 * -1860000 < -1 |
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <
     exp(skoXC1 * -1860000) |
     1 + skoXC1 * -1860000 <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (21743271936 <
     exp(skoXC1 * -1860000) *
     (-13176422793216/25 +
      skoXC1 *
      (476986505114419200 +
       skoXC1 *
       (-215836393564274688000000 +
        skoXC1 *
        (65110645391889530880000000000 +
         skoXC1 *
         (-14501107214916334387200000000000000 +
          skoXC1 *
          (2499714672285577641984000000000000000000 +
           skoXC1 *
           (-341692768615286381322240000000000000000000000 +
            skoXC1 *
            (37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (-3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (-11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 -7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) |
     0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) |
     21743271936 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) *
     (-13176422793216/25 +
      skoXC1 *
      (476986505114419200 +
       skoXC1 *
       (-215836393564274688000000 +
        skoXC1 *
        (65110645391889530880000000000 +
         skoXC1 *
         (-14501107214916334387200000000000000 +
          skoXC1 *
          (2499714672285577641984000000000000000000 +
           skoXC1 *
           (-341692768615286381322240000000000000000000000 +
            skoXC1 *
            (37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (-3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (-11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 -7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_7, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-19677661102080000 +
       skoXC1 *
       (8904141648691200000000 +
        skoXC1 *
        (-2686082730688512000000000000 +
         skoXC1 *
         (598230495664865280000000000000000 +
          skoXC1 *
          (-103123542586038681600000000000000000000 +
           skoXC1 *
           (14096236329013464576000000000000000000000000 +
            skoXC1 *
            (-1539649261108625829120000000000000000000000000000 +
             skoXC1 *
             (133739845973308424420100000000000000000000000000000000 +
              skoXC1 *
              (-9069773275373332296913500000000000000000000000000000000000 +
               skoXC1 *
               (459939149072272648591324687500000000000000000000000000000000000 +
                skoXC1 *
                (-16009420381169490268274955468750000000000000000000000000000000000000 +
                 skoXC1 *
                 301844280103299764433100722900390625000000000000000000000000000000000000)))))))))))) <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) *
     (21743271936 +
      skoXC1 *
      (-19677661102080000 +
       skoXC1 *
       (8904141648691200000000 +
        skoXC1 *
        (-2686082730688512000000000000 +
         skoXC1 *
         (598230495664865280000000000000000 +
          skoXC1 *
          (-103123542586038681600000000000000000000 +
           skoXC1 *
           (14096236329013464576000000000000000000000000 +
            skoXC1 *
            (-1539649261108625829120000000000000000000000000000 +
             skoXC1 *
             (133739845973308424420100000000000000000000000000000000 +
              skoXC1 *
              (-9069773275373332296913500000000000000000000000000000000000 +
               skoXC1 *
               (459939149072272648591324687500000000000000000000000000000000000 +
                skoXC1 *
                (-16009420381169490268274955468750000000000000000000000000000000000000 +
                 skoXC1 *
                 301844280103299764433100722900390625000000000000000000000000000000000000)))))))))))) <=
     21743271936 |
     21743271936 +
     skoXC1 *
     (-19677661102080000 +
      skoXC1 *
      (8904141648691200000000 +
       skoXC1 *
       (-2686082730688512000000000000 +
        skoXC1 *
        (598230495664865280000000000000000 +
         skoXC1 *
         (-103123542586038681600000000000000000000 +
          skoXC1 *
          (14096236329013464576000000000000000000000000 +
           skoXC1 *
           (-1539649261108625829120000000000000000000000000000 +
            skoXC1 *
            (133739845973308424420100000000000000000000000000000000 +
             skoXC1 *
             (-9069773275373332296913500000000000000000000000000000000000 +
              skoXC1 *
              (459939149072272648591324687500000000000000000000000000000000000 +
               skoXC1 *
               (-16009420381169490268274955468750000000000000000000000000000000000000 +
                skoXC1 *
                301844280103299764433100722900390625000000000000000000000000000000000000))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_8, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000))) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) |
     0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))) *
     (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) *
     (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_9, plain,
    (cos(X_000031) < X_000032 | ~ lgen(1, cos(X_000031), X_000032)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_10, plain,
    (~ lgen(1, 1 - X_000031 ^ 2 / 2 + X_000031 ^ 4 / 24, X_000032) |
     lgen(1, cos(X_000031), X_000032)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_11, plain,
    (cos(X_000031) < X_000032 |
     ~ lgen(1, 1 - X_000031 ^ 2 / 2 + X_000031 ^ 4 / 24, X_000032)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (cos(X_000031) < X_000032 |
     X_000032 <=
     1 + X_000031 * (X_000031 * (-1/2 + X_000031 * (X_000031 * 1/24)))),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (cos(skoXC1 * 1710000) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 + skoXC1 * 1710000 * (skoXC1 * 1710000 * 1/24)))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50) <
     skoXC1 * (1710000 + skoXC1 * (skoXC1 * -833368500000000000)) |
     skoXC1 * (1710000 + skoXC1 * (skoXC1 * -833368500000000000)) *
     (exp(skoXC1 * 905000) * 3221/50) <=
     exp(skoXC1 * -1860000) * -606/25 +
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25) |
     exp(skoXC1 * 905000) * 3221/50 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_15, plain,
    ((3221/50 * sin(1710000 * skoXC1) - 527/25 * cos(1710000 * skoXC1)) *
     exp(905000 * skoXC1) + 606/25 * exp(-1860000 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_16, plain,
    (sin(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 3221/50) <=
     exp(skoXC1 * -1860000) * -606/25 +
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (exp(skoXC1 * 905000) * 3221/50 <= 0 |
     sin(skoXC1 * 1710000) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(split, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (sin(skoXC1 * 1710000) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (X_000120 < sin(X_000119) | ~ lgen(1, X_000120, sin(X_000119))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_20, plain,
    (X_000119 < 0 | ~ lgen(1, X_000120, X_000119 - X_000119 ^ 3 / 6) |
     lgen(1, X_000120, sin(X_000119))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_21, plain,
    (X_000119 < 0 | X_000120 < sin(X_000119) |
     ~ lgen(1, X_000120, X_000119 - X_000119 ^ 3 / 6)),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (X_000119 < 0 | X_000120 < sin(X_000119) |
     X_000119 * (1 + X_000119 * (X_000119 * -1/6)) <= X_000120),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 * 1710000 < 0 |
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50) < sin(skoXC1 * 1710000) |
     skoXC1 * 1710000 *
     (1 + skoXC1 * 1710000 * (skoXC1 * 1710000 * -1/6)) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * 1710000 < 0 |
     skoXC1 * 1710000 *
     (1 + skoXC1 * 1710000 * (skoXC1 * 1710000 * -1/6)) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(resolve, [], [refute_0_18, refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 < 0 |
     skoXC1 * (1710000 + skoXC1 * (skoXC1 * -833368500000000000)) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain,
    (skoXC1 * (1710000 + skoXC1 * (skoXC1 * -833368500000000000)) <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(resolve, [], [refute_0_26, refute_0_25])).

cnf(refute_0_28, plain,
    (skoXC1 * (1710000 + skoXC1 * (skoXC1 * -833368500000000000)) *
     (exp(skoXC1 * 905000) * 3221/50) <=
     exp(skoXC1 * -1860000) * -606/25 +
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25) |
     exp(skoXC1 * 905000) * 3221/50 <= 0),
    inference(resolve, [], [refute_0_27, refute_0_14])).

cnf(refute_0_29, plain,
    (cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) *
     (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (0 <= exp(skoXC1 * 905000) * -527/25 |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <= cos(skoXC1 * 1710000)),
    inference(split, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <= cos(skoXC1 * 1710000)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 + skoXC1 * 1710000 * (skoXC1 * 1710000 * 1/24)))),
    inference(resolve, [], [refute_0_31, refute_0_13])).

cnf(refute_0_33, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))) *
     (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) *
     (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))),
    inference(resolve, [], [refute_0_33, refute_0_8])).

cnf(refute_0_35, plain,
    (exp(skoXC1 * 905000) *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) * -606/25),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)), inference(split, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 + skoXC1 * 7510066911450000000000000))) <=
     -527/25 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)), inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (-527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <
     0 |
     0 <
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) |
     0 <= exp(skoXC1 * -1860000) * -606/25),
    inference(split, [], [refute_0_35])).

cnf(refute_0_39, plain,
    (-527/25 <
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 + skoXC1 * 7510066911450000000000000))) |
     527/25 <
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (0 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <
     exp(skoXC1 * -1860000) * -606/25 |
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     0 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (X_000014 < exp(X_000013) | ~ lgen(1, X_000014, exp(X_000013))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_42, plain, (0 < X_000014 | lgen(1, X_000014, exp(X_000013))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_43, plain, (0 < X_000014 | X_000014 < exp(X_000013)),
    inference(resolve, [], [refute_0_42, refute_0_41])).

cnf(refute_0_44, plain,
    (0 <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <
     exp(skoXC1 * 905000)), inference(subst, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (-527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     0 |
     exp(skoXC1 * 905000) <=
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000))))),
    inference(split, [], [refute_0_35])).

cnf(refute_0_46, plain,
    (skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     527/25 |
     exp(skoXC1 * 905000) <=
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000))))),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (0 <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) |
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     527/25), inference(resolve, [], [refute_0_46, refute_0_44])).

cnf(refute_0_48, plain,
    (0 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <
     exp(skoXC1 * -1860000) * -606/25 |
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     527/25 |
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     0), inference(resolve, [], [refute_0_40, refute_0_47])).

cnf(refute_0_49, plain,
    (skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) <=
     527/25), inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (-527/25 <
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 + skoXC1 * 7510066911450000000000000)))),
    inference(resolve, [], [refute_0_49, refute_0_39])).

cnf(refute_0_51, plain,
    (exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)),
    inference(resolve, [], [refute_0_37, refute_0_50])).

cnf(refute_0_52, plain,
    (exp(X_000135) < X_000136 | ~ lgen(1, exp(X_000135), X_000136)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_53, plain,
    (3 < X_000135 |
     ~ lgen(1,
            21743271936 /
            (-(X_000135 ^ 3) + 12 * X_000135 ^ 2 - 96 * X_000135 + 384) ^
            4, X_000136) | lgen(1, exp(X_000135), X_000136)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_54, plain,
    (3 < X_000135 | exp(X_000135) < X_000136 |
     ~ lgen(1,
            21743271936 /
            (-(X_000135 ^ 3) + 12 * X_000135 ^ 2 - 96 * X_000135 + 384) ^
            4, X_000136)),
    inference(resolve, [], [refute_0_53, refute_0_52])).

cnf(refute_0_55, plain,
    (3 < X_000135 | exp(X_000135) < X_000136 |
     X_000136 <=
     21743271936 /
     (21743271936 +
      X_000135 *
      (-21743271936 +
       X_000135 *
       (10871635968 +
        X_000135 *
        (-3623878656 +
         X_000135 *
         (891813888 +
          X_000135 *
          (-169869312 +
           X_000135 *
           (25657344 +
            X_000135 *
            (-3096576 +
             X_000135 *
             (297216 +
              X_000135 *
              (-22272 +
               X_000135 * (1248 + X_000135 * (-48 + X_000135))))))))))))),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (3 < skoXC1 * 905000 |
     exp(skoXC1 * 905000) <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 905000 *
      (-21743271936 +
       skoXC1 * 905000 *
       (10871635968 +
        skoXC1 * 905000 *
        (-3623878656 +
         skoXC1 * 905000 *
         (891813888 +
          skoXC1 * 905000 *
          (-169869312 +
           skoXC1 * 905000 *
           (25657344 +
            skoXC1 * 905000 *
            (-3096576 +
             skoXC1 * 905000 *
             (297216 +
              skoXC1 * 905000 *
              (-22272 +
               skoXC1 * 905000 *
               (1248 +
                skoXC1 * 905000 * (-48 + skoXC1 * 905000))))))))))))),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (3 < skoXC1 * 905000 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936 /
     (21743271936 +
      skoXC1 * 905000 *
      (-21743271936 +
       skoXC1 * 905000 *
       (10871635968 +
        skoXC1 * 905000 *
        (-3623878656 +
         skoXC1 * 905000 *
         (891813888 +
          skoXC1 * 905000 *
          (-169869312 +
           skoXC1 * 905000 *
           (25657344 +
            skoXC1 * 905000 *
            (-3096576 +
             skoXC1 * 905000 *
             (297216 +
              skoXC1 * 905000 *
              (-22272 +
               skoXC1 * 905000 *
               (1248 +
                skoXC1 * 905000 * (-48 + skoXC1 * 905000))))))))))))),
    inference(resolve, [], [refute_0_51, refute_0_56])).

cnf(refute_0_58, plain,
    (3/905000 < skoXC1 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-19677661102080000 +
       skoXC1 *
       (8904141648691200000000 +
        skoXC1 *
        (-2686082730688512000000000000 +
         skoXC1 *
         (598230495664865280000000000000000 +
          skoXC1 *
          (-103123542586038681600000000000000000000 +
           skoXC1 *
           (14096236329013464576000000000000000000000000 +
            skoXC1 *
            (-1539649261108625829120000000000000000000000000000 +
             skoXC1 *
             (133739845973308424420100000000000000000000000000000000 +
              skoXC1 *
              (-9069773275373332296913500000000000000000000000000000000000 +
               skoXC1 *
               (459939149072272648591324687500000000000000000000000000000000000 +
                skoXC1 *
                (-16009420381169490268274955468750000000000000000000000000000000000000 +
                 skoXC1 *
                 301844280103299764433100722900390625000000000000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    ((1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <
     1 + skoXC1 * -1860000 |
     (1 + skoXC1 * -1860000) *
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <=
     1214208/25 +
     skoXC1 *
     (-253804492800 +
      skoXC1 *
      (-71009312256000000 +
       skoXC1 * (skoXC1 * 17303194163980800000000000000))) |
     1396224/25 +
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_60, plain,
    (skoXC1 * -1860000 < -1 |
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <
     exp(skoXC1 * -1860000) |
     1 + skoXC1 * -1860000 <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_61, plain,
    (2304 <
     exp(skoXC1 * -1860000) *
     (-1396224/25 +
      skoXC1 *
      (50543308800 +
       skoXC1 *
       (-22870847232000000 +
        skoXC1 *
        (6899372248320000000000 +
         skoXC1 *
         (-1365860099784600000000000000 +
          skoXC1 *
          (176586198615009000000000000000000 +
           skoXC1 * -13317542478881928750000000000000000000)))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) |
     0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 + skoXC1 * (skoXC1 * -7510066911450000000000000))) |
     2304 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) *
     (-1396224/25 +
      skoXC1 *
      (50543308800 +
       skoXC1 *
       (-22870847232000000 +
        skoXC1 *
        (6899372248320000000000 +
         skoXC1 *
         (-1365860099784600000000000000 +
          skoXC1 *
          (176586198615009000000000000000000 +
           skoXC1 * -13317542478881928750000000000000000000))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_62, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-2085120000 +
       skoXC1 *
       (943516800000000 +
        skoXC1 *
        (-284627568000000000000 +
         skoXC1 *
         (56347363852500000000000000 +
          skoXC1 *
          (-7284909183787500000000000000000 +
           skoXC1 * 549403567610640625000000000000000000)))))) <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) *
     (2304 +
      skoXC1 *
      (-2085120000 +
       skoXC1 *
       (943516800000000 +
        skoXC1 *
        (-284627568000000000000 +
         skoXC1 *
         (56347363852500000000000000 +
          skoXC1 *
          (-7284909183787500000000000000000 +
           skoXC1 * 549403567610640625000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (-2085120000 +
      skoXC1 *
      (943516800000000 +
       skoXC1 *
       (-284627568000000000000 +
        skoXC1 *
        (56347363852500000000000000 +
         skoXC1 *
         (-7284909183787500000000000000000 +
          skoXC1 * 549403567610640625000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_63, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000))) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) |
     0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))) *
     (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) * (skoXC1 * -110158200)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_64, plain,
    (cos(skoXC1 * 1710000) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 + skoXC1 * 1710000 * (skoXC1 * 1710000 * 1/24)))),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_65, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50) < skoXC1 * 1710000 |
     skoXC1 * 1710000 * (exp(skoXC1 * 905000) * 3221/50) <=
     exp(skoXC1 * -1860000) * -606/25 +
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25) |
     exp(skoXC1 * 905000) * 3221/50 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_66, plain,
    (X_000099 < sin(X_000098) | ~ lgen(1, X_000099, sin(X_000098))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_67, plain,
    (0 < X_000098 | ~ lgen(1, X_000099, X_000098) |
     lgen(1, X_000099, sin(X_000098))),
    inference(subst, [], [sin_lower_bound_0_neg])).

cnf(refute_0_68, plain,
    (0 < X_000098 | X_000099 < sin(X_000098) |
     ~ lgen(1, X_000099, X_000098)),
    inference(resolve, [], [refute_0_67, refute_0_66])).

cnf(refute_0_69, plain,
    (0 < X_000098 | X_000099 < sin(X_000098) | X_000098 <= X_000099),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (0 < skoXC1 * 1710000 |
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50) < sin(skoXC1 * 1710000) |
     skoXC1 * 1710000 <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(subst, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (0 < skoXC1 * 1710000 |
     skoXC1 * 1710000 <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(resolve, [], [refute_0_18, refute_0_70])).

cnf(refute_0_72, plain,
    (0 < skoXC1 |
     skoXC1 * 1710000 <=
     (exp(skoXC1 * -1860000) * -606/25 +
      cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25)) /
     (exp(skoXC1 * 905000) * 3221/50)),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (0 < skoXC1 |
     skoXC1 * 1710000 * (exp(skoXC1 * 905000) * 3221/50) <=
     exp(skoXC1 * -1860000) * -606/25 +
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * 527/25) |
     exp(skoXC1 * 905000) * 3221/50 <= 0),
    inference(resolve, [], [refute_0_72, refute_0_65])).

cnf(refute_0_74, plain,
    (0 < skoXC1 |
     cos(skoXC1 * 1710000) * (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) * (skoXC1 * -110158200)),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (0 < skoXC1 | 0 <= exp(skoXC1 * 905000) * -527/25 |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) <= cos(skoXC1 * 1710000)),
    inference(split, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (0 < skoXC1 |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) <= cos(skoXC1 * 1710000)),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (0 < skoXC1 |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 + skoXC1 * 1710000 * (skoXC1 * 1710000 * 1/24)))),
    inference(resolve, [], [refute_0_76, refute_0_64])).

cnf(refute_0_78, plain,
    (0 < skoXC1 |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) * (skoXC1 * -110158200)) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (0 < skoXC1 | 0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 + skoXC1 * (skoXC1 * 356265033750000000000000)))) *
     (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) * (skoXC1 * -110158200)),
    inference(resolve, [], [refute_0_78, refute_0_63])).

cnf(refute_0_80, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 905000) *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) * -606/25),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (0 < skoXC1 |
     0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 + skoXC1 * (skoXC1 * -7510066911450000000000000))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)), inference(split, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 + skoXC1 * (skoXC1 * 7510066911450000000000000))) <=
     -527/25 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)), inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain, (skoXC1 <= 73/50000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_84, plain,
    (-527/25 <
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 + skoXC1 * (skoXC1 * 7510066911450000000000000))) |
     0 < skoXC1), inference(decision, [], [refute_0_83, refute_0_26])).

cnf(refute_0_85, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * 905000)),
    inference(resolve, [], [refute_0_82, refute_0_84])).

cnf(refute_0_86, plain,
    (exp(X_000128) < X_000129 | ~ lgen(1, exp(X_000128), X_000129)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_87, plain,
    (0 < X_000128 |
     ~ lgen(1,
            2304 /
            (-(X_000128 ^ 3) + 6 * X_000128 ^ 2 - 24 * X_000128 + 48) ^ 2,
            X_000129) | lgen(1, exp(X_000128), X_000129)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_88, plain,
    (0 < X_000128 | exp(X_000128) < X_000129 |
     ~ lgen(1,
            2304 /
            (-(X_000128 ^ 3) + 6 * X_000128 ^ 2 - 24 * X_000128 + 48) ^ 2,
            X_000129)),
    inference(resolve, [], [refute_0_87, refute_0_86])).

cnf(refute_0_89, plain,
    (0 < X_000128 | exp(X_000128) < X_000129 |
     X_000129 <=
     2304 /
     (2304 +
      X_000128 *
      (-2304 +
       X_000128 *
       (1152 +
        X_000128 *
        (-384 + X_000128 * (84 + X_000128 * (-12 + X_000128))))))),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (0 < skoXC1 * 905000 |
     exp(skoXC1 * 905000) <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     2304 /
     (2304 +
      skoXC1 * 905000 *
      (-2304 +
       skoXC1 * 905000 *
       (1152 +
        skoXC1 * 905000 *
        (-384 +
         skoXC1 * 905000 *
         (84 + skoXC1 * 905000 * (-12 + skoXC1 * 905000))))))),
    inference(subst, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (0 < skoXC1 * 905000 | 0 < skoXC1 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     2304 /
     (2304 +
      skoXC1 * 905000 *
      (-2304 +
       skoXC1 * 905000 *
       (1152 +
        skoXC1 * 905000 *
        (-384 +
         skoXC1 * 905000 *
         (84 + skoXC1 * 905000 * (-12 + skoXC1 * 905000))))))),
    inference(resolve, [], [refute_0_85, refute_0_90])).

cnf(refute_0_92, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     2304 /
     (2304 +
      skoXC1 *
      (-2085120000 +
       skoXC1 *
       (943516800000000 +
        skoXC1 *
        (-284627568000000000000 +
         skoXC1 *
         (56347363852500000000000000 +
          skoXC1 *
          (-7284909183787500000000000000000 +
           skoXC1 * 549403567610640625000000000000000000))))))),
    inference(arithmetic, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) *
     (2304 +
      skoXC1 *
      (-2085120000 +
       skoXC1 *
       (943516800000000 +
        skoXC1 *
        (-284627568000000000000 +
         skoXC1 *
         (56347363852500000000000000 +
          skoXC1 *
          (-7284909183787500000000000000000 +
           skoXC1 * 549403567610640625000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (-2085120000 +
      skoXC1 *
      (943516800000000 +
       skoXC1 *
       (-284627568000000000000 +
        skoXC1 *
        (56347363852500000000000000 +
         skoXC1 *
         (-7284909183787500000000000000000 +
          skoXC1 * 549403567610640625000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_92, refute_0_62])).

cnf(refute_0_94, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-2085120000 +
      skoXC1 *
      (943516800000000 +
       skoXC1 *
       (-284627568000000000000 +
        skoXC1 *
        (56347363852500000000000000 +
         skoXC1 *
         (-7284909183787500000000000000000 +
          skoXC1 * 549403567610640625000000000000000000))))) <= -2304 |
     exp(skoXC1 * -1860000) *
     (-1396224/25 +
      skoXC1 *
      (50543308800 +
       skoXC1 *
       (-22870847232000000 +
        skoXC1 *
        (6899372248320000000000 +
         skoXC1 *
         (-1365860099784600000000000000 +
          skoXC1 *
          (176586198615009000000000000000000 +
           skoXC1 * -13317542478881928750000000000000000000)))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <= 2304),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 +
       skoXC1 *
       (skoXC1 *
        (356265033750000000000000 +
         skoXC1 *
         (skoXC1 *
          (-34725152839612500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (12692477427288863906250000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (-412378591612615188314062500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                100486353311204006012429179687500000000000000000000000000000000000000/
                77))))))))))) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) |
     0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 +
        skoXC1 *
        (skoXC1 *
         (356265033750000000000000 +
          skoXC1 *
          (skoXC1 *
           (-34725152839612500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (12692477427288863906250000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (-412378591612615188314062500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 100486353311204006012429179687500000000000000000000000000000000000000/
                 77)))))))))))) * (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) *
     (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_96, plain,
    (cos(X_000045) < X_000046 | ~ lgen(1, cos(X_000045), X_000046)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_97, plain,
    (~ lgen(1,
            1 - X_000045 ^ 2 / 2 + X_000045 ^ 4 / 24 - X_000045 ^ 6 / 720 +
            X_000045 ^ 8 / 40320 - X_000045 ^ 10 / 3628800 +
            X_000045 ^ 12 / 479001600, X_000046) |
     lgen(1, cos(X_000045), X_000046)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_98, plain,
    (cos(X_000045) < X_000046 |
     ~ lgen(1,
            1 - X_000045 ^ 2 / 2 + X_000045 ^ 4 / 24 - X_000045 ^ 6 / 720 +
            X_000045 ^ 8 / 40320 - X_000045 ^ 10 / 3628800 +
            X_000045 ^ 12 / 479001600, X_000046)),
    inference(resolve, [], [refute_0_97, refute_0_96])).

cnf(refute_0_99, plain,
    (cos(X_000045) < X_000046 |
     X_000046 <=
     1 +
     X_000045 *
     (X_000045 *
      (-1/2 +
       X_000045 *
       (X_000045 *
        (1/24 +
         X_000045 *
         (X_000045 *
          (-1/720 +
           X_000045 *
           (X_000045 *
            (1/40320 +
             X_000045 *
             (X_000045 *
              (-1/3628800 + X_000045 * (X_000045 * 1/479001600)))))))))))),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (cos(skoXC1 * 1710000) <
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) |
     (exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 +
       skoXC1 * 1710000 *
       (skoXC1 * 1710000 *
        (1/24 +
         skoXC1 * 1710000 *
         (skoXC1 * 1710000 *
          (-1/720 +
           skoXC1 * 1710000 *
           (skoXC1 * 1710000 *
            (1/40320 +
             skoXC1 * 1710000 *
             (skoXC1 * 1710000 *
              (-1/3628800 +
               skoXC1 * 1710000 *
               (skoXC1 * 1710000 * 1/479001600)))))))))))),
    inference(subst, [], [refute_0_99])).

cnf(refute_0_101, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 * 1710000 *
     (skoXC1 * 1710000 *
      (-1/2 +
       skoXC1 * 1710000 *
       (skoXC1 * 1710000 *
        (1/24 +
         skoXC1 * 1710000 *
         (skoXC1 * 1710000 *
          (-1/720 +
           skoXC1 * 1710000 *
           (skoXC1 * 1710000 *
            (1/40320 +
             skoXC1 * 1710000 *
             (skoXC1 * 1710000 *
              (-1/3628800 +
               skoXC1 * 1710000 *
               (skoXC1 * 1710000 * 1/479001600)))))))))))),
    inference(resolve, [], [refute_0_31, refute_0_100])).

cnf(refute_0_102, plain,
    ((exp(skoXC1 * -1860000) * -606/25 +
      exp(skoXC1 * 905000) *
      (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))) /
     (exp(skoXC1 * 905000) * -527/25) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1462050000000 +
       skoXC1 *
       (skoXC1 *
        (356265033750000000000000 +
         skoXC1 *
         (skoXC1 *
          (-34725152839612500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (12692477427288863906250000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (-412378591612615188314062500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                100486353311204006012429179687500000000000000000000000000000000000000/
                77)))))))))))), inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (0 <= exp(skoXC1 * 905000) * -527/25 |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1462050000000 +
        skoXC1 *
        (skoXC1 *
         (356265033750000000000000 +
          skoXC1 *
          (skoXC1 *
           (-34725152839612500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (12692477427288863906250000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (-412378591612615188314062500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 100486353311204006012429179687500000000000000000000000000000000000000/
                 77)))))))))))) * (exp(skoXC1 * 905000) * -527/25) <=
     exp(skoXC1 * -1860000) * -606/25 +
     exp(skoXC1 * 905000) *
     (skoXC1 * (-110158200 + skoXC1 * (skoXC1 * 53685598770000000000)))),
    inference(resolve, [], [refute_0_102, refute_0_95])).

cnf(refute_0_104, plain,
    (exp(skoXC1 * 905000) *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) <= exp(skoXC1 * -1860000) * -606/25),
    inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (-527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) < 0 |
     0 <
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) | 0 <= exp(skoXC1 * -1860000) * -606/25),
    inference(split, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (-527/25 <
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 +
        skoXC1 *
        (7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (-732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (-8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) |
     527/25 <
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77)))))))))))), inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (0 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) < exp(skoXC1 * -1860000) * -606/25 |
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 0 |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_108, plain,
    (0 <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) |
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) < exp(skoXC1 * 905000)),
    inference(subst, [], [refute_0_43])).

cnf(refute_0_109, plain,
    (-527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 0 |
     exp(skoXC1 * 905000) <=
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77))))))))))))), inference(split, [], [refute_0_104])).

cnf(refute_0_110, plain,
    (skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 527/25 |
     exp(skoXC1 * 905000) <=
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77))))))))))))),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (0 <
     exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) |
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 527/25),
    inference(resolve, [], [refute_0_110, refute_0_108])).

cnf(refute_0_112, plain,
    (0 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 +
         skoXC1 *
         (-7510066911450000000000000 +
          skoXC1 *
          (skoXC1 *
           (732006221859031500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-267557424167249251143750000000000000000000000000/7 +
              skoXC1 *
              (skoXC1 *
               (8692940711193928169660437500000000000000000000000000000000/
                7 +
                skoXC1 *
                (skoXC1 *
                 -2118252327800180446742007107812500000000000000000000000000000000000000/
                 77)))))))))))) < exp(skoXC1 * -1860000) * -606/25 |
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 527/25 |
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 0),
    inference(resolve, [], [refute_0_107, refute_0_111])).

cnf(refute_0_113, plain,
    (skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 +
        skoXC1 *
        (-7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                -2118252327800180446742007107812500000000000000000000000000000000000000/
                77))))))))))) <= 527/25),
    inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (-527/25 <
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 +
        skoXC1 *
        (7510066911450000000000000 +
         skoXC1 *
         (skoXC1 *
          (-732006221859031500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (267557424167249251143750000000000000000000000000/7 +
             skoXC1 *
             (skoXC1 *
              (-8692940711193928169660437500000000000000000000000000000000/
               7 +
               skoXC1 *
               (skoXC1 *
                2118252327800180446742007107812500000000000000000000000000000000000000/
                77)))))))))))),
    inference(resolve, [], [refute_0_113, refute_0_106])).

cnf(refute_0_115, plain,
    (-2304 <
     skoXC1 *
     (-2085120000 +
      skoXC1 *
      (943516800000000 +
       skoXC1 *
       (-284627568000000000000 +
        skoXC1 *
        (56347363852500000000000000 +
         skoXC1 *
         (-7284909183787500000000000000000 +
          skoXC1 * 549403567610640625000000000000000000))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_114, refute_0_113, refute_0_49, refute_0_39,
               refute_0_83, refute_0_26])).

cnf(refute_0_116, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) *
     (-1396224/25 +
      skoXC1 *
      (50543308800 +
       skoXC1 *
       (-22870847232000000 +
        skoXC1 *
        (6899372248320000000000 +
         skoXC1 *
         (-1365860099784600000000000000 +
          skoXC1 *
          (176586198615009000000000000000000 +
           skoXC1 * -13317542478881928750000000000000000000)))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <= 2304),
    inference(resolve, [], [refute_0_94, refute_0_115])).

cnf(refute_0_117, plain,
    (0 < skoXC1 |
     0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 + skoXC1 * (skoXC1 * -7510066911450000000000000))) |
     2304 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 * (skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) *
     (-1396224/25 +
      skoXC1 *
      (50543308800 +
       skoXC1 *
       (-22870847232000000 +
        skoXC1 *
        (6899372248320000000000 +
         skoXC1 *
         (-1365860099784600000000000000 +
          skoXC1 *
          (176586198615009000000000000000000 +
           skoXC1 * -13317542478881928750000000000000000000))))))),
    inference(resolve, [], [refute_0_116, refute_0_61])).

cnf(refute_0_118, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) *
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <=
     1214208/25 +
     skoXC1 *
     (-253804492800 +
      skoXC1 *
      (-71009312256000000 +
       skoXC1 * (skoXC1 * 17303194163980800000000000000))) |
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 + skoXC1 * (skoXC1 * 7510066911450000000000000))) <=
     -527/25), inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) *
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <=
     1214208/25 +
     skoXC1 *
     (-253804492800 +
      skoXC1 *
      (-71009312256000000 +
       skoXC1 * (skoXC1 * 17303194163980800000000000000)))),
    inference(resolve, [], [refute_0_118, refute_0_84])).

cnf(refute_0_120, plain,
    (0 < skoXC1 |
     1396224/25 +
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <= 0 |
     exp(skoXC1 * -1860000) <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(split, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <=
     -1396224/25 |
     exp(skoXC1 * -1860000) <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (-1396224/25 <
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_114, refute_0_113, refute_0_49, refute_0_39,
               refute_0_83, refute_0_26])).

cnf(refute_0_123, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -1860000) <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(resolve, [], [refute_0_121, refute_0_122])).

cnf(refute_0_124, plain,
    (skoXC1 * -1860000 < -1 | 0 < skoXC1 |
     1 + skoXC1 * -1860000 <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(resolve, [], [refute_0_123, refute_0_60])).

cnf(refute_0_125, plain,
    (0 < skoXC1 | 1/1860000 < skoXC1 |
     1 + skoXC1 * -1860000 <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain, (0 < skoXC1 | skoXC1 <= 1/1860000),
    inference(decision, [],
              [refute_0_114, refute_0_113, refute_0_49, refute_0_39,
               refute_0_83, refute_0_26])).

cnf(refute_0_127, plain,
    (0 < skoXC1 |
     1 + skoXC1 * -1860000 <=
     (1214208/25 +
      skoXC1 *
      (-253804492800 +
       skoXC1 *
       (-71009312256000000 +
        skoXC1 * (skoXC1 * 17303194163980800000000000000)))) /
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000))))))),
    inference(resolve, [], [refute_0_126, refute_0_125])).

cnf(refute_0_128, plain,
    (0 < skoXC1 |
     (1 + skoXC1 * -1860000) *
     (1396224/25 +
      skoXC1 *
      (-50543308800 +
       skoXC1 *
       (22870847232000000 +
        skoXC1 *
        (-6899372248320000000000 +
         skoXC1 *
         (1365860099784600000000000000 +
          skoXC1 *
          (-176586198615009000000000000000000 +
           skoXC1 * 13317542478881928750000000000000000000)))))) <=
     1214208/25 +
     skoXC1 *
     (-253804492800 +
      skoXC1 *
      (-71009312256000000 +
       skoXC1 * (skoXC1 * 17303194163980800000000000000))) |
     1396224/25 +
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_127, refute_0_59])).

cnf(refute_0_129, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <=
     -1396224/25 |
     skoXC1 *
     (99382118400 +
      skoXC1 *
      (187890713856000000 +
       skoXC1 *
       (-49439148099840000000000 +
        skoXC1 *
        (-3104501682321000000000000000 +
         skoXC1 *
         (-2717085984214365000000000000000000 +
          skoXC1 *
          (341767871902798668750000000000000000000 +
           skoXC1 * -24770629010720387475000000000000000000000000)))))) <=
     -182016/25), inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (-182016/25 <
     skoXC1 *
     (99382118400 +
      skoXC1 *
      (187890713856000000 +
       skoXC1 *
       (-49439148099840000000000 +
        skoXC1 *
        (-3104501682321000000000000000 +
         skoXC1 *
         (-2717085984214365000000000000000000 +
          skoXC1 *
          (341767871902798668750000000000000000000 +
           skoXC1 * -24770629010720387475000000000000000000000000)))))) |
     0 < skoXC1 |
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <=
     -1396224/25),
    inference(decision, [],
              [refute_0_114, refute_0_113, refute_0_49, refute_0_39,
               refute_0_83, refute_0_26])).

cnf(refute_0_131, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-50543308800 +
      skoXC1 *
      (22870847232000000 +
       skoXC1 *
       (-6899372248320000000000 +
        skoXC1 *
        (1365860099784600000000000000 +
         skoXC1 *
         (-176586198615009000000000000000000 +
          skoXC1 * 13317542478881928750000000000000000000))))) <=
     -1396224/25), inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_131, refute_0_122])).

cnf(refute_0_133, plain, (skoXC1 <= 3/905000),
    inference(decision, [],
              [refute_0_132, refute_0_114, refute_0_113, refute_0_49,
               refute_0_39, refute_0_83, refute_0_26])).

cnf(refute_0_134, plain,
    (exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-19677661102080000 +
       skoXC1 *
       (8904141648691200000000 +
        skoXC1 *
        (-2686082730688512000000000000 +
         skoXC1 *
         (598230495664865280000000000000000 +
          skoXC1 *
          (-103123542586038681600000000000000000000 +
           skoXC1 *
           (14096236329013464576000000000000000000000000 +
            skoXC1 *
            (-1539649261108625829120000000000000000000000000000 +
             skoXC1 *
             (133739845973308424420100000000000000000000000000000000 +
              skoXC1 *
              (-9069773275373332296913500000000000000000000000000000000000 +
               skoXC1 *
               (459939149072272648591324687500000000000000000000000000000000000 +
                skoXC1 *
                (-16009420381169490268274955468750000000000000000000000000000000000000 +
                 skoXC1 *
                 301844280103299764433100722900390625000000000000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_133, refute_0_58])).

cnf(refute_0_135, plain,
    (exp(skoXC1 * -1860000) * -606/25 /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) *
     (21743271936 +
      skoXC1 *
      (-19677661102080000 +
       skoXC1 *
       (8904141648691200000000 +
        skoXC1 *
        (-2686082730688512000000000000 +
         skoXC1 *
         (598230495664865280000000000000000 +
          skoXC1 *
          (-103123542586038681600000000000000000000 +
           skoXC1 *
           (14096236329013464576000000000000000000000000 +
            skoXC1 *
            (-1539649261108625829120000000000000000000000000000 +
             skoXC1 *
             (133739845973308424420100000000000000000000000000000000 +
              skoXC1 *
              (-9069773275373332296913500000000000000000000000000000000000 +
               skoXC1 *
               (459939149072272648591324687500000000000000000000000000000000000 +
                skoXC1 *
                (-16009420381169490268274955468750000000000000000000000000000000000000 +
                 skoXC1 *
                 301844280103299764433100722900390625000000000000000000000000000000000000)))))))))))) <=
     21743271936 |
     21743271936 +
     skoXC1 *
     (-19677661102080000 +
      skoXC1 *
      (8904141648691200000000 +
       skoXC1 *
       (-2686082730688512000000000000 +
        skoXC1 *
        (598230495664865280000000000000000 +
         skoXC1 *
         (-103123542586038681600000000000000000000 +
          skoXC1 *
          (14096236329013464576000000000000000000000000 +
           skoXC1 *
           (-1539649261108625829120000000000000000000000000000 +
            skoXC1 *
            (133739845973308424420100000000000000000000000000000000 +
             skoXC1 *
             (-9069773275373332296913500000000000000000000000000000000000 +
              skoXC1 *
              (459939149072272648591324687500000000000000000000000000000000000 +
               skoXC1 *
               (-16009420381169490268274955468750000000000000000000000000000000000000 +
                skoXC1 *
                301844280103299764433100722900390625000000000000000000000000000000000000))))))))))) <=
     0), inference(resolve, [], [refute_0_134, refute_0_7])).

cnf(refute_0_136, plain,
    (skoXC1 *
     (-19677661102080000 +
      skoXC1 *
      (8904141648691200000000 +
       skoXC1 *
       (-2686082730688512000000000000 +
        skoXC1 *
        (598230495664865280000000000000000 +
         skoXC1 *
         (-103123542586038681600000000000000000000 +
          skoXC1 *
          (14096236329013464576000000000000000000000000 +
           skoXC1 *
           (-1539649261108625829120000000000000000000000000000 +
            skoXC1 *
            (133739845973308424420100000000000000000000000000000000 +
             skoXC1 *
             (-9069773275373332296913500000000000000000000000000000000000 +
              skoXC1 *
              (459939149072272648591324687500000000000000000000000000000000000 +
               skoXC1 *
               (-16009420381169490268274955468750000000000000000000000000000000000000 +
                skoXC1 *
                301844280103299764433100722900390625000000000000000000000000000000000000))))))))))) <=
     -21743271936 |
     exp(skoXC1 * -1860000) *
     (-13176422793216/25 +
      skoXC1 *
      (476986505114419200 +
       skoXC1 *
       (-215836393564274688000000 +
        skoXC1 *
        (65110645391889530880000000000 +
         skoXC1 *
         (-14501107214916334387200000000000000 +
          skoXC1 *
          (2499714672285577641984000000000000000000 +
           skoXC1 *
           (-341692768615286381322240000000000000000000000 +
            skoXC1 *
            (37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (-3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (-11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 -7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936), inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain,
    (-21743271936 <
     skoXC1 *
     (-19677661102080000 +
      skoXC1 *
      (8904141648691200000000 +
       skoXC1 *
       (-2686082730688512000000000000 +
        skoXC1 *
        (598230495664865280000000000000000 +
         skoXC1 *
         (-103123542586038681600000000000000000000 +
          skoXC1 *
          (14096236329013464576000000000000000000000000 +
           skoXC1 *
           (-1539649261108625829120000000000000000000000000000 +
            skoXC1 *
            (133739845973308424420100000000000000000000000000000000 +
             skoXC1 *
             (-9069773275373332296913500000000000000000000000000000000000 +
              skoXC1 *
              (459939149072272648591324687500000000000000000000000000000000000 +
               skoXC1 *
               (-16009420381169490268274955468750000000000000000000000000000000000000 +
                skoXC1 *
                301844280103299764433100722900390625000000000000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_132, refute_0_114, refute_0_113, refute_0_49,
               refute_0_39, refute_0_83, refute_0_26])).

cnf(refute_0_138, plain,
    (exp(skoXC1 * -1860000) *
     (-13176422793216/25 +
      skoXC1 *
      (476986505114419200 +
       skoXC1 *
       (-215836393564274688000000 +
        skoXC1 *
        (65110645391889530880000000000 +
         skoXC1 *
         (-14501107214916334387200000000000000 +
          skoXC1 *
          (2499714672285577641984000000000000000000 +
           skoXC1 *
           (-341692768615286381322240000000000000000000000 +
            skoXC1 *
            (37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (-3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (-11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 -7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) /
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     21743271936), inference(resolve, [], [refute_0_136, refute_0_137])).

cnf(refute_0_139, plain,
    (0 <=
     -527/25 +
     skoXC1 *
     (110158200 +
      skoXC1 *
      (30820014000000 +
       skoXC1 *
       (-53685598770000000000 + skoXC1 * -7510066911450000000000000))) |
     21743271936 *
     (-527/25 +
      skoXC1 *
      (110158200 +
       skoXC1 *
       (30820014000000 +
        skoXC1 *
        (-53685598770000000000 + skoXC1 * -7510066911450000000000000)))) <=
     exp(skoXC1 * -1860000) *
     (-13176422793216/25 +
      skoXC1 *
      (476986505114419200 +
       skoXC1 *
       (-215836393564274688000000 +
        skoXC1 *
        (65110645391889530880000000000 +
         skoXC1 *
         (-14501107214916334387200000000000000 +
          skoXC1 *
          (2499714672285577641984000000000000000000 +
           skoXC1 *
           (-341692768615286381322240000000000000000000000 +
            skoXC1 *
            (37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (-3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (-11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 -7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_138, refute_0_6])).

cnf(refute_0_140, plain,
    (exp(skoXC1 * -1860000) *
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <=
     11458704310272/25 +
     skoXC1 *
     (-2395199698580275200 +
      skoXC1 *
      (-670127945473327104000000 +
       skoXC1 *
       (1167300573103097118720000000000 +
        skoXC1 * 163293427113212982067200000000000000))) |
     skoXC1 *
     (-110158200 +
      skoXC1 *
      (-30820014000000 +
       skoXC1 *
       (53685598770000000000 + skoXC1 * 7510066911450000000000000))) <=
     -527/25), inference(arithmetic, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (exp(skoXC1 * -1860000) *
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <=
     11458704310272/25 +
     skoXC1 *
     (-2395199698580275200 +
      skoXC1 *
      (-670127945473327104000000 +
       skoXC1 *
       (1167300573103097118720000000000 +
        skoXC1 * 163293427113212982067200000000000000)))),
    inference(resolve, [], [refute_0_140, refute_0_50])).

cnf(refute_0_142, plain,
    (13176422793216/25 +
     skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     0 |
     exp(skoXC1 * -1860000) <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(split, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     -13176422793216/25 |
     exp(skoXC1 * -1860000) <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (-13176422793216/25 <
     skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))),
    inference(decision, [],
              [refute_0_132, refute_0_114, refute_0_113, refute_0_49,
               refute_0_39, refute_0_83, refute_0_26])).

cnf(refute_0_145, plain,
    (exp(skoXC1 * -1860000) <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain,
    (skoXC1 * -1860000 < -1 |
     1 + skoXC1 * -1860000 <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_145, refute_0_5])).

cnf(refute_0_147, plain,
    (1/1860000 < skoXC1 |
     1 + skoXC1 * -1860000 <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_146])).

cnf(refute_0_148, plain, (skoXC1 <= 1/1860000),
    inference(decision, [],
              [refute_0_132, refute_0_114, refute_0_113, refute_0_49,
               refute_0_39, refute_0_83, refute_0_26])).

cnf(refute_0_149, plain,
    (1 + skoXC1 * -1860000 <=
     (11458704310272/25 +
      skoXC1 *
      (-2395199698580275200 +
       skoXC1 *
       (-670127945473327104000000 +
        skoXC1 *
        (1167300573103097118720000000000 +
         skoXC1 * 163293427113212982067200000000000000)))) /
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_148, refute_0_147])).

cnf(refute_0_150, plain,
    ((1 + skoXC1 * -1860000) *
     (13176422793216/25 +
      skoXC1 *
      (-476986505114419200 +
       skoXC1 *
       (215836393564274688000000 +
        skoXC1 *
        (-65110645391889530880000000000 +
         skoXC1 *
         (14501107214916334387200000000000000 +
          skoXC1 *
          (-2499714672285577641984000000000000000000 +
           skoXC1 *
           (341692768615286381322240000000000000000000000 +
            skoXC1 *
            (-37321098089273090097868800000000000000000000000000 +
             skoXC1 *
             (3241853866392996207943224000000000000000000000000000000 +
              skoXC1 *
              (-219851304195049574877183240000000000000000000000000000000000 +
               skoXC1 *
               (11148924973511889001853710425000000000000000000000000000000000000 +
                skoXC1 *
                (-388068350039548444102984920562500000000000000000000000000000000000000 +
                 skoXC1 *
                 7316705349703986289858361523105468750000000000000000000000000000000000000)))))))))))) <=
     11458704310272/25 +
     skoXC1 *
     (-2395199698580275200 +
      skoXC1 *
      (-670127945473327104000000 +
       skoXC1 *
       (1167300573103097118720000000000 +
        skoXC1 * 163293427113212982067200000000000000))) |
     13176422793216/25 +
     skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     0), inference(resolve, [], [refute_0_149, refute_0_0])).

cnf(refute_0_151, plain,
    (skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     -13176422793216/25 |
     skoXC1 *
     (937887337650585600 +
      skoXC1 *
      (1773159238550421504000000 +
       skoXC1 *
       (-1633866910524537569280000000000 +
        skoXC1 *
        (-27686519469382120243200000000000000 +
         skoXC1 *
         (-29471774092029959602176000000000000000000 +
          skoXC1 *
          (4991162059066460795412480000000000000000000000 +
           skoXC1 *
           (-672869647713705759357235200000000000000000000000000 +
            skoXC1 *
            (72659096312440943789979192000000000000000000000000000000 +
             skoXC1 *
             (-6249699495686022521651579880000000000000000000000000000000000 +
              skoXC1 *
              (420072350776304098273414536825000000000000000000000000000000000000 +
               skoXC1 *
               (-21125068800771661987550886311062500000000000000000000000000000000000000 +
                skoXC1 *
                (729123836423264092321410313769355468750000000000000000000000000000000000000 +
                 skoXC1 *
                 -13609071950449414499136552432976171875000000000000000000000000000000000000000000)))))))))))) <=
     -1717718482944/25), inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (-1717718482944/25 <
     skoXC1 *
     (937887337650585600 +
      skoXC1 *
      (1773159238550421504000000 +
       skoXC1 *
       (-1633866910524537569280000000000 +
        skoXC1 *
        (-27686519469382120243200000000000000 +
         skoXC1 *
         (-29471774092029959602176000000000000000000 +
          skoXC1 *
          (4991162059066460795412480000000000000000000000 +
           skoXC1 *
           (-672869647713705759357235200000000000000000000000000 +
            skoXC1 *
            (72659096312440943789979192000000000000000000000000000000 +
             skoXC1 *
             (-6249699495686022521651579880000000000000000000000000000000000 +
              skoXC1 *
              (420072350776304098273414536825000000000000000000000000000000000000 +
               skoXC1 *
               (-21125068800771661987550886311062500000000000000000000000000000000000000 +
                skoXC1 *
                (729123836423264092321410313769355468750000000000000000000000000000000000000 +
                 skoXC1 *
                 -13609071950449414499136552432976171875000000000000000000000000000000000000000000)))))))))))) |
     skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     -13176422793216/25),
    inference(decision, [],
              [refute_0_132, refute_0_114, refute_0_113, refute_0_49,
               refute_0_39, refute_0_83, refute_0_26])).

cnf(refute_0_153, plain,
    (skoXC1 *
     (-476986505114419200 +
      skoXC1 *
      (215836393564274688000000 +
       skoXC1 *
       (-65110645391889530880000000000 +
        skoXC1 *
        (14501107214916334387200000000000000 +
         skoXC1 *
         (-2499714672285577641984000000000000000000 +
          skoXC1 *
          (341692768615286381322240000000000000000000000 +
           skoXC1 *
           (-37321098089273090097868800000000000000000000000000 +
            skoXC1 *
            (3241853866392996207943224000000000000000000000000000000 +
             skoXC1 *
             (-219851304195049574877183240000000000000000000000000000000000 +
              skoXC1 *
              (11148924973511889001853710425000000000000000000000000000000000000 +
               skoXC1 *
               (-388068350039548444102984920562500000000000000000000000000000000000000 +
                skoXC1 *
                7316705349703986289858361523105468750000000000000000000000000000000000000))))))))))) <=
     -13176422793216/25),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain, ($false),
    inference(resolve, [], [refute_0_153, refute_0_144])).
% SZS output end CNFRefutation for Colpitts-BJT-2-VC2+VC1.tptp
