% SZS output start CNFRefutation for CONVOI.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(sin_upper_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, X - X ^ 3 / 6, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

fof('CONVOI', conjecture,
    (! [X] :
       (0 <= X =>
        0 <
        12 - 71/5 * exp(-159/500 * X) +
        (13/4 * cos(29/25 * X) - 31/200 * sin(29/25 * X)) *
        exp(-67/50 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        0 <
        12 - 71/5 * exp(-159/500 * X) +
        (13/4 * cos(29/25 * X) - 31/200 * sin(29/25 * X)) *
        exp(-67/50 * X))), inference(strip, [], ['CONVOI'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          0 <
          12 - 71/5 * exp(-159/500 * X) +
          (13/4 * cos(29/25 * X) - 31/200 * sin(29/25 * X)) *
          exp(-67/50 * X))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        12 - 71/5 * exp(-159/500 * X) +
        (13/4 * cos(29/25 * X) - 31/200 * sin(29/25 * X)) *
        exp(-67/50 * X) <= 0)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     12 - 71/5 * exp(-159/500 * skoXC1) +
     (13/4 * cos(29/25 * skoXC1) - 31/200 * sin(29/25 * skoXC1)) *
     exp(-67/50 * skoXC1) <= 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (12 - 71/5 * exp(-159/500 * skoXC1) +
     (13/4 * cos(29/25 * skoXC1) - 31/200 * sin(29/25 * skoXC1)) *
     exp(-67/50 * skoXC1) <= 0), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (2304 <
     (58392576/71 +
      skoXC1 *
      (-575963136/1775 +
       skoXC1 *
       (-310627270656/221875 +
        skoXC1 *
        (-35470156511232/27734375 +
         skoXC1 *
         (-2471558410610352/3466796875 +
          skoXC1 *
          (-114422930315294436/433349609375 +
           skoXC1 *
           (-14721538578463263627/216674804687500 +
            skoXC1 *
            (-132781241741472160389/10833740234375000 +
             skoXC1 *
             (-2903074685934985879641/2166748046875000000 +
              skoXC1 *
              (-23354642279990984757993/216674804687500000000 +
               skoXC1 *
               (-532922371819915718143683/86669921875000000000000 +
                skoXC1 *
                (-1023662219745311879046273/4333496093750000000000000 +
                 skoXC1 *
                 -4384830569741378145347787/
                 866699218750000000000000000)))))))))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) |
     0 <=
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     2304 *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     58392576/71 +
     skoXC1 *
     (-575963136/1775 +
      skoXC1 *
      (-310627270656/221875 +
       skoXC1 *
       (-35470156511232/27734375 +
        skoXC1 *
        (-2471558410610352/3466796875 +
         skoXC1 *
         (-114422930315294436/433349609375 +
          skoXC1 *
          (-14721538578463263627/216674804687500 +
           skoXC1 *
           (-132781241741472160389/10833740234375000 +
            skoXC1 *
            (-2903074685934985879641/2166748046875000000 +
             skoXC1 *
             (-23354642279990984757993/216674804687500000000 +
              skoXC1 *
              (-532922371819915718143683/86669921875000000000000 +
               skoXC1 *
               (-1023662219745311879046273/4333496093750000000000000 +
                skoXC1 *
                -4384830569741378145347787/
                866699218750000000000000000)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (77184/25 +
       skoXC1 *
       (1292832/625 +
        skoXC1 *
        (14436624/15625 +
         skoXC1 *
         (423173541/1562500 +
          skoXC1 *
          (4050375321/78125000 + skoXC1 * 90458382169/15625000000)))))) <
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) *
     (2304 +
      skoXC1 *
      (77184/25 +
       skoXC1 *
       (1292832/625 +
        skoXC1 *
        (14436624/15625 +
         skoXC1 *
         (423173541/1562500 +
          skoXC1 *
          (4050375321/78125000 + skoXC1 * 90458382169/15625000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (77184/25 +
      skoXC1 *
      (1292832/625 +
       skoXC1 *
       (14436624/15625 +
        skoXC1 *
        (423173541/1562500 +
         skoXC1 *
         (4050375321/78125000 + skoXC1 * 90458382169/15625000000))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (exp(X_000122) < X_000123 | ~ lgen(1, exp(X_000122), X_000123)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (0 < X_000122 |
     ~ lgen(1,
            2304 /
            (-(X_000122 ^ 3) + 6 * X_000122 ^ 2 - 24 * X_000122 + 48) ^ 2,
            X_000123) | lgen(1, exp(X_000122), X_000123)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_4, plain,
    (0 < X_000122 | exp(X_000122) < X_000123 |
     ~ lgen(1,
            2304 /
            (-(X_000122 ^ 3) + 6 * X_000122 ^ 2 - 24 * X_000122 + 48) ^ 2,
            X_000123)), inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (0 < X_000122 | exp(X_000122) < X_000123 |
     X_000123 <=
     2304 /
     (2304 +
      X_000122 *
      (-2304 +
       X_000122 *
       (1152 +
        X_000122 *
        (-384 + X_000122 * (84 + X_000122 * (-12 + X_000122))))))),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (0 < skoXC1 * -67/50 |
     exp(skoXC1 * -67/50) <
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304 /
     (2304 +
      skoXC1 * -67/50 *
      (-2304 +
       skoXC1 * -67/50 *
       (1152 +
        skoXC1 * -67/50 *
        (-384 +
         skoXC1 * -67/50 *
         (84 + skoXC1 * -67/50 * (-12 + skoXC1 * -67/50))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) |
     (60/71 +
      exp(skoXC1 * -67/50) *
      (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000))) *
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_8, plain,
    ((-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4) <
     1 + skoXC1 * (skoXC1 * -841/1250) |
     (1 + skoXC1 * (skoXC1 * -841/1250)) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
     exp(skoXC1 * -159/500) * 71/5 | exp(skoXC1 * -67/50) * 13/4 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_9, plain,
    (X_000017 < cos(X_000016) | ~ lgen(1, X_000017, cos(X_000016))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_10, plain,
    (~ lgen(1, X_000017, 1 - X_000016 ^ 2 / 2) |
     lgen(1, X_000017, cos(X_000016))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_11, plain,
    (X_000017 < cos(X_000016) | ~ lgen(1, X_000017, 1 - X_000016 ^ 2 / 2)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (X_000017 < cos(X_000016) |
     1 + X_000016 * (X_000016 * -1/2) <= X_000017),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4) <
     cos(skoXC1 * 29/25) |
     1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/2) <=
     (-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 * 29/25 <
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) |
     0 <= exp(skoXC1 * -67/50) * -31/200 |
     skoXC1 * 29/25 * (exp(skoXC1 * -67/50) * -31/200) <=
     -12 + exp(skoXC1 * -159/500) * 71/5 +
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_15, plain,
    (12 - 71/5 * exp(-159/500 * skoXC1) +
     (13/4 * cos(29/25 * skoXC1) - 31/200 * sin(29/25 * skoXC1)) *
     exp(-67/50 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_16, plain,
    (sin(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -31/200) <=
     -12 + exp(skoXC1 * -159/500) * 71/5 +
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (0 <= exp(skoXC1 * -67/50) * -31/200 |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= sin(skoXC1 * 29/25)),
    inference(split, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= sin(skoXC1 * 29/25)),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (sin(X_000087) < X_000088 | ~ lgen(1, sin(X_000087), X_000088)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_20, plain,
    (X_000087 < 0 | ~ lgen(1, X_000087, X_000088) |
     lgen(1, sin(X_000087), X_000088)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_21, plain,
    (sin(X_000087) < X_000088 | X_000087 < 0 |
     ~ lgen(1, X_000087, X_000088)),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (sin(X_000087) < X_000088 | X_000087 < 0 | X_000088 <= X_000087),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 * 29/25 < 0 |
     sin(skoXC1 * 29/25) <
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= skoXC1 * 29/25),
    inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * 29/25 < 0 |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= skoXC1 * 29/25),
    inference(resolve, [], [refute_0_18, refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 < 0 |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= skoXC1 * 29/25),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain,
    ((-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <= skoXC1 * 29/25),
    inference(resolve, [], [refute_0_26, refute_0_25])).

cnf(refute_0_28, plain,
    (0 <= exp(skoXC1 * -67/50) * -31/200 |
     skoXC1 * 29/25 * (exp(skoXC1 * -67/50) * -31/200) <=
     -12 + exp(skoXC1 * -159/500) * 71/5 +
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)),
    inference(resolve, [], [refute_0_27, refute_0_14])).

cnf(refute_0_29, plain,
    (cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
     exp(skoXC1 * -159/500) * 71/5),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (exp(skoXC1 * -67/50) * 13/4 <= 0 |
     cos(skoXC1 * 29/25) <=
     (-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(split, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (cos(skoXC1 * 29/25) <=
     (-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/2) <=
     (-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(resolve, [], [refute_0_31, refute_0_13])).

cnf(refute_0_33, plain,
    (1 + skoXC1 * (skoXC1 * -841/1250) <=
     (-12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    ((1 + skoXC1 * (skoXC1 * -841/1250)) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 + exp(skoXC1 * -67/50) * (skoXC1 * 899/5000) +
     exp(skoXC1 * -159/500) * 71/5 | exp(skoXC1 * -67/50) * 13/4 <= 0),
    inference(resolve, [], [refute_0_33, refute_0_8])).

cnf(refute_0_35, plain,
    (60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) <=
     exp(skoXC1 * -159/500)), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (0 < skoXC1 * -159/500 |
     exp(skoXC1 * -159/500) <
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) <=
     2304 /
     (2304 +
      skoXC1 * -159/500 *
      (-2304 +
       skoXC1 * -159/500 *
       (1152 +
        skoXC1 * -159/500 *
        (-384 +
         skoXC1 * -159/500 *
         (84 + skoXC1 * -159/500 * (-12 + skoXC1 * -159/500))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_37, plain,
    (0 < skoXC1 * -159/500 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) <=
     2304 /
     (2304 +
      skoXC1 * -159/500 *
      (-2304 +
       skoXC1 * -159/500 *
       (1152 +
        skoXC1 * -159/500 *
        (-384 +
         skoXC1 * -159/500 *
         (84 + skoXC1 * -159/500 * (-12 + skoXC1 * -159/500))))))),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 < 0 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) <=
     2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000))))))),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000)) <=
     2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000))))))),
    inference(resolve, [], [refute_0_26, refute_0_38])).

cnf(refute_0_40, plain,
    ((60/71 +
      exp(skoXC1 * -67/50) *
      (65/284 + skoXC1 * (-899/71000 + skoXC1 * -10933/71000))) *
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= 0),
    inference(resolve, [], [refute_0_39, refute_0_7])).

cnf(refute_0_41, plain,
    (exp(skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (-2304 <
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000)))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_43, plain,
    (exp(skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000)))))),
    inference(resolve, [], [refute_0_41, refute_0_42])).

cnf(refute_0_44, plain,
    (0 <=
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     exp(skoXC1 * -67/50)), inference(split, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) <=
     37440/71 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     exp(skoXC1 * -67/50)), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <
     0 |
     0 <
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     0 <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000)))))),
    inference(split, [], [refute_0_43])).

cnf(refute_0_47, plain,
    (37440/71 <
     skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) |
     -37440/71 <
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     skoXC1 *
     (1099008/1775 +
      skoXC1 *
      (21842784/221875 +
       skoXC1 *
       (289416888/27734375 +
        skoXC1 *
        (40265124543/55468750000 +
         skoXC1 *
         (914593543191/27734375000000 +
          skoXC1 * 48473457789123/55468750000000000))))) <= 25344/71),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (37440/71 <
     skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) |
     -37440/71 <
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     25344/71 <
     skoXC1 *
     (1099008/1775 +
      skoXC1 *
      (21842784/221875 +
       skoXC1 *
       (289416888/27734375 +
        skoXC1 *
        (40265124543/55468750000 +
         skoXC1 *
         (914593543191/27734375000000 +
          skoXC1 * 48473457789123/55468750000000000)))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_49, plain,
    (37440/71 <
     skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) |
     -37440/71 <
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000)))))))),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (0 *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     0 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_51, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_52, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_53, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_52, refute_0_51])).

cnf(refute_0_54, plain,
    (0 <
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <
     exp(skoXC1 * -67/50)), inference(subst, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     0 |
     exp(skoXC1 * -67/50) <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(split, [], [refute_0_43])).

cnf(refute_0_56, plain,
    (skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71 |
     exp(skoXC1 * -67/50) <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (0 <
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(resolve, [], [refute_0_56, refute_0_54])).

cnf(refute_0_58, plain,
    (0 *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71 |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     0), inference(resolve, [], [refute_0_50, refute_0_57])).

cnf(refute_0_59, plain,
    (-25344/71 <
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    ((25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <
     1 + skoXC1 * -67/50 |
     (1 + skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_61, plain,
    (X_000102 < exp(X_000101) | ~ lgen(1, X_000102, exp(X_000101))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_62, plain,
    (X_000101 < -1 | ~ lgen(1, X_000102, 1 + X_000101) |
     lgen(1, X_000102, exp(X_000101))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_63, plain,
    (X_000101 < -1 | X_000102 < exp(X_000101) |
     ~ lgen(1, X_000102, 1 + X_000101)),
    inference(resolve, [], [refute_0_62, refute_0_61])).

cnf(refute_0_64, plain,
    (X_000101 < -1 | X_000102 < exp(X_000101) | 1 + X_000101 <= X_000102),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 * -67/50 < -1 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <
     exp(skoXC1 * -67/50) |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(subst, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 * -67/50 < -1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(resolve, [], [refute_0_56, refute_0_65])).

cnf(refute_0_67, plain,
    (50/67 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    ((25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <
     1 + skoXC1 * -67/50 |
     (1 + skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_69, plain,
    (skoXC1 * -67/50 < -1 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <
     exp(skoXC1 * -67/50) |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(subst, [], [refute_0_64])).

cnf(refute_0_70, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 + skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) |
     (60/71 +
      exp(skoXC1 * -67/50) *
      (65/284 +
       skoXC1 *
       (-899/71000 +
        skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000)))) *
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_71, plain,
    ((-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4) <
     1 + skoXC1 * (skoXC1 * -841/1250) |
     (1 + skoXC1 * (skoXC1 * -841/1250)) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 +
     exp(skoXC1 * -67/50) *
     (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
     exp(skoXC1 * -159/500) * 71/5 | exp(skoXC1 * -67/50) * 13/4 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_72, plain,
    ((-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4) <
     cos(skoXC1 * 29/25) |
     1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/2) <=
     (-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_73, plain,
    (skoXC1 * (29/25 + skoXC1 * (skoXC1 * -24389/93750)) <
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) |
     0 <= exp(skoXC1 * -67/50) * -31/200 |
     skoXC1 * (29/25 + skoXC1 * (skoXC1 * -24389/93750)) *
     (exp(skoXC1 * -67/50) * -31/200) <=
     -12 + exp(skoXC1 * -159/500) * 71/5 +
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_74, plain,
    (sin(X_000108) < X_000109 | ~ lgen(1, sin(X_000108), X_000109)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_75, plain,
    (0 < X_000108 | ~ lgen(1, X_000108 - X_000108 ^ 3 / 6, X_000109) |
     lgen(1, sin(X_000108), X_000109)),
    inference(subst, [], [sin_upper_bound_0_neg])).

cnf(refute_0_76, plain,
    (0 < X_000108 | sin(X_000108) < X_000109 |
     ~ lgen(1, X_000108 - X_000108 ^ 3 / 6, X_000109)),
    inference(resolve, [], [refute_0_75, refute_0_74])).

cnf(refute_0_77, plain,
    (0 < X_000108 | sin(X_000108) < X_000109 |
     X_000109 <= X_000108 * (1 + X_000108 * (X_000108 * -1/6))),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (0 < skoXC1 * 29/25 |
     sin(skoXC1 * 29/25) <
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <=
     skoXC1 * 29/25 * (1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/6))),
    inference(subst, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (0 < skoXC1 * 29/25 |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <=
     skoXC1 * 29/25 * (1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/6))),
    inference(resolve, [], [refute_0_18, refute_0_78])).

cnf(refute_0_80, plain,
    (0 < skoXC1 |
     (-12 + exp(skoXC1 * -159/500) * 71/5 +
      cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)) /
     (exp(skoXC1 * -67/50) * -31/200) <=
     skoXC1 * (29/25 + skoXC1 * (skoXC1 * -24389/93750))),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (0 < skoXC1 | 0 <= exp(skoXC1 * -67/50) * -31/200 |
     skoXC1 * (29/25 + skoXC1 * (skoXC1 * -24389/93750)) *
     (exp(skoXC1 * -67/50) * -31/200) <=
     -12 + exp(skoXC1 * -159/500) * 71/5 +
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * -13/4)),
    inference(resolve, [], [refute_0_80, refute_0_73])).

cnf(refute_0_82, plain,
    (0 < skoXC1 |
     cos(skoXC1 * 29/25) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 +
     exp(skoXC1 * -67/50) *
     (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
     exp(skoXC1 * -159/500) * 71/5),
    inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (0 < skoXC1 | exp(skoXC1 * -67/50) * 13/4 <= 0 |
     cos(skoXC1 * 29/25) <=
     (-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(split, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (0 < skoXC1 |
     cos(skoXC1 * 29/25) <=
     (-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (0 < skoXC1 |
     1 + skoXC1 * 29/25 * (skoXC1 * 29/25 * -1/2) <=
     (-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(resolve, [], [refute_0_84, refute_0_72])).

cnf(refute_0_86, plain,
    (0 < skoXC1 |
     1 + skoXC1 * (skoXC1 * -841/1250) <=
     (-12 +
      exp(skoXC1 * -67/50) *
      (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
      exp(skoXC1 * -159/500) * 71/5) / (exp(skoXC1 * -67/50) * 13/4)),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 < skoXC1 |
     (1 + skoXC1 * (skoXC1 * -841/1250)) * (exp(skoXC1 * -67/50) * 13/4) <=
     -12 +
     exp(skoXC1 * -67/50) *
     (skoXC1 * (899/5000 + skoXC1 * (skoXC1 * -756059/18750000))) +
     exp(skoXC1 * -159/500) * 71/5 | exp(skoXC1 * -67/50) * 13/4 <= 0),
    inference(resolve, [], [refute_0_86, refute_0_71])).

cnf(refute_0_88, plain,
    (0 < skoXC1 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 +
       skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) <=
     exp(skoXC1 * -159/500)), inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (0 < skoXC1 * -159/500 |
     exp(skoXC1 * -159/500) <
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 + skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 +
       skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) <=
     2304 /
     (2304 +
      skoXC1 * -159/500 *
      (-2304 +
       skoXC1 * -159/500 *
       (1152 +
        skoXC1 * -159/500 *
        (-384 +
         skoXC1 * -159/500 *
         (84 + skoXC1 * -159/500 * (-12 + skoXC1 * -159/500))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_90, plain,
    (0 < skoXC1 * -159/500 | 0 < skoXC1 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 +
       skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) <=
     2304 /
     (2304 +
      skoXC1 * -159/500 *
      (-2304 +
       skoXC1 * -159/500 *
       (1152 +
        skoXC1 * -159/500 *
        (-384 +
         skoXC1 * -159/500 *
         (84 + skoXC1 * -159/500 * (-12 + skoXC1 * -159/500))))))),
    inference(resolve, [], [refute_0_88, refute_0_89])).

cnf(refute_0_91, plain,
    (skoXC1 < 0 | 0 < skoXC1 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 +
       skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) <=
     2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000))))))),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain, (0 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_93, plain,
    (0 < skoXC1 |
     60/71 +
     exp(skoXC1 * -67/50) *
     (65/284 +
      skoXC1 *
      (-899/71000 +
       skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000))) <=
     2304 /
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000))))))),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (0 < skoXC1 |
     (60/71 +
      exp(skoXC1 * -67/50) *
      (65/284 +
       skoXC1 *
       (-899/71000 +
        skoXC1 * (-10933/71000 + skoXC1 * 756059/266250000)))) *
     (2304 +
      skoXC1 *
      (91584/125 +
       skoXC1 *
       (1820232/15625 +
        skoXC1 *
        (24118074/1953125 +
         skoXC1 *
         (13421708181/15625000000 +
          skoXC1 *
          (304864514397/7812500000000 +
           skoXC1 * 16157819263041/15625000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= 0),
    inference(resolve, [], [refute_0_93, refute_0_70])).

cnf(refute_0_95, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (-2304 <
     skoXC1 *
     (91584/125 +
      skoXC1 *
      (1820232/15625 +
       skoXC1 *
       (24118074/1953125 +
        skoXC1 *
        (13421708181/15625000000 +
         skoXC1 *
         (304864514397/7812500000000 +
          skoXC1 * 16157819263041/15625000000000000))))) | 0 < skoXC1),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_97, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000)))))),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (0 < skoXC1 |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <= 0 |
     exp(skoXC1 * -67/50) <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(split, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (0 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <=
     -37440/71 |
     exp(skoXC1 * -67/50) <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (-37440/71 <
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) |
     0 < skoXC1), inference(decision, [], [refute_0_26])).

cnf(refute_0_101, plain,
    (0 < skoXC1 |
     exp(skoXC1 * -67/50) <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain,
    (skoXC1 * -67/50 < -1 | 0 < skoXC1 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(resolve, [], [refute_0_101, refute_0_69])).

cnf(refute_0_103, plain,
    (0 < skoXC1 | 50/67 < skoXC1 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain, (0 < skoXC1 | skoXC1 <= 50/67),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_105, plain,
    (0 < skoXC1 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000)))))))))),
    inference(resolve, [], [refute_0_104, refute_0_103])).

cnf(refute_0_106, plain,
    (0 < skoXC1 |
     (1 + skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-29100964137/277343750 +
         skoXC1 *
         (-70191415884267/4437500000000 +
          skoXC1 *
          (-3489267667512033/2218750000000000 +
           skoXC1 *
           (-432497807871814359/4437500000000000000 +
            skoXC1 *
            (-3974665813823961459/1109375000000000000000 +
             skoXC1 *
             (-268612363639519137/5546875000000000000000 +
              skoXC1 *
              4072088224731838473/1386718750000000000000000))))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_105, refute_0_68])).

cnf(refute_0_107, plain,
    (0 < skoXC1 |
     skoXC1 *
     (453168/8875 +
      skoXC1 *
      (-470999502/1109375 +
       skoXC1 *
       (99184033713/277343750 +
        skoXC1 *
        (556954465176453/4437500000000 +
         skoXC1 *
         (43612148458402137/2218750000000000 +
          skoXC1 *
          (8922617417683563921/4437500000000000000 +
           skoXC1 *
           (70456049911616924403/554687500000000000000 +
            skoXC1 *
            (263616485889810226383/55468750000000000000000 +
             skoXC1 *
             (11757153755496343671/173339843750000000000000 +
              skoXC1 *
              -272829911057033177691/69335937500000000000000000))))))))) <=
     -12096/71 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <=
     -37440/71), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (-12096/71 <
     skoXC1 *
     (453168/8875 +
      skoXC1 *
      (-470999502/1109375 +
       skoXC1 *
       (99184033713/277343750 +
        skoXC1 *
        (556954465176453/4437500000000 +
         skoXC1 *
         (43612148458402137/2218750000000000 +
          skoXC1 *
          (8922617417683563921/4437500000000000000 +
           skoXC1 *
           (70456049911616924403/554687500000000000000 +
            skoXC1 *
            (263616485889810226383/55468750000000000000000 +
             skoXC1 *
             (11757153755496343671/173339843750000000000000 +
              skoXC1 *
              -272829911057033177691/69335937500000000000000000))))))))) |
     0 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <=
     -37440/71), inference(decision, [], [refute_0_26])).

cnf(refute_0_109, plain,
    (0 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-29100964137/277343750 +
        skoXC1 *
        (-70191415884267/4437500000000 +
         skoXC1 *
         (-3489267667512033/2218750000000000 +
          skoXC1 *
          (-432497807871814359/4437500000000000000 +
           skoXC1 *
           (-3974665813823961459/1109375000000000000000 +
            skoXC1 *
            (-268612363639519137/5546875000000000000000 +
             skoXC1 *
             4072088224731838473/1386718750000000000000000)))))))) <=
     -37440/71), inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_109, refute_0_100])).

cnf(refute_0_111, plain,
    (-37440/71 <
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     50/67 < skoXC1),
    inference(decision, [], [refute_0_110, refute_0_26])).

cnf(refute_0_112, plain,
    (50/67 < skoXC1 |
     1 + skoXC1 * -67/50 <=
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000))))))))),
    inference(resolve, [], [refute_0_67, refute_0_111])).

cnf(refute_0_113, plain,
    (50/67 < skoXC1 |
     (1 + skoXC1 * -67/50) *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     25344/71 +
     skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) |
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     0), inference(resolve, [], [refute_0_112, refute_0_60])).

cnf(refute_0_114, plain,
    (50/67 < skoXC1 |
     skoXC1 *
     (453168/8875 +
      skoXC1 *
      (-470999502/1109375 +
       skoXC1 *
       (97369492113/277343750 +
        skoXC1 *
        (586625849419653/4437500000000 +
         skoXC1 *
         (49063873372104537/2218750000000000 +
          skoXC1 *
          (10734067271058488721/4437500000000000000 +
           skoXC1 *
           (95166481966750903203/554687500000000000000 +
            skoXC1 *
            (438773782638451119903/55468750000000000000000 +
             skoXC1 *
             11835780346189425951/55468750000000000000000)))))))) <=
     -12096/71 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (-12096/71 <
     skoXC1 *
     (453168/8875 +
      skoXC1 *
      (-470999502/1109375 +
       skoXC1 *
       (97369492113/277343750 +
        skoXC1 *
        (586625849419653/4437500000000 +
         skoXC1 *
         (49063873372104537/2218750000000000 +
          skoXC1 *
          (10734067271058488721/4437500000000000000 +
           skoXC1 *
           (95166481966750903203/554687500000000000000 +
            skoXC1 *
            (438773782638451119903/55468750000000000000000 +
             skoXC1 *
             11835780346189425951/55468750000000000000000)))))))) |
     50/67 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(decision, [], [refute_0_110, refute_0_26])).

cnf(refute_0_116, plain,
    (50/67 < skoXC1 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain, (50/67 < skoXC1),
    inference(resolve, [], [refute_0_116, refute_0_111])).

cnf(refute_0_118, plain,
    (skoXC1 *
     (-1099008/1775 +
      skoXC1 *
      (-21842784/221875 +
       skoXC1 *
       (-289416888/27734375 +
        skoXC1 *
        (-40265124543/55468750000 +
         skoXC1 *
         (-914593543191/27734375000000 +
          skoXC1 * -48473457789123/55468750000000000))))) <= -25344/71 |
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71),
    inference(decision, [], [refute_0_117, refute_0_110, refute_0_26])).

cnf(refute_0_119, plain,
    (skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) <=
     -37440/71), inference(resolve, [], [refute_0_118, refute_0_59])).

cnf(refute_0_120, plain,
    (37440/71 <
     skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000)))))))),
    inference(resolve, [], [refute_0_119, refute_0_49])).

cnf(refute_0_121, plain,
    ((25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     exp(skoXC1 * -67/50)),
    inference(resolve, [], [refute_0_45, refute_0_120])).

cnf(refute_0_122, plain,
    (0 < skoXC1 * -67/50 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304 /
     (2304 +
      skoXC1 * -67/50 *
      (-2304 +
       skoXC1 * -67/50 *
       (1152 +
        skoXC1 * -67/50 *
        (-384 +
         skoXC1 * -67/50 *
         (84 + skoXC1 * -67/50 * (-12 + skoXC1 * -67/50))))))),
    inference(resolve, [], [refute_0_121, refute_0_6])).

cnf(refute_0_123, plain,
    (skoXC1 < 0 |
     (25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (77184/25 +
       skoXC1 *
       (1292832/625 +
        skoXC1 *
        (14436624/15625 +
         skoXC1 *
         (423173541/1562500 +
          skoXC1 *
          (4050375321/78125000 + skoXC1 * 90458382169/15625000000))))))),
    inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    ((25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (77184/25 +
       skoXC1 *
       (1292832/625 +
        skoXC1 *
        (14436624/15625 +
         skoXC1 *
         (423173541/1562500 +
          skoXC1 *
          (4050375321/78125000 + skoXC1 * 90458382169/15625000000))))))),
    inference(resolve, [], [refute_0_26, refute_0_123])).

cnf(refute_0_125, plain,
    ((25344/71 +
      skoXC1 *
      (-1099008/1775 +
       skoXC1 *
       (-21842784/221875 +
        skoXC1 *
        (-289416888/27734375 +
         skoXC1 *
         (-40265124543/55468750000 +
          skoXC1 *
          (-914593543191/27734375000000 +
           skoXC1 * -48473457789123/55468750000000000)))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) *
     (2304 +
      skoXC1 *
      (77184/25 +
       skoXC1 *
       (1292832/625 +
        skoXC1 *
        (14436624/15625 +
         skoXC1 *
         (423173541/1562500 +
          skoXC1 *
          (4050375321/78125000 + skoXC1 * 90458382169/15625000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (77184/25 +
      skoXC1 *
      (1292832/625 +
       skoXC1 *
       (14436624/15625 +
        skoXC1 *
        (423173541/1562500 +
         skoXC1 *
         (4050375321/78125000 + skoXC1 * 90458382169/15625000000))))) <=
     0), inference(resolve, [], [refute_0_124, refute_0_1])).

cnf(refute_0_126, plain,
    (skoXC1 *
     (77184/25 +
      skoXC1 *
      (1292832/625 +
       skoXC1 *
       (14436624/15625 +
        skoXC1 *
        (423173541/1562500 +
         skoXC1 *
         (4050375321/78125000 + skoXC1 * 90458382169/15625000000))))) <=
     -2304 |
     (58392576/71 +
      skoXC1 *
      (-575963136/1775 +
       skoXC1 *
       (-310627270656/221875 +
        skoXC1 *
        (-35470156511232/27734375 +
         skoXC1 *
         (-2471558410610352/3466796875 +
          skoXC1 *
          (-114422930315294436/433349609375 +
           skoXC1 *
           (-14721538578463263627/216674804687500 +
            skoXC1 *
            (-132781241741472160389/10833740234375000 +
             skoXC1 *
             (-2903074685934985879641/2166748046875000000 +
              skoXC1 *
              (-23354642279990984757993/216674804687500000000 +
               skoXC1 *
               (-532922371819915718143683/86669921875000000000000 +
                skoXC1 *
                (-1023662219745311879046273/4333496093750000000000000 +
                 skoXC1 *
                 -4384830569741378145347787/
                 866699218750000000000000000)))))))))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304), inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (-2304 <
     skoXC1 *
     (77184/25 +
      skoXC1 *
      (1292832/625 +
       skoXC1 *
       (14436624/15625 +
        skoXC1 *
        (423173541/1562500 +
         skoXC1 *
         (4050375321/78125000 + skoXC1 * 90458382169/15625000000)))))),
    inference(decision, [],
              [refute_0_120, refute_0_119, refute_0_117, refute_0_110,
               refute_0_26])).

cnf(refute_0_128, plain,
    ((58392576/71 +
      skoXC1 *
      (-575963136/1775 +
       skoXC1 *
       (-310627270656/221875 +
        skoXC1 *
        (-35470156511232/27734375 +
         skoXC1 *
         (-2471558410610352/3466796875 +
          skoXC1 *
          (-114422930315294436/433349609375 +
           skoXC1 *
           (-14721538578463263627/216674804687500 +
            skoXC1 *
            (-132781241741472160389/10833740234375000 +
             skoXC1 *
             (-2903074685934985879641/2166748046875000000 +
              skoXC1 *
              (-23354642279990984757993/216674804687500000000 +
               skoXC1 *
               (-532922371819915718143683/86669921875000000000000 +
                skoXC1 *
                (-1023662219745311879046273/4333496093750000000000000 +
                 skoXC1 *
                 -4384830569741378145347787/
                 866699218750000000000000000)))))))))))) /
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     2304), inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain,
    (0 <=
     37440/71 +
     skoXC1 *
     (1229328/8875 +
      skoXC1 *
      (-374300982/1109375 +
       skoXC1 *
       (-30915505737/277343750 +
        skoXC1 *
        (-79423803545067/4437500000000 +
         skoXC1 *
         (-4223242486545633/2218750000000000 +
          skoXC1 *
          (-588100469506937559/4437500000000000000 +
           skoXC1 *
           (-6680693351322275859/1109375000000000000000 +
            skoXC1 * -176653438002827253/1109375000000000000000))))))) |
     2304 *
     (37440/71 +
      skoXC1 *
      (1229328/8875 +
       skoXC1 *
       (-374300982/1109375 +
        skoXC1 *
        (-30915505737/277343750 +
         skoXC1 *
         (-79423803545067/4437500000000 +
          skoXC1 *
          (-4223242486545633/2218750000000000 +
           skoXC1 *
           (-588100469506937559/4437500000000000000 +
            skoXC1 *
            (-6680693351322275859/1109375000000000000000 +
             skoXC1 * -176653438002827253/1109375000000000000000)))))))) <=
     58392576/71 +
     skoXC1 *
     (-575963136/1775 +
      skoXC1 *
      (-310627270656/221875 +
       skoXC1 *
       (-35470156511232/27734375 +
        skoXC1 *
        (-2471558410610352/3466796875 +
         skoXC1 *
         (-114422930315294436/433349609375 +
          skoXC1 *
          (-14721538578463263627/216674804687500 +
           skoXC1 *
           (-132781241741472160389/10833740234375000 +
            skoXC1 *
            (-2903074685934985879641/2166748046875000000 +
             skoXC1 *
             (-23354642279990984757993/216674804687500000000 +
              skoXC1 *
              (-532922371819915718143683/86669921875000000000000 +
               skoXC1 *
               (-1023662219745311879046273/4333496093750000000000000 +
                skoXC1 *
                -4384830569741378145347787/
                866699218750000000000000000)))))))))))),
    inference(resolve, [], [refute_0_128, refute_0_0])).

cnf(refute_0_130, plain,
    (skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) <=
     37440/71 |
     skoXC1 *
     (5712187392/8875 +
      skoXC1 *
      (690746890752/1109375 +
       skoXC1 *
       (141736119947136/138671875 +
        skoXC1 *
        (11642977821146157/17333984375 +
         skoXC1 *
         (2250449423926978023/8666992187500 +
          skoXC1 *
          (1172430182051498652129/17333984375000000 +
           skoXC1 *
           (53052370456426963672869/4333496093750000000 +
            skoXC1 *
            (1160911898185589262801/866699218750000000 +
             skoXC1 *
             (23354642279990984757993/216674804687500000000 +
              skoXC1 *
              (532922371819915718143683/86669921875000000000000 +
               skoXC1 *
               (1023662219745311879046273/4333496093750000000000000 +
                skoXC1 *
                4384830569741378145347787/
                866699218750000000000000000))))))))))) <= -27869184/71),
    inference(arithmetic, [], [refute_0_129])).

cnf(refute_0_131, plain,
    (-27869184/71 <
     skoXC1 *
     (5712187392/8875 +
      skoXC1 *
      (690746890752/1109375 +
       skoXC1 *
       (141736119947136/138671875 +
        skoXC1 *
        (11642977821146157/17333984375 +
         skoXC1 *
         (2250449423926978023/8666992187500 +
          skoXC1 *
          (1172430182051498652129/17333984375000000 +
           skoXC1 *
           (53052370456426963672869/4333496093750000000 +
            skoXC1 *
            (1160911898185589262801/866699218750000000 +
             skoXC1 *
             (23354642279990984757993/216674804687500000000 +
              skoXC1 *
              (532922371819915718143683/86669921875000000000000 +
               skoXC1 *
               (1023662219745311879046273/4333496093750000000000000 +
                skoXC1 *
                4384830569741378145347787/
                866699218750000000000000000))))))))))) |
     skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) <=
     37440/71),
    inference(decision, [],
              [refute_0_120, refute_0_119, refute_0_117, refute_0_110,
               refute_0_26])).

cnf(refute_0_132, plain,
    (skoXC1 *
     (-1229328/8875 +
      skoXC1 *
      (374300982/1109375 +
       skoXC1 *
       (30915505737/277343750 +
        skoXC1 *
        (79423803545067/4437500000000 +
         skoXC1 *
         (4223242486545633/2218750000000000 +
          skoXC1 *
          (588100469506937559/4437500000000000000 +
           skoXC1 *
           (6680693351322275859/1109375000000000000000 +
            skoXC1 * 176653438002827253/1109375000000000000000))))))) <=
     37440/71), inference(resolve, [], [refute_0_130, refute_0_131])).

cnf(refute_0_133, plain, ($false),
    inference(resolve, [], [refute_0_132, refute_0_120])).
% SZS output end CNFRefutation for CONVOI.tptp
