% SZS output start CNFRefutation for cos-error-analysis-7.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

cnf(cos_upper_bound_1, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320,
            Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_3, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200 +
            X ^ 16 / 20922789888000, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_1, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_3, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200) |
     lgen(R, Y, cos(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(cos_error_analysis_7, conjecture,
    (! [X] :
       ((0 <= X & X <= pi / 2) =>
        abs(cos(X) -
            (1 - 4999999963/10000000000 * X ^ 2 +
             208333209/5000000000 * X ^ 4 - 13888397/10000000000 * X ^ 6 +
             247609/10000000000 * X ^ 8 - 521/2000000000 * X ^ 10)) <=
        3/1250000000))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= pi / 2) =>
        abs(cos(X) -
            (1 - 4999999963/10000000000 * X ^ 2 +
             208333209/5000000000 * X ^ 4 - 13888397/10000000000 * X ^ 6 +
             247609/10000000000 * X ^ 8 - 521/2000000000 * X ^ 10)) <=
        3/1250000000)), inference(strip, [], [cos_error_analysis_7])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= pi / 2) =>
          abs(cos(X) -
              (1 - 4999999963/10000000000 * X ^ 2 +
               208333209/5000000000 * X ^ 4 -
               13888397/10000000000 * X ^ 6 + 247609/10000000000 * X ^ 8 -
               521/2000000000 * X ^ 10)) <= 3/1250000000)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (3/1250000000 <
        abs(cos(X) -
            (1 - 4999999963/10000000000 * X ^ 2 +
             208333209/5000000000 * X ^ 4 - 13888397/10000000000 * X ^ 6 +
             247609/10000000000 * X ^ 8 - 521/2000000000 * X ^ 10)) &
        X <= pi / 2 & 0 <= X)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (3/1250000000 <
     abs(cos(skoXC1) -
         (1 - 4999999963/10000000000 * skoXC1 ^ 2 +
          208333209/5000000000 * skoXC1 ^ 4 -
          13888397/10000000000 * skoXC1 ^ 6 +
          247609/10000000000 * skoXC1 ^ 8 -
          521/2000000000 * skoXC1 ^ 10)) & 0 <= skoXC1 & skoXC1 <= pi / 2),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (3/1250000000 <
     abs(cos(skoXC1) -
         (1 - 4999999963/10000000000 * skoXC1 ^ 2 +
          208333209/5000000000 * skoXC1 ^ 4 -
          13888397/10000000000 * skoXC1 ^ 6 +
          247609/10000000000 * skoXC1 ^ 8 -
          521/2000000000 * skoXC1 ^ 10))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_4, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_5, plain, (skoXC1 <= pi / 2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (3/1250000000 <
     abs(cos(skoXC1) -
         (1 - 4999999963/10000000000 * skoXC1 ^ 2 +
          208333209/5000000000 * skoXC1 ^ 4 -
          13888397/10000000000 * skoXC1 ^ 6 +
          247609/10000000000 * skoXC1 ^ 8 -
          521/2000000000 * skoXC1 ^ 10))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain,
    (3/1250000000 <
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1))), inference(arithmetic, [], [refute_0_0])).

cnf(refute_0_2, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) < 0 |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) =
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_3, plain,
    (3/1250000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) !=
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) <= 3/1250000000), introduced(tautology, [equality])).

cnf(refute_0_4, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) < 0 |
     3/1250000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) <= 3/1250000000),
    inference(resolve, [], [refute_0_2, refute_0_3])).

cnf(refute_0_5, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) < 0 |
     3/1250000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)),
    inference(resolve, [], [refute_0_4, refute_0_1])).

cnf(refute_0_6, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) < cos(skoXC1)),
    inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) =
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (4999999963/10000000000 +
        skoXC1 *
        (skoXC1 *
         (-208333209/5000000000 +
          skoXC1 *
          (skoXC1 *
           (13888397/10000000000 +
            skoXC1 *
            (skoXC1 *
             (-247609/10000000000 +
              skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_8, plain,
    (3/1250000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (4999999963/10000000000 +
        skoXC1 *
        (skoXC1 *
         (-208333209/5000000000 +
          skoXC1 *
          (skoXC1 *
           (13888397/10000000000 +
            skoXC1 *
            (skoXC1 *
             (-247609/10000000000 +
              skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) !=
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (4999999963/10000000000 +
        skoXC1 *
        (skoXC1 *
         (-208333209/5000000000 +
          skoXC1 *
          (skoXC1 *
           (13888397/10000000000 +
            skoXC1 *
            (skoXC1 *
             (-247609/10000000000 +
              skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) <= 3/1250000000), introduced(tautology, [equality])).

cnf(refute_0_9, plain,
    (3/1250000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (4999999963/10000000000 +
        skoXC1 *
        (skoXC1 *
         (-208333209/5000000000 +
          skoXC1 *
          (skoXC1 *
           (13888397/10000000000 +
            skoXC1 *
            (skoXC1 *
             (-247609/10000000000 +
              skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)) |
     0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (4999999963/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-208333209/5000000000 +
             skoXC1 *
             (skoXC1 *
              (13888397/10000000000 +
               skoXC1 *
               (skoXC1 *
                (-247609/10000000000 +
                 skoXC1 * (skoXC1 * 521/2000000000))))))))) +
         cos(skoXC1)) <= 3/1250000000),
    inference(resolve, [], [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain,
    (3/1250000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (4999999963/10000000000 +
        skoXC1 *
        (skoXC1 *
         (-208333209/5000000000 +
          skoXC1 *
          (skoXC1 *
           (13888397/10000000000 +
            skoXC1 *
            (skoXC1 *
             (-247609/10000000000 +
              skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)) |
     0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (-247609/10000000000 +
             skoXC1 * (skoXC1 * 521/2000000000))))))))) + cos(skoXC1)),
    inference(resolve, [], [refute_0_9, refute_0_1])).

cnf(refute_0_11, plain,
    (cos(skoXC1) <
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (~ lgen(0, X_000058, cos(X_000057)) | X_000058 <= cos(X_000057)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (~ lgen(0, X_000058,
            1 - X_000057 ^ 2 / 2 + X_000057 ^ 4 / 24 - X_000057 ^ 6 / 720 +
            X_000057 ^ 8 / 40320 - X_000057 ^ 10 / 3628800 +
            X_000057 ^ 12 / 479001600 - X_000057 ^ 14 / 87178291200) |
     lgen(0, X_000058, cos(X_000057))),
    inference(subst, [], [cos_lower_bound_3])).

cnf(refute_0_14, plain,
    (~ lgen(0, X_000058,
            1 - X_000057 ^ 2 / 2 + X_000057 ^ 4 / 24 - X_000057 ^ 6 / 720 +
            X_000057 ^ 8 / 40320 - X_000057 ^ 10 / 3628800 +
            X_000057 ^ 12 / 479001600 - X_000057 ^ 14 / 87178291200) |
     X_000058 <= cos(X_000057)),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (1 +
     X_000057 *
     (X_000057 *
      (-1/2 +
       X_000057 *
       (X_000057 *
        (1/24 +
         X_000057 *
         (X_000057 *
          (-1/720 +
           X_000057 *
           (X_000057 *
            (1/40320 +
             X_000057 *
             (X_000057 *
              (-1/3628800 +
               X_000057 *
               (X_000057 *
                (1/479001600 +
                 X_000057 * (X_000057 * -1/87178291200))))))))))))) <
     X_000058 | X_000058 <= cos(X_000057)),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (-25633/630000000000 +
             skoXC1 *
             (skoXC1 *
              (17093/1134000000000 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000040, cos(X_000039)) | X_000040 <= cos(X_000039)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000040,
            1 - X_000039 ^ 2 / 2 + X_000039 ^ 4 / 24 - X_000039 ^ 6 / 720 +
            X_000039 ^ 8 / 40320 - X_000039 ^ 10 / 3628800) |
     lgen(0, X_000040, cos(X_000039))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000040,
            1 - X_000039 ^ 2 / 2 + X_000039 ^ 4 / 24 - X_000039 ^ 6 / 720 +
            X_000039 ^ 8 / 40320 - X_000039 ^ 10 / 3628800) |
     X_000040 <= cos(X_000039)),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (1 +
     X_000039 *
     (X_000039 *
      (-1/2 +
       X_000039 *
       (X_000039 *
        (1/24 +
         X_000039 *
         (X_000039 *
          (-1/720 +
           X_000039 *
           (X_000039 *
            (1/40320 + X_000039 * (X_000039 * -1/3628800))))))))) <
     X_000040 | X_000040 <= cos(X_000039)),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (-25633/630000000000 +
             skoXC1 * (skoXC1 * 17093/1134000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_25, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (-25633/630000000000 +
             skoXC1 * (skoXC1 * 17093/1134000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (~ lgen(0, X_000026, cos(X_000025)) | X_000026 <= cos(X_000025)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 -
            X_000025 ^ 6 / 720) | lgen(0, X_000026, cos(X_000025))),
    inference(subst, [], [cos_lower_bound_1])).

cnf(refute_0_28, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 -
            X_000025 ^ 6 / 720) | X_000026 <= cos(X_000025)),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (1 +
     X_000025 *
     (X_000025 *
      (-1/2 +
       X_000025 * (X_000025 * (1/24 + X_000025 * (X_000025 * -1/720))))) <
     X_000026 | X_000026 <= cos(X_000025)),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (~ lgen(0, X_000012, cos(X_000011)) | X_000012 <= cos(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_33, plain,
    (~ lgen(0, X_000012, 1 - X_000011 ^ 2 / 2) |
     lgen(0, X_000012, cos(X_000011))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_34, plain,
    (~ lgen(0, X_000012, 1 - X_000011 ^ 2 / 2) |
     X_000012 <= cos(X_000011)),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (1 + X_000011 * (X_000011 * -1/2) < X_000012 |
     X_000012 <= cos(X_000011)), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_39, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_41, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_42, plain, (skoXC1 <= pi / 2),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_43, plain, (skoXC1 <= pi * 1/2),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_45, plain,
    (skoXC1 *
     (skoXC1 *
      (37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (-373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (-25633/630000000000 +
             skoXC1 *
             (skoXC1 *
              (17093/1134000000000 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) <=
     3/1250000000 |
     1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_23, refute_0_25, refute_0_31, refute_0_37,
               refute_0_39, refute_0_40, refute_0_41, refute_0_43,
               refute_0_44])).

cnf(refute_0_46, plain,
    (1249999997/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(resolve, [], [refute_0_45, refute_0_17])).

cnf(refute_0_47, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <= cos(skoXC1)),
    inference(resolve, [], [refute_0_46, refute_0_11])).

cnf(refute_0_48, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) < cos(skoXC1)),
    inference(resolve, [], [refute_0_47, refute_0_6])).

cnf(refute_0_49, plain,
    (~ lgen(0, cos(X_000064), X_000065) | cos(X_000064) <= X_000065),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_50, plain,
    (~ lgen(0,
            1 - X_000064 ^ 2 / 2 + X_000064 ^ 4 / 24 - X_000064 ^ 6 / 720 +
            X_000064 ^ 8 / 40320 - X_000064 ^ 10 / 3628800 +
            X_000064 ^ 12 / 479001600 - X_000064 ^ 14 / 87178291200 +
            X_000064 ^ 16 / 20922789888000, X_000065) |
     lgen(0, cos(X_000064), X_000065)),
    inference(subst, [], [cos_upper_bound_3])).

cnf(refute_0_51, plain,
    (~ lgen(0,
            1 - X_000064 ^ 2 / 2 + X_000064 ^ 4 / 24 - X_000064 ^ 6 / 720 +
            X_000064 ^ 8 / 40320 - X_000064 ^ 10 / 3628800 +
            X_000064 ^ 12 / 479001600 - X_000064 ^ 14 / 87178291200 +
            X_000064 ^ 16 / 20922789888000, X_000065) |
     cos(X_000064) <= X_000065),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (X_000065 <
     1 +
     X_000064 *
     (X_000064 *
      (-1/2 +
       X_000064 *
       (X_000064 *
        (1/24 +
         X_000064 *
         (X_000064 *
          (-1/720 +
           X_000064 *
           (X_000064 *
            (1/40320 +
             X_000064 *
             (X_000064 *
              (-1/3628800 +
               X_000064 *
               (X_000064 *
                (1/479001600 +
                 X_000064 *
                 (X_000064 *
                  (-1/87178291200 +
                   X_000064 * (X_000064 * 1/20922789888000))))))))))))))) |
     cos(X_000064) <= X_000065), inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     cos(skoXC1) <=
     1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000)))))))))),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(resolve, [], [refute_0_53, refute_0_48])).

cnf(refute_0_55, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (-37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (-4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (25633/630000000000 +
             skoXC1 *
             (skoXC1 *
              (-17093/1134000000000 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (~ lgen(0, cos(X_000032), X_000033) | cos(X_000032) <= X_000033),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_57, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320, X_000033) |
     lgen(0, cos(X_000032), X_000033)),
    inference(subst, [], [cos_upper_bound_1])).

cnf(refute_0_58, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320, X_000033) | cos(X_000032) <= X_000033),
    inference(resolve, [], [refute_0_57, refute_0_56])).

cnf(refute_0_59, plain,
    (X_000033 <
     1 +
     X_000032 *
     (X_000032 *
      (-1/2 +
       X_000032 *
       (X_000032 *
        (1/24 +
         X_000032 *
         (X_000032 * (-1/720 + X_000032 * (X_000032 * 1/40320))))))) |
     cos(X_000032) <= X_000033), inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <=
     1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000)))))))))),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_60, refute_0_48])).

cnf(refute_0_62, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (-37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (-4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (25633/630000000000 +
             skoXC1 * (skoXC1 * 521/2000000000)))))))))),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (~ lgen(0, cos(X_000046), X_000047) | cos(X_000046) <= X_000047),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_64, plain,
    (~ lgen(0,
            1 - X_000046 ^ 2 / 2 + X_000046 ^ 4 / 24 - X_000046 ^ 6 / 720 +
            X_000046 ^ 8 / 40320 - X_000046 ^ 10 / 3628800 +
            X_000046 ^ 12 / 479001600, X_000047) |
     lgen(0, cos(X_000046), X_000047)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_65, plain,
    (~ lgen(0,
            1 - X_000046 ^ 2 / 2 + X_000046 ^ 4 / 24 - X_000046 ^ 6 / 720 +
            X_000046 ^ 8 / 40320 - X_000046 ^ 10 / 3628800 +
            X_000046 ^ 12 / 479001600, X_000047) |
     cos(X_000046) <= X_000047),
    inference(resolve, [], [refute_0_64, refute_0_63])).

cnf(refute_0_66, plain,
    (X_000047 <
     1 +
     X_000046 *
     (X_000046 *
      (-1/2 +
       X_000046 *
       (X_000046 *
        (1/24 +
         X_000046 *
         (X_000046 *
          (-1/720 +
           X_000046 *
           (X_000046 *
            (1/40320 +
             X_000046 *
             (X_000046 *
              (-1/3628800 + X_000046 * (X_000046 * 1/479001600))))))))))) |
     cos(X_000046) <= X_000047), inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000)))))))))),
    inference(subst, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (1250000003/1250000000 +
     skoXC1 *
     (skoXC1 *
      (-4999999963/10000000000 +
       skoXC1 *
       (skoXC1 *
        (208333209/5000000000 +
         skoXC1 *
         (skoXC1 *
          (-13888397/10000000000 +
           skoXC1 *
           (skoXC1 *
            (247609/10000000000 +
             skoXC1 * (skoXC1 * -521/2000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))),
    inference(resolve, [], [refute_0_67, refute_0_48])).

cnf(refute_0_69, plain,
    (3/1250000000 <
     skoXC1 *
     (skoXC1 *
      (-37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (-4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (25633/630000000000 +
             skoXC1 *
             (skoXC1 *
              (-17093/1134000000000 +
               skoXC1 * (skoXC1 * 1/479001600)))))))))))),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (skoXC1 *
     (skoXC1 *
      (-37/10000000000 +
       skoXC1 *
       (skoXC1 *
        (373/15000000000 +
         skoXC1 *
         (skoXC1 *
          (-4427/90000000000 +
           skoXC1 *
           (skoXC1 *
            (25633/630000000000 +
             skoXC1 *
             (skoXC1 *
              (-17093/1134000000000 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) <=
     3/1250000000),
    inference(decision, [],
              [refute_0_44, refute_0_43, refute_0_41, refute_0_40,
               refute_0_62, refute_0_69])).

cnf(refute_0_71, plain, ($false),
    inference(resolve, [], [refute_0_70, refute_0_55])).
% SZS output end CNFRefutation for cos-error-analysis-7.tptp
