% SZS output start CNFRefutation for cos-error-analysis-5.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_1, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_3, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200) |
     lgen(R, Y, cos(X)))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(cos_error_analysis_5, conjecture,
    (! [X] :
       (abs(X) <= 1/32 =>
        abs(cos(X) - 1 +
            1562499999999999999999999857053/
            3125000000000000000000000000000 * X ^ 2 -
            52083333333333333332503437/1250000000000000000000000000 *
            X ^ 4 + 86805555555555359977/62500000000000000000000 * X ^ 6 -
            62003968237876409/2500000000000000000000 * X ^ 8 +
            34445897759/125000000000000000 * X ^ 10) <=
        99/100000000000000000000000000000000))).

fof(subgoal_0, plain,
    (! [X] :
       (abs(X) <= 1/32 =>
        abs(cos(X) - 1 +
            1562499999999999999999999857053/
            3125000000000000000000000000000 * X ^ 2 -
            52083333333333333332503437/1250000000000000000000000000 *
            X ^ 4 + 86805555555555359977/62500000000000000000000 * X ^ 6 -
            62003968237876409/2500000000000000000000 * X ^ 8 +
            34445897759/125000000000000000 * X ^ 10) <=
        99/100000000000000000000000000000000)),
    inference(strip, [], [cos_error_analysis_5])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (abs(X) <= 1/32 =>
          abs(cos(X) - 1 +
              1562499999999999999999999857053/
              3125000000000000000000000000000 * X ^ 2 -
              52083333333333333332503437/1250000000000000000000000000 *
              X ^ 4 +
              86805555555555359977/62500000000000000000000 * X ^ 6 -
              62003968237876409/2500000000000000000000 * X ^ 8 +
              34445897759/125000000000000000 * X ^ 10) <=
          99/100000000000000000000000000000000)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (99/100000000000000000000000000000000 <
        abs(cos(X) - 1 +
            1562499999999999999999999857053/
            3125000000000000000000000000000 * X ^ 2 -
            52083333333333333332503437/1250000000000000000000000000 *
            X ^ 4 + 86805555555555359977/62500000000000000000000 * X ^ 6 -
            62003968237876409/2500000000000000000000 * X ^ 8 +
            34445897759/125000000000000000 * X ^ 10) & abs(X) <= 1/32)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (99/100000000000000000000000000000000 <
     abs(cos(skoXC1) - 1 +
         1562499999999999999999999857053/3125000000000000000000000000000 *
         skoXC1 ^ 2 -
         52083333333333333332503437/1250000000000000000000000000 *
         skoXC1 ^ 4 +
         86805555555555359977/62500000000000000000000 * skoXC1 ^ 6 -
         62003968237876409/2500000000000000000000 * skoXC1 ^ 8 +
         34445897759/125000000000000000 * skoXC1 ^ 10) &
     abs(skoXC1) <= 1/32), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (99/100000000000000000000000000000000 <
     abs(cos(skoXC1) - 1 +
         1562499999999999999999999857053/3125000000000000000000000000000 *
         skoXC1 ^ 2 -
         52083333333333333332503437/1250000000000000000000000000 *
         skoXC1 ^ 4 +
         86805555555555359977/62500000000000000000000 * skoXC1 ^ 6 -
         62003968237876409/2500000000000000000000 * skoXC1 ^ 8 +
         34445897759/125000000000000000 * skoXC1 ^ 10)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (abs(skoXC1) <= 1/32),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (~ lgen(0, cos(X_000046), X_000047) | cos(X_000046) <= X_000047),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0,
            1 - X_000046 ^ 2 / 2 + X_000046 ^ 4 / 24 - X_000046 ^ 6 / 720 +
            X_000046 ^ 8 / 40320 - X_000046 ^ 10 / 3628800 +
            X_000046 ^ 12 / 479001600, X_000047) |
     lgen(0, cos(X_000046), X_000047)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_2, plain,
    (~ lgen(0,
            1 - X_000046 ^ 2 / 2 + X_000046 ^ 4 / 24 - X_000046 ^ 6 / 720 +
            X_000046 ^ 8 / 40320 - X_000046 ^ 10 / 3628800 +
            X_000046 ^ 12 / 479001600, X_000047) |
     cos(X_000046) <= X_000047),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000047 <
     1 +
     X_000046 *
     (X_000046 *
      (-1/2 +
       X_000046 *
       (X_000046 *
        (1/24 +
         X_000046 *
         (X_000046 *
          (-1/720 +
           X_000046 *
           (X_000046 *
            (1/40320 +
             X_000046 *
             (X_000046 *
              (-1/3628800 + X_000046 * (X_000046 * 1/479001600))))))))))) |
     cos(X_000046) <= X_000047), inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (100000000000000000000000000000099/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     100000000000000000000000000000099/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000)))))))))),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (99/100000000000000000000000000000000 <
     abs(cos(skoXC1) - 1 +
         1562499999999999999999999857053/3125000000000000000000000000000 *
         skoXC1 ^ 2 -
         52083333333333333332503437/1250000000000000000000000000 *
         skoXC1 ^ 4 +
         86805555555555359977/62500000000000000000000 * skoXC1 ^ 6 -
         62003968237876409/2500000000000000000000 * skoXC1 ^ 8 +
         34445897759/125000000000000000 * skoXC1 ^ 10)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_6, plain,
    (99/100000000000000000000000000000000 <
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1))), inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) < 0 |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) =
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1)), inference(subst, [], [abs_nonnegative])).

cnf(refute_0_8, plain,
    (99/100000000000000000000000000000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) !=
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) <= 99/100000000000000000000000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_9, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) < 0 |
     99/100000000000000000000000000000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) <= 99/100000000000000000000000000000000),
    inference(resolve, [], [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain,
    (-1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) < 0 |
     99/100000000000000000000000000000000 <
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1)), inference(resolve, [], [refute_0_9, refute_0_6])).

cnf(refute_0_11, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     100000000000000000000000000000099/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <
     cos(skoXC1)), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) =
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (1562499999999999999999999857053/3125000000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-52083333333333333332503437/1250000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (86805555555555359977/62500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-62003968237876409/2500000000000000000000 +
              skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
      cos(skoXC1))), inference(subst, [], [abs_negative])).

cnf(refute_0_13, plain,
    (99/100000000000000000000000000000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (1562499999999999999999999857053/3125000000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-52083333333333333332503437/1250000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (86805555555555359977/62500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-62003968237876409/2500000000000000000000 +
              skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
      cos(skoXC1)) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) !=
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (1562499999999999999999999857053/3125000000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-52083333333333333332503437/1250000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (86805555555555359977/62500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-62003968237876409/2500000000000000000000 +
              skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
      cos(skoXC1)) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) <= 99/100000000000000000000000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_14, plain,
    (99/100000000000000000000000000000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (1562499999999999999999999857053/3125000000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-52083333333333333332503437/1250000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (86805555555555359977/62500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-62003968237876409/2500000000000000000000 +
              skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
      cos(skoXC1)) |
     0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1) |
     abs(-1 +
         skoXC1 *
         (skoXC1 *
          (1562499999999999999999999857053/
           3125000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-52083333333333333332503437/1250000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (86805555555555359977/62500000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-62003968237876409/2500000000000000000000 +
                 skoXC1 *
                 (skoXC1 * 34445897759/125000000000000000))))))))) +
         cos(skoXC1)) <= 99/100000000000000000000000000000000),
    inference(resolve, [], [refute_0_12, refute_0_13])).

cnf(refute_0_15, plain,
    (99/100000000000000000000000000000000 <
     -(-1 +
      skoXC1 *
      (skoXC1 *
       (1562499999999999999999999857053/3125000000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-52083333333333333332503437/1250000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (86805555555555359977/62500000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (-62003968237876409/2500000000000000000000 +
              skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
      cos(skoXC1)) |
     0 <=
     -1 +
     skoXC1 *
     (skoXC1 *
      (1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * 34445897759/125000000000000000))))))))) +
     cos(skoXC1)), inference(resolve, [], [refute_0_14, refute_0_6])).

cnf(refute_0_16, plain,
    (cos(skoXC1) <
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (~ lgen(0, X_000058, cos(X_000057)) | X_000058 <= cos(X_000057)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000058,
            1 - X_000057 ^ 2 / 2 + X_000057 ^ 4 / 24 - X_000057 ^ 6 / 720 +
            X_000057 ^ 8 / 40320 - X_000057 ^ 10 / 3628800 +
            X_000057 ^ 12 / 479001600 - X_000057 ^ 14 / 87178291200) |
     lgen(0, X_000058, cos(X_000057))),
    inference(subst, [], [cos_lower_bound_3])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000058,
            1 - X_000057 ^ 2 / 2 + X_000057 ^ 4 / 24 - X_000057 ^ 6 / 720 +
            X_000057 ^ 8 / 40320 - X_000057 ^ 10 / 3628800 +
            X_000057 ^ 12 / 479001600 - X_000057 ^ 14 / 87178291200) |
     X_000058 <= cos(X_000057)),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (1 +
     X_000057 *
     (X_000057 *
      (-1/2 +
       X_000057 *
       (X_000057 *
        (1/24 +
         X_000057 *
         (X_000057 *
          (-1/720 +
           X_000057 *
           (X_000057 *
            (1/40320 +
             X_000057 *
             (X_000057 *
              (-1/3628800 +
               X_000057 *
               (X_000057 *
                (1/479001600 +
                 X_000057 * (X_000057 * -1/87178291200))))))))))))) <
     X_000058 | X_000058 <= cos(X_000057)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (99/100000000000000000000000000000000 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-1013786233/157500000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (425970647/70875000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000040, cos(X_000039)) | X_000040 <= cos(X_000039)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_24, plain,
    (~ lgen(0, X_000040,
            1 - X_000039 ^ 2 / 2 + X_000039 ^ 4 / 24 - X_000039 ^ 6 / 720 +
            X_000039 ^ 8 / 40320 - X_000039 ^ 10 / 3628800) |
     lgen(0, X_000040, cos(X_000039))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_25, plain,
    (~ lgen(0, X_000040,
            1 - X_000039 ^ 2 / 2 + X_000039 ^ 4 / 24 - X_000039 ^ 6 / 720 +
            X_000039 ^ 8 / 40320 - X_000039 ^ 10 / 3628800) |
     X_000040 <= cos(X_000039)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (1 +
     X_000039 *
     (X_000039 *
      (-1/2 +
       X_000039 *
       (X_000039 *
        (1/24 +
         X_000039 *
         (X_000039 *
          (-1/720 +
           X_000039 *
           (X_000039 *
            (1/40320 + X_000039 * (X_000039 * -1/3628800))))))))) <
     X_000040 | X_000040 <= cos(X_000039)),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (99/100000000000000000000000000000000 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-1013786233/157500000000000000000000 +
             skoXC1 * (skoXC1 * 425970647/70875000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_26])).

cnf(refute_0_30, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-1013786233/157500000000000000000000 +
             skoXC1 * (skoXC1 * 425970647/70875000000000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (~ lgen(0, X_000026, cos(X_000025)) | X_000026 <= cos(X_000025)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_32, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 -
            X_000025 ^ 6 / 720) | lgen(0, X_000026, cos(X_000025))),
    inference(subst, [], [cos_lower_bound_1])).

cnf(refute_0_33, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 -
            X_000025 ^ 6 / 720) | X_000026 <= cos(X_000025)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (1 +
     X_000025 *
     (X_000025 *
      (-1/2 +
       X_000025 * (X_000025 * (1/24 + X_000025 * (X_000025 * -1/720))))) <
     X_000026 | X_000026 <= cos(X_000025)),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (99/100000000000000000000000000000000 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000012, cos(X_000011)) | X_000012 <= cos(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_38, plain,
    (~ lgen(0, X_000012, 1 - X_000011 ^ 2 / 2) |
     lgen(0, X_000012, cos(X_000011))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_39, plain,
    (~ lgen(0, X_000012, 1 - X_000011 ^ 2 / 2) |
     X_000012 <= cos(X_000011)),
    inference(resolve, [], [refute_0_38, refute_0_37])).

cnf(refute_0_40, plain,
    (1 + X_000011 * (X_000011 * -1/2) < X_000012 |
     X_000012 <= cos(X_000011)), inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (99/100000000000000000000000000000000 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_40])).

cnf(refute_0_44, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain, (abs(skoXC1) <= 1/32),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_46, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_47, plain,
    (1/32 < abs(skoXC1) | abs(skoXC1) != -skoXC1 | -skoXC1 <= 1/32),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain,
    (1/32 < abs(skoXC1) | 0 <= skoXC1 | -skoXC1 <= 1/32),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain, (0 <= skoXC1 | -skoXC1 <= 1/32),
    inference(resolve, [], [refute_0_45, refute_0_48])).

cnf(refute_0_50, plain, (-1/32 <= skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_52, plain,
    (1/32 < abs(skoXC1) | abs(skoXC1) != skoXC1 | skoXC1 <= 1/32),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain, (skoXC1 < 0 | 1/32 < abs(skoXC1) | skoXC1 <= 1/32),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (skoXC1 < 0 | skoXC1 <= 1/32),
    inference(resolve, [], [refute_0_45, refute_0_53])).

cnf(refute_0_55, plain, (0 <= skoXC1 | skoXC1 <= 1/32),
    introduced(tautology, [decision])).

cnf(refute_0_56, plain, (skoXC1 <= 1/32),
    inference(resolve, [], [refute_0_55, refute_0_54])).

cnf(refute_0_57, plain, (skoXC1 < 0 | -1/32 <= skoXC1),
    inference(decision, [], [refute_0_56])).

cnf(refute_0_58, plain, (-1/32 <= skoXC1),
    inference(resolve, [], [refute_0_50, refute_0_57])).

cnf(refute_0_59, plain,
    (skoXC1 *
     (skoXC1 *
      (142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-1013786233/157500000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (425970647/70875000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) <=
     99/100000000000000000000000000000000 |
     99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)),
    inference(decision, [],
              [refute_0_28, refute_0_30, refute_0_36, refute_0_42,
               refute_0_44, refute_0_58, refute_0_56])).

cnf(refute_0_60, plain,
    (99999999999999999999999999999901/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(resolve, [], [refute_0_59, refute_0_22])).

cnf(refute_0_61, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <=
     cos(skoXC1)), inference(resolve, [], [refute_0_60, refute_0_16])).

cnf(refute_0_62, plain,
    (100000000000000000000000000000099/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <
     cos(skoXC1)), inference(resolve, [], [refute_0_61, refute_0_11])).

cnf(refute_0_63, plain,
    (100000000000000000000000000000099/100000000000000000000000000000000 +
     skoXC1 *
     (skoXC1 *
      (-1562499999999999999999999857053/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (52083333333333333332503437/1250000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-86805555555555359977/62500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (62003968237876409/2500000000000000000000 +
             skoXC1 * (skoXC1 * -34445897759/125000000000000000))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))),
    inference(resolve, [], [refute_0_4, refute_0_62])).

cnf(refute_0_64, plain,
    (99/100000000000000000000000000000000 <
     skoXC1 *
     (skoXC1 *
      (-142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (1013786233/157500000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-425970647/70875000000000000000 +
               skoXC1 * (skoXC1 * 1/479001600)))))))))))),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 *
     (skoXC1 *
      (-142947/3125000000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (2489689/3750000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-1760207/562500000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (1013786233/157500000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-425970647/70875000000000000000 +
               skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     99/100000000000000000000000000000000),
    inference(decision, [], [refute_0_58, refute_0_56])).

cnf(refute_0_66, plain, ($false),
    inference(resolve, [], [refute_0_65, refute_0_64])).
% SZS output end CNFRefutation for cos-error-analysis-5.tptp
