% SZS output start CNFRefutation for cosh-3612a-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(cosh_taylor_lower_bound_pos1, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            1/2 *
            (1 + 1/3 * X + 1/18 * X ^ 2 + 1/162 * X ^ 3 + 1/1944 * X ^ 4 +
             1/29160 * X ^ 5) ^ 3 -
            1/2 * (X ^ 3 - 12 * X ^ 2 + 60 * X - 120) /
            (X ^ 3 + 12 * X ^ 2 + 60 * X + 120)) | lgen(R, Y, cosh(X)))).

cnf(cosh_lower_bound_pos1, axiom,
    (X < 0 | 5/2 < X |
     ~ lgen(R, Y,
            6 * (3 * X ^ 4 + 104 * X ^ 2 + 240) /
            ((X ^ 2 - 6 * X + 12) * (X ^ 3 + 12 * X ^ 2 + 60 * X + 120))) |
     lgen(R, Y, cosh(X)))).

cnf(sinh_upper_bound_pos2, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -24 * X * (X ^ 4 + 70 * X ^ 2 + 600) /
            ((X ^ 3 - 12 * X ^ 2 + 60 * X - 120) *
             (X ^ 3 + 12 * X ^ 2 + 60 * X + 120)), Y) |
     lgen(R, sinh(X), Y))).

fof(cosh_3612, conjecture,
    (! [X] :
       ((0 < X & X < 12/5) =>
        (3 + X ^ 2 / 11) * sinh(X) / (2 + cosh(X) + X ^ 2 / 11) < X))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 12/5) =>
        (3 + X ^ 2 / 11) * sinh(X) / (2 + cosh(X) + X ^ 2 / 11) < X)),
    inference(strip, [], [cosh_3612])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 12/5) =>
          (3 + X ^ 2 / 11) * sinh(X) / (2 + cosh(X) + X ^ 2 / 11) < X)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 12/5 &
        X <= (3 + X ^ 2 / 11) * sinh(X) / (2 + cosh(X) + X ^ 2 / 11))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 12/5 & 0 < skoXC1 &
     skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * sinh(skoXC1) /
     (2 + cosh(skoXC1) + skoXC1 ^ 2 / 11)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * sinh(skoXC1) /
     (2 + cosh(skoXC1) + skoXC1 ^ 2 / 11)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 12/5),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) <
     (1728000 +
      skoXC1 *
      (864000 +
       skoXC1 *
       (950400 +
        skoXC1 *
        (403200 +
         skoXC1 *
         (109440 +
          skoXC1 *
          (22320 +
           skoXC1 *
           (265928/81 +
            skoXC1 *
            (25064/81 +
             skoXC1 *
             (-1004/243 +
              skoXC1 *
              (-184280/19683 +
               skoXC1 *
               (-29861/6561 +
                skoXC1 *
                (-292979/196830 +
                 skoXC1 *
                 (-4061857/10628820 +
                  skoXC1 *
                  (-3394169/42515280 +
                   skoXC1 *
                   (-887213/63772920 +
                    skoXC1 *
                    (-4720681/2295825120 +
                     skoXC1 *
                     (-2989033/11479125600 +
                      skoXC1 *
                      (-260783/9183300480 +
                       skoXC1 *
                       (-5741/2152336050 +
                        skoXC1 *
                        (-5909/27549901440 +
                         skoXC1 *
                         (-4013/275499014400 +
                          skoXC1 *
                          (-13489/16529940864000 +
                           skoXC1 *
                           (-199/5509980288000 +
                            skoXC1 *
                            (-19/16529940864000 +
                             skoXC1 *
                             -1/49589822592000)))))))))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) |
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= 0 |
     1728000 +
     skoXC1 *
     (864000 +
      skoXC1 *
      (950400 +
       skoXC1 *
       (403200 +
        skoXC1 *
        (109440 +
         skoXC1 *
         (22320 +
          skoXC1 *
          (265928/81 +
           skoXC1 *
           (25064/81 +
            skoXC1 *
            (-1004/243 +
             skoXC1 *
             (-184280/19683 +
              skoXC1 *
              (-29861/6561 +
               skoXC1 *
               (-292979/196830 +
                skoXC1 *
                (-4061857/10628820 +
                 skoXC1 *
                 (-3394169/42515280 +
                  skoXC1 *
                  (-887213/63772920 +
                   skoXC1 *
                   (-4720681/2295825120 +
                    skoXC1 *
                    (-2989033/11479125600 +
                     skoXC1 *
                     (-260783/9183300480 +
                      skoXC1 *
                      (-5741/2152336050 +
                       skoXC1 *
                       (-5909/27549901440 +
                        skoXC1 *
                        (-4013/275499014400 +
                         skoXC1 *
                         (-13489/16529940864000 +
                          skoXC1 *
                          (-199/5509980288000 +
                           skoXC1 *
                           (-19/16529940864000 +
                            skoXC1 *
                            -1/49589822592000))))))))))))))))))))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) *
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <
     (120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) |
     (120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) |
     14400 +
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (X_000040 < cosh(X_000039) | ~ lgen(1, X_000040, cosh(X_000039))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (X_000039 < 0 |
     ~ lgen(1, X_000040,
            1/2 *
            (1 + 1/3 * X_000039 + 1/18 * X_000039 ^ 2 +
             1/162 * X_000039 ^ 3 + 1/1944 * X_000039 ^ 4 +
             1/29160 * X_000039 ^ 5) ^ 3 -
            1/2 *
            (X_000039 ^ 3 - 12 * X_000039 ^ 2 + 60 * X_000039 - 120) /
            (X_000039 ^ 3 + 12 * X_000039 ^ 2 + 60 * X_000039 + 120)) |
     lgen(1, X_000040, cosh(X_000039))),
    inference(subst, [], [cosh_taylor_lower_bound_pos1])).

cnf(refute_0_4, plain,
    (X_000039 < 0 | X_000040 < cosh(X_000039) |
     ~ lgen(1, X_000040,
            1/2 *
            (1 + 1/3 * X_000039 + 1/18 * X_000039 ^ 2 +
             1/162 * X_000039 ^ 3 + 1/1944 * X_000039 ^ 4 +
             1/29160 * X_000039 ^ 5) ^ 3 -
            1/2 *
            (X_000039 ^ 3 - 12 * X_000039 ^ 2 + 60 * X_000039 - 120) /
            (X_000039 ^ 3 + 12 * X_000039 ^ 2 + 60 * X_000039 + 120))),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000039 < 0 | X_000040 < cosh(X_000039) |
     1/2 +
     X_000039 *
     (1/2 +
      X_000039 *
      (1/4 +
       X_000039 *
       (1/12 +
        X_000039 *
        (1/48 +
         X_000039 *
         (1/240 +
          X_000039 *
          (121/174960 +
           X_000039 *
           (17/174960 +
            X_000039 *
            (49/4199040 +
             X_000039 *
             (409/340122240 +
              X_000039 *
              (361/3401222400 +
               X_000039 *
               (1/125971200 +
                X_000039 *
                (181/367332019200 +
                 X_000039 *
                 (1/40814668800 +
                  X_000039 *
                  (1/1101996057600 +
                   X_000039 * 1/49589822592000)))))))))))))) +
     (-60 + X_000039 * (30 + X_000039 * (-6 + X_000039 * 1/2))) /
     (120 + X_000039 * (60 + X_000039 * (12 + X_000039))) * -1 <=
     X_000040), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 < 0 |
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <
     cosh(skoXC1) |
     1/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (1/4 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) +
     (-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) * -1 <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1) <
     skoXC1 *
     (43200 +
      skoXC1 *
      (skoXC1 *
       (69840/11 +
        skoXC1 * (skoXC1 * (2472/11 + skoXC1 * (skoXC1 * 24/11)))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) |
     0 <=
     -14400 +
     skoXC1 *
     (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) *
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) <=
     skoXC1 *
     (43200 +
      skoXC1 *
      (skoXC1 *
       (69840/11 +
        skoXC1 * (skoXC1 * (2472/11 + skoXC1 * (skoXC1 * 24/11))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_8, plain,
    (skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) <
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) |
     0 <= -3 + skoXC1 * (skoXC1 * -1/11) |
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) *
     (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_9, plain,
    (skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * sinh(skoXC1) /
     (2 + cosh(skoXC1) + skoXC1 ^ 2 / 11)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    (skoXC1 <=
     sinh(skoXC1) * (3 + skoXC1 * (skoXC1 * 1/11)) /
     (2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1))),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (sinh(skoXC1) * (3 + skoXC1 * (skoXC1 * 1/11)) /
     (2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1)) < skoXC1 |
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1)) <=
     sinh(skoXC1) * (3 + skoXC1 * (skoXC1 * 1/11)) |
     2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_12, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1)) <=
     sinh(skoXC1) * (3 + skoXC1 * (skoXC1 * 1/11)) |
     2 + skoXC1 * (skoXC1 * 1/11) + cosh(skoXC1) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (sinh(skoXC1) * (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1) |
     cosh(skoXC1) <= -2 + skoXC1 * (skoXC1 * -1/11)),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <
     5/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (15/44 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) |
     (5/2 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (15/44 +
        skoXC1 *
        (1/12 +
         skoXC1 *
         (1/48 +
          skoXC1 *
          (1/240 +
           skoXC1 *
           (121/174960 +
            skoXC1 *
            (17/174960 +
             skoXC1 *
             (49/4199040 +
              skoXC1 *
              (409/340122240 +
               skoXC1 *
               (361/3401222400 +
                skoXC1 *
                (1/125971200 +
                 skoXC1 *
                 (181/367332019200 +
                  skoXC1 *
                  (1/40814668800 +
                   skoXC1 *
                   (1/1101996057600 +
                    skoXC1 * 1/49589822592000))))))))))))))) *
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     -60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2)) |
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_15, plain,
    (skoXC1 < 0 | -2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     1/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (1/4 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) +
     (-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) * -1 <=
     -2 + skoXC1 * (skoXC1 * -1/11)), inference(subst, [], [refute_0_5])).

cnf(refute_0_16, plain,
    (skoXC1 < 0 | -2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     5/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (15/44 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) <=
     (-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_18, plain, (skoXC1 < 12/5),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_19, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_17, refute_0_18])).

cnf(refute_0_20, plain,
    (-2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     5/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (15/44 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) <=
     (-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)))),
    inference(resolve, [], [refute_0_19, refute_0_16])).

cnf(refute_0_21, plain,
    (-2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     (5/2 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (15/44 +
        skoXC1 *
        (1/12 +
         skoXC1 *
         (1/48 +
          skoXC1 *
          (1/240 +
           skoXC1 *
           (121/174960 +
            skoXC1 *
            (17/174960 +
             skoXC1 *
             (49/4199040 +
              skoXC1 *
              (409/340122240 +
               skoXC1 *
               (361/3401222400 +
                skoXC1 *
                (1/125971200 +
                 skoXC1 *
                 (181/367332019200 +
                  skoXC1 *
                  (1/40814668800 +
                   skoXC1 *
                   (1/1101996057600 +
                    skoXC1 * 1/49589822592000))))))))))))))) *
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     -60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2)) |
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_20, refute_0_14])).

cnf(refute_0_22, plain,
    (-2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120 |
     skoXC1 *
     (180 +
      skoXC1 *
      (1176/11 +
       skoXC1 *
       (423/11 +
        skoXC1 *
        (133/11 +
         skoXC1 *
         (34/11 +
          skoXC1 *
          (1943/2916 +
           skoXC1 *
           (7231/58320 +
            skoXC1 *
            (1723/87480 +
             skoXC1 *
             (15317/5668704 +
              skoXC1 *
              (3043/9447840 +
               skoXC1 *
               (421/12597120 +
                skoXC1 *
                (23047/7652750400 +
                 skoXC1 *
                 (179/765275040 +
                  skoXC1 *
                  (1417/91833004800 +
                   skoXC1 *
                   (2789/3305988172800 +
                    skoXC1 *
                    (121/3305988172800 +
                     skoXC1 *
                     (19/16529940864000 +
                      skoXC1 * 1/49589822592000))))))))))))))))) <= -360),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-360 <
     skoXC1 *
     (180 +
      skoXC1 *
      (1176/11 +
       skoXC1 *
       (423/11 +
        skoXC1 *
        (133/11 +
         skoXC1 *
         (34/11 +
          skoXC1 *
          (1943/2916 +
           skoXC1 *
           (7231/58320 +
            skoXC1 *
            (1723/87480 +
             skoXC1 *
             (15317/5668704 +
              skoXC1 *
              (3043/9447840 +
               skoXC1 *
               (421/12597120 +
                skoXC1 *
                (23047/7652750400 +
                 skoXC1 *
                 (179/765275040 +
                  skoXC1 *
                  (1417/91833004800 +
                   skoXC1 *
                   (2789/3305988172800 +
                    skoXC1 *
                    (121/3305988172800 +
                     skoXC1 *
                     (19/16529940864000 +
                      skoXC1 * 1/49589822592000))))))))))))))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(decision, [], [refute_0_17, refute_0_18])).

cnf(refute_0_24, plain,
    (-2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain, (-120 < skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(decision, [], [refute_0_17, refute_0_18])).

cnf(refute_0_26, plain, (-2 + skoXC1 * (skoXC1 * -1/11) < cosh(skoXC1)),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain,
    (sinh(skoXC1) * (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_13, refute_0_26])).

cnf(refute_0_28, plain,
    (0 <= -3 + skoXC1 * (skoXC1 * -1/11) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     sinh(skoXC1)), inference(split, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 * (skoXC1 * 1/11) <= -3 |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     sinh(skoXC1)), inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain, (-3 < skoXC1 * (skoXC1 * 1/11)),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_31, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     sinh(skoXC1)), inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain,
    (sinh(X_000141) < X_000142 | ~ lgen(1, sinh(X_000141), X_000142)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_33, plain,
    (4 < X_000141 | X_000141 < 0 |
     ~ lgen(1,
            -24 * X_000141 * (X_000141 ^ 4 + 70 * X_000141 ^ 2 + 600) /
            ((X_000141 ^ 3 - 12 * X_000141 ^ 2 + 60 * X_000141 - 120) *
             (X_000141 ^ 3 + 12 * X_000141 ^ 2 + 60 * X_000141 + 120)),
            X_000142) | lgen(1, sinh(X_000141), X_000142)),
    inference(subst, [], [sinh_upper_bound_pos2])).

cnf(refute_0_34, plain,
    (4 < X_000141 | sinh(X_000141) < X_000142 | X_000141 < 0 |
     ~ lgen(1,
            -24 * X_000141 * (X_000141 ^ 4 + 70 * X_000141 ^ 2 + 600) /
            ((X_000141 ^ 3 - 12 * X_000141 ^ 2 + 60 * X_000141 - 120) *
             (X_000141 ^ 3 + 12 * X_000141 ^ 2 + 60 * X_000141 + 120)),
            X_000142)),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (4 < X_000141 | sinh(X_000141) < X_000142 | X_000141 < 0 |
     X_000142 <=
     X_000141 *
     (-14400 +
      X_000141 * (X_000141 * (-1680 + X_000141 * (X_000141 * -24)))) /
     (-14400 +
      X_000141 *
      (X_000141 *
       (720 + X_000141 * (X_000141 * (-24 + X_000141 * X_000141)))))),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 < 0 |
     sinh(skoXC1) <
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) |
     4 < skoXC1 |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))))),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (skoXC1 < 0 | 4 < skoXC1 |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_31, refute_0_36])).

cnf(refute_0_38, plain, (4 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_39, plain,
    (4 < skoXC1 |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_38, refute_0_37])).

cnf(refute_0_40, plain, (skoXC1 <= 4),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_41, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) / (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_40, refute_0_39])).

cnf(refute_0_42, plain,
    (0 <= -3 + skoXC1 * (skoXC1 * -1/11) |
     skoXC1 *
     (-14400 + skoXC1 * (skoXC1 * (-1680 + skoXC1 * (skoXC1 * -24)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) *
     (-3 + skoXC1 * (skoXC1 * -1/11)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_41, refute_0_8])).

cnf(refute_0_43, plain,
    (skoXC1 * (skoXC1 * 1/11) <= -3 |
     skoXC1 *
     (43200 +
      skoXC1 *
      (skoXC1 *
       (69840/11 +
        skoXC1 * (skoXC1 * (2472/11 + skoXC1 * (skoXC1 * 24/11)))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (skoXC1 *
     (43200 +
      skoXC1 *
      (skoXC1 *
       (69840/11 +
        skoXC1 * (skoXC1 * (2472/11 + skoXC1 * (skoXC1 * 24/11)))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     cosh(skoXC1) * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_43, refute_0_30])).

cnf(refute_0_45, plain,
    (0 <=
     -14400 +
     skoXC1 *
     (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1)))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      cosh(skoXC1) * (skoXC1 * -1)) *
     (-14400 +
      skoXC1 *
      (skoXC1 * (720 + skoXC1 * (skoXC1 * (-24 + skoXC1 * skoXC1))))) <=
     skoXC1 *
     (43200 +
      skoXC1 *
      (skoXC1 *
       (69840/11 +
        skoXC1 * (skoXC1 * (2472/11 + skoXC1 * (skoXC1 * 24/11))))))),
    inference(resolve, [], [refute_0_44, refute_0_7])).

cnf(refute_0_46, plain,
    (cosh(skoXC1) *
     (skoXC1 *
      (14400 +
       skoXC1 *
       (skoXC1 *
        (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))) <=
     skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) |
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <=
     -14400), inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (-14400 <
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_48, plain,
    (cosh(skoXC1) *
     (skoXC1 *
      (14400 +
       skoXC1 *
       (skoXC1 *
        (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))) <=
     skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))))),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <= 0 |
     cosh(skoXC1) <=
     skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (skoXC1 *
      (14400 +
       skoXC1 *
       (skoXC1 *
        (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))))),
    inference(split, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <= 0 |
     cosh(skoXC1) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (0 <
     skoXC1 *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 = 0), inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_52, plain,
    (cosh(skoXC1) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 = 0), inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain, (skoXC1 != 0),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_54, plain,
    (cosh(skoXC1) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 < 0 |
     1/2 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (1/4 +
       skoXC1 *
       (1/12 +
        skoXC1 *
        (1/48 +
         skoXC1 *
         (1/240 +
          skoXC1 *
          (121/174960 +
           skoXC1 *
           (17/174960 +
            skoXC1 *
            (49/4199040 +
             skoXC1 *
             (409/340122240 +
              skoXC1 *
              (361/3401222400 +
               skoXC1 *
               (1/125971200 +
                skoXC1 *
                (181/367332019200 +
                 skoXC1 *
                 (1/40814668800 +
                  skoXC1 *
                  (1/1101996057600 +
                   skoXC1 * 1/49589822592000)))))))))))))) +
     (-60 + skoXC1 * (30 + skoXC1 * (-6 + skoXC1 * 1/2))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) * -1 <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_54, refute_0_6])).

cnf(refute_0_56, plain,
    (skoXC1 < 0 |
     (120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) = -120),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 | skoXC1 * (60 + skoXC1 * (12 + skoXC1)) != -120),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_58, plain,
    (skoXC1 < 0 |
     (120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_56, refute_0_57])).

cnf(refute_0_59, plain,
    ((120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_19, refute_0_58])).

cnf(refute_0_60, plain,
    ((120 +
      skoXC1 *
      (60 +
       skoXC1 *
       (72 +
        skoXC1 *
        (31 +
         skoXC1 *
         (11 +
          skoXC1 *
          (3 +
           skoXC1 *
           (1943/2916 +
            skoXC1 *
            (7231/58320 +
             skoXC1 *
             (1723/87480 +
              skoXC1 *
              (15317/5668704 +
               skoXC1 *
               (3043/9447840 +
                skoXC1 *
                (421/12597120 +
                 skoXC1 *
                 (23047/7652750400 +
                  skoXC1 *
                  (179/765275040 +
                   skoXC1 *
                   (1417/91833004800 +
                    skoXC1 *
                    (2789/3305988172800 +
                     skoXC1 *
                     (121/3305988172800 +
                      skoXC1 *
                      (19/16529940864000 +
                       skoXC1 * 1/49589822592000)))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) |
     14400 +
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <= 0),
    inference(resolve, [], [refute_0_59, refute_0_1])).

cnf(refute_0_61, plain,
    (skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <=
     -14400 |
     (1728000 +
      skoXC1 *
      (864000 +
       skoXC1 *
       (950400 +
        skoXC1 *
        (403200 +
         skoXC1 *
         (109440 +
          skoXC1 *
          (22320 +
           skoXC1 *
           (265928/81 +
            skoXC1 *
            (25064/81 +
             skoXC1 *
             (-1004/243 +
              skoXC1 *
              (-184280/19683 +
               skoXC1 *
               (-29861/6561 +
                skoXC1 *
                (-292979/196830 +
                 skoXC1 *
                 (-4061857/10628820 +
                  skoXC1 *
                  (-3394169/42515280 +
                   skoXC1 *
                   (-887213/63772920 +
                    skoXC1 *
                    (-4720681/2295825120 +
                     skoXC1 *
                     (-2989033/11479125600 +
                      skoXC1 *
                      (-260783/9183300480 +
                       skoXC1 *
                       (-5741/2152336050 +
                        skoXC1 *
                        (-5909/27549901440 +
                         skoXC1 *
                         (-4013/275499014400 +
                          skoXC1 *
                          (-13489/16529940864000 +
                           skoXC1 *
                           (-199/5509980288000 +
                            skoXC1 *
                            (-19/16529940864000 +
                             skoXC1 *
                             -1/49589822592000)))))))))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    ((1728000 +
      skoXC1 *
      (864000 +
       skoXC1 *
       (950400 +
        skoXC1 *
        (403200 +
         skoXC1 *
         (109440 +
          skoXC1 *
          (22320 +
           skoXC1 *
           (265928/81 +
            skoXC1 *
            (25064/81 +
             skoXC1 *
             (-1004/243 +
              skoXC1 *
              (-184280/19683 +
               skoXC1 *
               (-29861/6561 +
                skoXC1 *
                (-292979/196830 +
                 skoXC1 *
                 (-4061857/10628820 +
                  skoXC1 *
                  (-3394169/42515280 +
                   skoXC1 *
                   (-887213/63772920 +
                    skoXC1 *
                    (-4720681/2295825120 +
                     skoXC1 *
                     (-2989033/11479125600 +
                      skoXC1 *
                      (-260783/9183300480 +
                       skoXC1 *
                       (-5741/2152336050 +
                        skoXC1 *
                        (-5909/27549901440 +
                         skoXC1 *
                         (-4013/275499014400 +
                          skoXC1 *
                          (-13489/16529940864000 +
                           skoXC1 *
                           (-199/5509980288000 +
                            skoXC1 *
                            (-19/16529940864000 +
                             skoXC1 *
                             -1/49589822592000)))))))))))))))))))))))) /
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))),
    inference(resolve, [], [refute_0_61, refute_0_47])).

cnf(refute_0_63, plain,
    (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= 0 |
     1728000 +
     skoXC1 *
     (864000 +
      skoXC1 *
      (950400 +
       skoXC1 *
       (403200 +
        skoXC1 *
        (109440 +
         skoXC1 *
         (22320 +
          skoXC1 *
          (265928/81 +
           skoXC1 *
           (25064/81 +
            skoXC1 *
            (-1004/243 +
             skoXC1 *
             (-184280/19683 +
              skoXC1 *
              (-29861/6561 +
               skoXC1 *
               (-292979/196830 +
                skoXC1 *
                (-4061857/10628820 +
                 skoXC1 *
                 (-3394169/42515280 +
                  skoXC1 *
                  (-887213/63772920 +
                   skoXC1 *
                   (-4720681/2295825120 +
                    skoXC1 *
                    (-2989033/11479125600 +
                     skoXC1 *
                     (-260783/9183300480 +
                      skoXC1 *
                      (-5741/2152336050 +
                       skoXC1 *
                       (-5909/27549901440 +
                        skoXC1 *
                        (-4013/275499014400 +
                         skoXC1 *
                         (-13489/16529940864000 +
                          skoXC1 *
                          (-199/5509980288000 +
                           skoXC1 *
                           (-19/16529940864000 +
                            skoXC1 *
                            -1/49589822592000))))))))))))))))))))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) *
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)))),
    inference(resolve, [], [refute_0_62, refute_0_0])).

cnf(refute_0_64, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (28800/11 +
         skoXC1 *
         (14400/11 +
          skoXC1 *
          (121960/891 +
           skoXC1 *
           (-47000/891 +
            skoXC1 *
            (-104356/2673 +
             skoXC1 *
             (-3641086/216513 +
              skoXC1 *
              (-407203/72171 +
               skoXC1 *
               (-3419599/2165130 +
                skoXC1 *
                (-4061857/10628820 +
                 skoXC1 *
                 (-3394169/42515280 +
                  skoXC1 *
                  (-887213/63772920 +
                   skoXC1 *
                   (-4720681/2295825120 +
                    skoXC1 *
                    (-2989033/11479125600 +
                     skoXC1 *
                     (-260783/9183300480 +
                      skoXC1 *
                      (-5741/2152336050 +
                       skoXC1 *
                       (-5909/27549901440 +
                        skoXC1 *
                        (-4013/275499014400 +
                         skoXC1 *
                         (-13489/16529940864000 +
                          skoXC1 *
                          (-199/5509980288000 +
                           skoXC1 *
                           (-19/16529940864000 +
                            skoXC1 *
                            -1/49589822592000))))))))))))))))))))))) <= 0 |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) <
     (20736000 +
      skoXC1 *
      (skoXC1 *
       (7948800 +
        skoXC1 *
        (skoXC1 *
         (-155520 +
          skoXC1 *
          (skoXC1 *
           (576 +
            skoXC1 * (skoXC1 * (-192 + skoXC1 * (skoXC1 * -18)))))))))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) |
     1440 + skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) <=
     0 |
     20736000 +
     skoXC1 *
     (skoXC1 *
      (7948800 +
       skoXC1 *
       (skoXC1 *
        (-155520 +
         skoXC1 *
         (skoXC1 *
          (576 +
           skoXC1 * (skoXC1 * (-192 + skoXC1 * (skoXC1 * -18))))))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) *
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_66, plain,
    ((14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <
     (1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) |
     (1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) |
     14400 +
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_67, plain,
    (X_000121 < cosh(X_000120) | ~ lgen(1, X_000121, cosh(X_000120))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_68, plain,
    (5/2 < X_000120 | X_000120 < 0 |
     ~ lgen(1, X_000121,
            6 * (3 * X_000120 ^ 4 + 104 * X_000120 ^ 2 + 240) /
            ((X_000120 ^ 2 - 6 * X_000120 + 12) *
             (X_000120 ^ 3 + 12 * X_000120 ^ 2 + 60 * X_000120 + 120))) |
     lgen(1, X_000121, cosh(X_000120))),
    inference(subst, [], [cosh_lower_bound_pos1])).

cnf(refute_0_69, plain,
    (5/2 < X_000120 | X_000120 < 0 | X_000121 < cosh(X_000120) |
     ~ lgen(1, X_000121,
            6 * (3 * X_000120 ^ 4 + 104 * X_000120 ^ 2 + 240) /
            ((X_000120 ^ 2 - 6 * X_000120 + 12) *
             (X_000120 ^ 3 + 12 * X_000120 ^ 2 + 60 * X_000120 + 120)))),
    inference(resolve, [], [refute_0_68, refute_0_67])).

cnf(refute_0_70, plain,
    (5/2 < X_000120 | X_000120 < 0 | X_000121 < cosh(X_000120) |
     (1440 + X_000120 * (X_000120 * (624 + X_000120 * (X_000120 * 18)))) /
     (1440 +
      X_000120 *
      (X_000120 * (-96 + X_000120 * (X_000120 * (6 + X_000120))))) <=
     X_000121), inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 |
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <
     cosh(skoXC1) | 5/2 < skoXC1 |
     (1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(subst, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 < 0 | 5/2 < skoXC1 |
     (1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_54, refute_0_71])).

cnf(refute_0_73, plain, (5/2 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_74, plain,
    (5/2 < skoXC1 |
     (1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain, (skoXC1 <= 5/2),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_76, plain,
    ((1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) /
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_75, refute_0_74])).

cnf(refute_0_77, plain,
    ((1440 + skoXC1 * (skoXC1 * (624 + skoXC1 * (skoXC1 * 18)))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) *
     (14400 +
      skoXC1 *
      (skoXC1 *
       (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1)))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))))))) |
     14400 +
     skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <= 0),
    inference(resolve, [], [refute_0_76, refute_0_66])).

cnf(refute_0_78, plain,
    (skoXC1 *
     (skoXC1 *
      (-720 + skoXC1 * (skoXC1 * (24 + skoXC1 * (skoXC1 * -1))))) <=
     -14400 |
     (20736000 +
      skoXC1 *
      (skoXC1 *
       (7948800 +
        skoXC1 *
        (skoXC1 *
         (-155520 +
          skoXC1 *
          (skoXC1 *
           (576 +
            skoXC1 * (skoXC1 * (-192 + skoXC1 * (skoXC1 * -18)))))))))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    ((20736000 +
      skoXC1 *
      (skoXC1 *
       (7948800 +
        skoXC1 *
        (skoXC1 *
         (-155520 +
          skoXC1 *
          (skoXC1 *
           (576 +
            skoXC1 * (skoXC1 * (-192 + skoXC1 * (skoXC1 * -18)))))))))) /
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1))))) <=
     14400 +
     skoXC1 *
     (skoXC1 *
      (6480 +
       skoXC1 *
       (skoXC1 *
        (2664/11 + skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))),
    inference(resolve, [], [refute_0_78, refute_0_47])).

cnf(refute_0_80, plain,
    (1440 + skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) <=
     0 |
     20736000 +
     skoXC1 *
     (skoXC1 *
      (7948800 +
       skoXC1 *
       (skoXC1 *
        (-155520 +
         skoXC1 *
         (skoXC1 *
          (576 +
           skoXC1 * (skoXC1 * (-192 + skoXC1 * (skoXC1 * -18))))))))) <=
     (14400 +
      skoXC1 *
      (skoXC1 *
       (6480 +
        skoXC1 *
        (skoXC1 *
         (2664/11 +
          skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)))))))) *
     (1440 +
      skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))))),
    inference(resolve, [], [refute_0_79, refute_0_65])).

cnf(refute_0_81, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (345600/11 +
         skoXC1 *
         (-14400 +
          skoXC1 *
          (-197280/11 +
           skoXC1 *
           (-6480 +
            skoXC1 *
            (-1584 +
             skoXC1 *
             (-2664/11 +
              skoXC1 *
              (-234/11 +
               skoXC1 *
               (-2 + skoXC1 * (-6/11 + skoXC1 * -1/11)))))))))))) <= 0 |
     skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) <=
     -1440), inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (-1440 < skoXC1 * (skoXC1 * (-96 + skoXC1 * (skoXC1 * (6 + skoXC1)))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (345600/11 +
         skoXC1 *
         (-14400 +
          skoXC1 *
          (-197280/11 +
           skoXC1 *
           (-6480 +
            skoXC1 *
            (-1584 +
             skoXC1 *
             (-2664/11 +
              skoXC1 *
              (-234/11 +
               skoXC1 *
               (-2 + skoXC1 * (-6/11 + skoXC1 * -1/11)))))))))))) <= 0),
    inference(decision, [], [refute_0_18, refute_0_17])).

cnf(refute_0_83, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (345600/11 +
         skoXC1 *
         (-14400 +
          skoXC1 *
          (-197280/11 +
           skoXC1 *
           (-6480 +
            skoXC1 *
            (-1584 +
             skoXC1 *
             (-2664/11 +
              skoXC1 *
              (-234/11 +
               skoXC1 *
               (-2 + skoXC1 * (-6/11 + skoXC1 * -1/11)))))))))))) <= 0),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (28800/11 +
         skoXC1 *
         (14400/11 +
          skoXC1 *
          (121960/891 +
           skoXC1 *
           (-47000/891 +
            skoXC1 *
            (-104356/2673 +
             skoXC1 *
             (-3641086/216513 +
              skoXC1 *
              (-407203/72171 +
               skoXC1 *
               (-3419599/2165130 +
                skoXC1 *
                (-4061857/10628820 +
                 skoXC1 *
                 (-3394169/42515280 +
                  skoXC1 *
                  (-887213/63772920 +
                   skoXC1 *
                   (-4720681/2295825120 +
                    skoXC1 *
                    (-2989033/11479125600 +
                     skoXC1 *
                     (-260783/9183300480 +
                      skoXC1 *
                      (-5741/2152336050 +
                       skoXC1 *
                       (-5909/27549901440 +
                        skoXC1 *
                        (-4013/275499014400 +
                         skoXC1 *
                         (-13489/16529940864000 +
                          skoXC1 *
                          (-199/5509980288000 +
                           skoXC1 *
                           (-19/16529940864000 +
                            skoXC1 *
                            -1/49589822592000))))))))))))))))))))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(decision, [], [refute_0_83, refute_0_18, refute_0_17])).

cnf(refute_0_85, plain, (skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(resolve, [], [refute_0_64, refute_0_84])).

cnf(refute_0_86, plain, ($false),
    inference(resolve, [], [refute_0_85, refute_0_25])).
% SZS output end CNFRefutation for cosh-3612a-weak.tptp
