% SZS output start CNFRefutation for exp-upper-17.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf3, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120), Y) | lgen(R, exp(X), Y))).

fof(exp_upper_bound_problem_17, conjecture,
    (! [X] :
       ((-4 < X & X < 4) =>
        exp(X) <= (16 + X) ^ 16 / (16 ^ 15 * (16 - X ^ 2))))).

fof(subgoal_0, plain,
    (! [X] :
       ((-4 < X & X < 4) =>
        exp(X) <= (16 + X) ^ 16 / (16 ^ 15 * (16 - X ^ 2)))),
    inference(strip, [], [exp_upper_bound_problem_17])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((-4 < X & X < 4) =>
          exp(X) <= (16 + X) ^ 16 / (16 ^ 15 * (16 - X ^ 2)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-4 < X & X < 4 &
        (16 + X) ^ 16 / (16 ^ 15 * (16 - X ^ 2)) < exp(X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 4 &
     (16 + skoXC1) ^ 16 / (16 ^ 15 * (16 - skoXC1 ^ 2)) < exp(skoXC1) &
     -4 < skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    ((16 + skoXC1) ^ 16 / (16 ^ 15 * (16 - skoXC1 ^ 2)) < exp(skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (-4 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 4),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-2213609288845146193920 +
     skoXC1 *
     (-1106804644422573096960 +
      skoXC1 *
      (-83010348331692982272 +
       skoXC1 *
       (50728546202701266944 +
        skoXC1 * (13835058055282163712 + skoXC1 * 1152921504606846976)))) <
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     (-2213609288845146193920 +
      skoXC1 *
      (-1106804644422573096960 +
       skoXC1 *
       (-83010348331692982272 +
        skoXC1 *
        (50728546202701266944 +
         skoXC1 *
         (13835058055282163712 + skoXC1 * 1152921504606846976))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <=
     18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1)))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0 |
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    ((16 + skoXC1) ^ 16 / (16 ^ 15 * (16 - skoXC1 ^ 2)) < exp(skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     exp(skoXC1)), inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (~ lgen(0, exp(X_000061), X_000062) | exp(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_5, plain,
    (4 < X_000061 | X_000061 < 0 |
     ~ lgen(0,
            -(X_000061 ^ 3 + 12 * X_000061 ^ 2 + 60 * X_000061 + 120) /
            (X_000061 ^ 3 - 12 * X_000061 ^ 2 + 60 * X_000061 - 120),
            X_000062) | lgen(0, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_cf3])).

cnf(refute_0_6, plain,
    (4 < X_000061 | X_000061 < 0 |
     ~ lgen(0,
            -(X_000061 ^ 3 + 12 * X_000061 ^ 2 + 60 * X_000061 + 120) /
            (X_000061 ^ 3 - 12 * X_000061 ^ 2 + 60 * X_000061 - 120),
            X_000062) | exp(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (4 < X_000061 | X_000061 < 0 |
     X_000062 <
     (-120 + X_000061 * (-60 + X_000061 * (-12 + X_000061 * -1))) /
     (-120 + X_000061 * (60 + X_000061 * (-12 + X_000061))) |
     exp(X_000061) <= X_000062), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (skoXC1 < 0 |
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1 |
     exp(skoXC1) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (skoXC1 < 0 |
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1),
    inference(resolve, [], [refute_0_8, refute_0_3])).

cnf(refute_0_10, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     401092572728463209067316248576 +
     skoXC1 * (skoXC1 * -25068285795528950566707265536) |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0 |
     (401092572728463209067316248576 +
      skoXC1 * (skoXC1 * -25068285795528950566707265536)) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1)))))))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_11, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain,
    (~ lgen(0, exp(X_000040), X_000041) | exp(X_000040) <= X_000041),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (3 < X_000040 |
     ~ lgen(0,
            21743271936 /
            (-(X_000040 ^ 3) + 12 * X_000040 ^ 2 - 96 * X_000040 + 384) ^
            4, X_000041) | lgen(0, exp(X_000040), X_000041)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_14, plain,
    (3 < X_000040 |
     ~ lgen(0,
            21743271936 /
            (-(X_000040 ^ 3) + 12 * X_000040 ^ 2 - 96 * X_000040 + 384) ^
            4, X_000041) | exp(X_000040) <= X_000041),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (3 < X_000040 |
     X_000041 <
     21743271936 /
     (21743271936 +
      X_000040 *
      (-21743271936 +
       X_000040 *
       (10871635968 +
        X_000040 *
        (-3623878656 +
         X_000040 *
         (891813888 +
          X_000040 *
          (-169869312 +
           X_000040 *
           (25657344 +
            X_000040 *
            (-3096576 +
             X_000040 *
             (297216 +
              X_000040 *
              (-22272 +
               X_000040 * (1248 + X_000040 * (-48 + X_000040)))))))))))) |
     exp(X_000040) <= X_000041), inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 |
     exp(skoXC1) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1), inference(resolve, [], [refute_0_16, refute_0_3])).

cnf(refute_0_18, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     3 < skoXC1 |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0),
    inference(resolve, [], [refute_0_11, refute_0_17])).

cnf(refute_0_19, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (401092572728463209067316248576 +
      skoXC1 * (skoXC1 * -25068285795528950566707265536)) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 |
     skoXC1 * (skoXC1 * -1152921504606846976) <= -18446744073709551616),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <
     42501298345826806923264 +
     skoXC1 * (skoXC1 * -2656331146614175432704) |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0 |
     (42501298345826806923264 +
      skoXC1 * (skoXC1 * -2656331146614175432704)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <=
     18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1)))))))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_21, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_22, plain,
    (~ lgen(0, exp(X_000033), X_000034) | exp(X_000033) <= X_000034),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_23, plain,
    (0 < X_000033 |
     ~ lgen(0,
            2304 /
            (-(X_000033 ^ 3) + 6 * X_000033 ^ 2 - 24 * X_000033 + 48) ^ 2,
            X_000034) | lgen(0, exp(X_000033), X_000034)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_24, plain,
    (0 < X_000033 |
     ~ lgen(0,
            2304 /
            (-(X_000033 ^ 3) + 6 * X_000033 ^ 2 - 24 * X_000033 + 48) ^ 2,
            X_000034) | exp(X_000033) <= X_000034),
    inference(resolve, [], [refute_0_23, refute_0_22])).

cnf(refute_0_25, plain,
    (0 < X_000033 |
     X_000034 <
     2304 /
     (2304 +
      X_000033 *
      (-2304 +
       X_000033 *
       (1152 +
        X_000033 *
        (-384 + X_000033 * (84 + X_000033 * (-12 + X_000033)))))) |
     exp(X_000033) <= X_000034), inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1 |
     exp(skoXC1) <=
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976))),
    inference(subst, [], [refute_0_25])).

cnf(refute_0_27, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_26, refute_0_3])).

cnf(refute_0_28, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     2304 /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     0 < skoXC1 |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0),
    inference(resolve, [], [refute_0_21, refute_0_27])).

cnf(refute_0_29, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (42501298345826806923264 +
      skoXC1 * (skoXC1 * -2656331146614175432704)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1 |
     skoXC1 * (skoXC1 * -1152921504606846976) <= -18446744073709551616),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain, (-4 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_31, plain, (skoXC1 < 4),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_32, plain,
    (-18446744073709551616 < skoXC1 * (skoXC1 * -1152921504606846976) |
     0 < skoXC1), inference(decision, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (42501298345826806923264 +
      skoXC1 * (skoXC1 * -2656331146614175432704)) /
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_29, refute_0_32])).

cnf(refute_0_34, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (2304 +
      skoXC1 *
      (-2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1)))))) <
     42501298345826806923264 +
     skoXC1 * (skoXC1 * -2656331146614175432704) | 0 < skoXC1 |
     2304 +
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <= 0),
    inference(resolve, [], [refute_0_20, refute_0_33])).

cnf(refute_0_35, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1328165573307087716352 +
       skoXC1 *
       (-55340232221128654848 +
        skoXC1 *
        (203202415186956779520 +
         skoXC1 *
         (90144050141447847936 +
          skoXC1 *
          (11635049637311676416 +
           skoXC1 *
           (136515363704668160 +
            skoXC1 *
            (-366328687072641024 +
             skoXC1 *
             (-233673708693094400 +
              skoXC1 *
              (-84813612547309568 +
               skoXC1 *
               (-19233280163315712 +
                skoXC1 *
                (-2979551890112512 +
                 skoXC1 *
                 (-334207494651904 +
                  skoXC1 *
                  (-28084375388160 +
                   skoXC1 *
                   (-1800183808000 +
                    skoXC1 *
                    (-88562927872 +
                     skoXC1 *
                     (-3330045696 +
                      skoXC1 *
                      (-94233728 +
                       skoXC1 *
                       (-1946240 +
                        skoXC1 *
                        (-27732 +
                         skoXC1 *
                         (-244 + skoXC1 * -1))))))))))))))))))))) |
     0 < skoXC1 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-1328165573307087716352 +
       skoXC1 *
       (-55340232221128654848 +
        skoXC1 *
        (203202415186956779520 +
         skoXC1 *
         (90144050141447847936 +
          skoXC1 *
          (11635049637311676416 +
           skoXC1 *
           (136515363704668160 +
            skoXC1 *
            (-366328687072641024 +
             skoXC1 *
             (-233673708693094400 +
              skoXC1 *
              (-84813612547309568 +
               skoXC1 *
               (-19233280163315712 +
                skoXC1 *
                (-2979551890112512 +
                 skoXC1 *
                 (-334207494651904 +
                  skoXC1 *
                  (-28084375388160 +
                   skoXC1 *
                   (-1800183808000 +
                    skoXC1 *
                    (-88562927872 +
                     skoXC1 *
                     (-3330045696 +
                      skoXC1 *
                      (-94233728 +
                       skoXC1 *
                       (-1946240 +
                        skoXC1 *
                        (-27732 +
                         skoXC1 *
                         (-244 + skoXC1 * -1))))))))))))))))))))) <= 0 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(decision, [], [refute_0_30, refute_0_31])).

cnf(refute_0_37, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 +
       skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) <=
     -2304), inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (-2304 <
     skoXC1 *
     (-2304 +
      skoXC1 *
      (1152 + skoXC1 * (-384 + skoXC1 * (84 + skoXC1 * (-12 + skoXC1))))) |
     0 < skoXC1), inference(decision, [], [refute_0_30, refute_0_31])).

cnf(refute_0_39, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (-18446744073709551616 < skoXC1 * (skoXC1 * -1152921504606846976) |
     3 < skoXC1),
    inference(decision, [], [refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_41, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (401092572728463209067316248576 +
      skoXC1 * (skoXC1 * -25068285795528950566707265536)) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1), inference(resolve, [], [refute_0_19, refute_0_40])).

cnf(refute_0_42, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     401092572728463209067316248576 +
     skoXC1 * (skoXC1 * -25068285795528950566707265536) | 3 < skoXC1 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_41])).

cnf(refute_0_43, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-12534142897764475283353632768 +
       skoXC1 *
       (-522255954073519803473068032 +
        skoXC1 *
        (89762742106386216221933568 +
         skoXC1 *
         (67322056579789662166450176 +
          skoXC1 *
          (-1721302583005985680392192 +
           skoXC1 *
           (-1091752101258426102841344 +
            skoXC1 *
            (-30444045250648401248256 +
             skoXC1 *
             (-15962198231281796382720 +
              skoXC1 *
              (-3709849200245700820992 +
               skoXC1 *
               (-56029282964116340736 +
                skoXC1 *
                (37677747801279168512 +
                 skoXC1 *
                 (5011538815030919168 +
                  skoXC1 *
                  (884869915603894272 +
                   skoXC1 *
                   (135312979020218368 +
                    skoXC1 *
                    (9643726024474624 +
                     skoXC1 *
                     (-93161867182080 +
                      skoXC1 *
                      (-105068992921600 +
                       skoXC1 *
                       (-18905640730624 +
                        skoXC1 *
                        (-2627801186304 +
                         skoXC1 *
                         (-274582208512 +
                          skoXC1 *
                          (-20635942912 +
                           skoXC1 *
                           (-1106362368 +
                            skoXC1 *
                            (-42109184 +
                             skoXC1 *
                             (-1116416 +
                              skoXC1 *
                              (-19680 +
                               skoXC1 *
                               (-208 +
                                skoXC1 * -1))))))))))))))))))))))))))) |
     3 < skoXC1 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (3 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-12534142897764475283353632768 +
       skoXC1 *
       (-522255954073519803473068032 +
        skoXC1 *
        (89762742106386216221933568 +
         skoXC1 *
         (67322056579789662166450176 +
          skoXC1 *
          (-1721302583005985680392192 +
           skoXC1 *
           (-1091752101258426102841344 +
            skoXC1 *
            (-30444045250648401248256 +
             skoXC1 *
             (-15962198231281796382720 +
              skoXC1 *
              (-3709849200245700820992 +
               skoXC1 *
               (-56029282964116340736 +
                skoXC1 *
                (37677747801279168512 +
                 skoXC1 *
                 (5011538815030919168 +
                  skoXC1 *
                  (884869915603894272 +
                   skoXC1 *
                   (135312979020218368 +
                    skoXC1 *
                    (9643726024474624 +
                     skoXC1 *
                     (-93161867182080 +
                      skoXC1 *
                      (-105068992921600 +
                       skoXC1 *
                       (-18905640730624 +
                        skoXC1 *
                        (-2627801186304 +
                         skoXC1 *
                         (-274582208512 +
                          skoXC1 *
                          (-20635942912 +
                           skoXC1 *
                           (-1106362368 +
                            skoXC1 *
                            (-42109184 +
                             skoXC1 *
                             (-1116416 +
                              skoXC1 *
                              (-19680 +
                               skoXC1 *
                               (-208 +
                                skoXC1 * -1))))))))))))))))))))))))))) <=
     0 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936),
    inference(decision, [], [refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_45, plain,
    (3 < skoXC1 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (-21743271936 <
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 + skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) |
     3 < skoXC1),
    inference(decision, [], [refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_47, plain, (3 < skoXC1),
    inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain, (4 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_47, refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_49, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) | 4 < skoXC1),
    inference(resolve, [], [refute_0_48, refute_0_9])).

cnf(refute_0_50, plain, (skoXC1 <= 4),
    inference(decision, [],
              [refute_0_47, refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_51, plain,
    ((18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) /
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) *
     (18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976)) |
     18446744073709551616 + skoXC1 * (skoXC1 * -1152921504606846976) <= 0),
    inference(resolve, [], [refute_0_1, refute_0_51])).

cnf(refute_0_53, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (-2213609288845146193920 +
      skoXC1 *
      (-1106804644422573096960 +
       skoXC1 *
       (-83010348331692982272 +
        skoXC1 *
        (50728546202701266944 +
         skoXC1 *
         (13835058055282163712 + skoXC1 * 1152921504606846976))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     skoXC1 * (skoXC1 * -1152921504606846976) <= -18446744073709551616),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (-18446744073709551616 < skoXC1 * (skoXC1 * -1152921504606846976)),
    inference(decision, [],
              [refute_0_47, refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_55, plain,
    (18446744073709551616 +
     skoXC1 *
     (18446744073709551616 +
      skoXC1 *
      (8646911284551352320 +
       skoXC1 *
       (2522015791327477760 +
        skoXC1 *
        (512284457613393920 +
         skoXC1 *
         (76842668642009088 +
          skoXC1 *
          (8804889115230208 +
           skoXC1 *
           (786150813859840 +
            skoXC1 *
            (55276229099520 +
             skoXC1 *
             (3070901616640 +
              skoXC1 *
              (134351945728 +
               skoXC1 *
               (4580179968 +
                skoXC1 *
                (119275520 +
                 skoXC1 *
                 (2293760 +
                  skoXC1 * (30720 + skoXC1 * (256 + skoXC1))))))))))))))) <
     (-2213609288845146193920 +
      skoXC1 *
      (-1106804644422573096960 +
       skoXC1 *
       (-83010348331692982272 +
        skoXC1 *
        (50728546202701266944 +
         skoXC1 *
         (13835058055282163712 + skoXC1 * 1152921504606846976))))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (-2213609288845146193920 +
     skoXC1 *
     (-1106804644422573096960 +
      skoXC1 *
      (-83010348331692982272 +
       skoXC1 *
       (50728546202701266944 +
        skoXC1 * (13835058055282163712 + skoXC1 * 1152921504606846976)))) <
     (18446744073709551616 +
      skoXC1 *
      (18446744073709551616 +
       skoXC1 *
       (8646911284551352320 +
        skoXC1 *
        (2522015791327477760 +
         skoXC1 *
         (512284457613393920 +
          skoXC1 *
          (76842668642009088 +
           skoXC1 *
           (8804889115230208 +
            skoXC1 *
            (786150813859840 +
             skoXC1 *
             (55276229099520 +
              skoXC1 *
              (3070901616640 +
               skoXC1 *
               (134351945728 +
                skoXC1 *
                (4580179968 +
                 skoXC1 *
                 (119275520 +
                  skoXC1 *
                  (2293760 +
                   skoXC1 *
                   (30720 + skoXC1 * (256 + skoXC1)))))))))))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))),
    inference(resolve, [], [refute_0_0, refute_0_55])).

cnf(refute_0_57, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-69175290276410818560 +
       skoXC1 *
       (-37469948899722526720 +
        skoXC1 *
        (-9304436830147444736 +
         skoXC1 *
         (-1254252496222683136 +
          skoXC1 *
          (-71424275340328960 +
           skoXC1 *
           (24127683159916544 +
            skoXC1 *
            (11719900598894592 +
             skoXC1 *
             (2319144900886528 +
              skoXC1 *
              (290967928176640 +
               skoXC1 *
               (25936904847360 +
                skoXC1 *
                (1719176003584 +
                 skoXC1 *
                 (86271066112 +
                  skoXC1 *
                  (3282812928 +
                   skoXC1 *
                   (93562880 +
                    skoXC1 *
                    (1940360 +
                     skoXC1 *
                     (27708 + skoXC1 * (244 + skoXC1)))))))))))))))))) |
     skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 *
     (skoXC1 *
      (-69175290276410818560 +
       skoXC1 *
       (-37469948899722526720 +
        skoXC1 *
        (-9304436830147444736 +
         skoXC1 *
         (-1254252496222683136 +
          skoXC1 *
          (-71424275340328960 +
           skoXC1 *
           (24127683159916544 +
            skoXC1 *
            (11719900598894592 +
             skoXC1 *
             (2319144900886528 +
              skoXC1 *
              (290967928176640 +
               skoXC1 *
               (25936904847360 +
                skoXC1 *
                (1719176003584 +
                 skoXC1 *
                 (86271066112 +
                  skoXC1 *
                  (3282812928 +
                   skoXC1 *
                   (93562880 +
                    skoXC1 *
                    (1940360 +
                     skoXC1 *
                     (27708 + skoXC1 * (244 + skoXC1)))))))))))))))))) <=
     0 | skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(decision, [],
              [refute_0_47, refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_59, plain,
    (skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (-120 < skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1))),
    inference(decision, [],
              [refute_0_47, refute_0_39, refute_0_30, refute_0_31])).

cnf(refute_0_61, plain, ($false),
    inference(resolve, [], [refute_0_59, refute_0_60])).
% SZS output end CNFRefutation for exp-upper-17.tptp
