% SZS output start CNFRefutation for exp-361-neg-4.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

fof(a, axiom, (a = 4)).

fof(exp_361, conjecture,
    (! [X] : (0 < X => exp(-X) <= 7/200 + (a / (exp(1) * X)) ^ a))).

fof(subgoal_0, plain,
    (! [X] : (0 < X => exp(-X) <= 7/200 + (a / (exp(1) * X)) ^ a)),
    inference(strip, [], [exp_361])).

fof(negate_0_0, plain,
    (~ ! [X] : (0 < X => exp(-X) <= 7/200 + (a / (exp(1) * X)) ^ a)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (0 < X & 7/200 + (a / (exp(1) * X)) ^ a < exp(-X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a < exp(-skoXC1) & 0 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a < exp(-skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (a = 4), inference(canonicalize, [], [a])).

fof(normalize_0_4, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) |
     21743271936/7992538801 <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, exp(X_000061), X_000062) | exp(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(0, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_3, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | exp(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (3 < X_000061 |
     X_000062 <
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061)))))))))))) |
     exp(X_000061) <= X_000062), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1)))))))))))) |
     exp(1) <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) |
     21743271936/7992538801 <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1)))))))))))) |
     exp(1) <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025)))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_8, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) |
     21743271936/7992538801 <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_9, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1)))))))))))) |
     exp(1) <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025))))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_10, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     21743271936/7992538801 <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_11, plain,
    (-2304 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     -589824 +
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0 |
     (-589824 +
      skoXC1 *
      (-589824 +
       skoXC1 *
       (-294912 +
        skoXC1 *
        (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-2016/25 +
               skoXC1 *
               (-1008/25 +
                skoXC1 *
                (-336/25 +
                 skoXC1 *
                 (-147/50 +
                  skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <=
     -2304), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_12, plain,
    ((256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) *
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <
     2304 |
     2304 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     0 |
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <=
     (256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_13, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a < exp(-skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain, (a = 4),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_15, plain,
    ((4 / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ a),
    introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    ((4 / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ a | a != 4 |
     (4 / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (a != 4 | (4 / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    ((4 / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (a / (exp(1) * skoXC1) = a / (exp(1) * skoXC1)),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (a / (exp(1) * skoXC1) != a / (exp(1) * skoXC1) | a != 4 |
     a / (exp(1) * skoXC1) = 4 / (exp(1) * skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (a != 4 | a / (exp(1) * skoXC1) = 4 / (exp(1) * skoXC1)),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain, (a / (exp(1) * skoXC1) = 4 / (exp(1) * skoXC1)),
    inference(resolve, [], [refute_0_14, refute_0_21])).

cnf(refute_0_23, plain,
    ((a / (exp(1) * skoXC1)) ^ a = (a / (exp(1) * skoXC1)) ^ a),
    introduced(tautology, [refl])).

cnf(refute_0_24, plain,
    (a / (exp(1) * skoXC1) != 4 / (exp(1) * skoXC1) |
     (a / (exp(1) * skoXC1)) ^ a != (a / (exp(1) * skoXC1)) ^ a |
     (a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ a),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (a / (exp(1) * skoXC1) != 4 / (exp(1) * skoXC1) |
     (a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ a),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    ((a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ a),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_28, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_29, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_31, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain,
    ((4 / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ 4 |
     (a / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ a |
     (a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_33, plain,
    ((4 / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ 4 |
     (a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_26, refute_0_32])).

cnf(refute_0_34, plain,
    ((a / (exp(1) * skoXC1)) ^ a = (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_18, refute_0_33])).

cnf(refute_0_35, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a =
     7/200 + (a / (exp(1) * skoXC1)) ^ a), introduced(tautology, [refl])).

cnf(refute_0_36, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a !=
     7/200 + (a / (exp(1) * skoXC1)) ^ a |
     (a / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ 4 |
     7/200 + (a / (exp(1) * skoXC1)) ^ a =
     7/200 + (4 / (exp(1) * skoXC1)) ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    ((a / (exp(1) * skoXC1)) ^ a != (4 / (exp(1) * skoXC1)) ^ 4 |
     7/200 + (a / (exp(1) * skoXC1)) ^ a =
     7/200 + (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain,
    (7/200 + (a / (exp(1) * skoXC1)) ^ a =
     7/200 + (4 / (exp(1) * skoXC1)) ^ 4),
    inference(resolve, [], [refute_0_34, refute_0_37])).

cnf(refute_0_39, plain,
    (7/200 + (4 / (exp(1) * skoXC1)) ^ 4 < exp(-skoXC1) |
     7/200 + (a / (exp(1) * skoXC1)) ^ a !=
     7/200 + (4 / (exp(1) * skoXC1)) ^ 4 |
     exp(-skoXC1) <= 7/200 + (a / (exp(1) * skoXC1)) ^ a),
    introduced(tautology, [equality])).

cnf(refute_0_40, plain,
    (7/200 + (4 / (exp(1) * skoXC1)) ^ 4 < exp(-skoXC1) |
     exp(-skoXC1) <= 7/200 + (a / (exp(1) * skoXC1)) ^ a),
    inference(resolve, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain,
    (7/200 + (4 / (exp(1) * skoXC1)) ^ 4 < exp(-skoXC1)),
    inference(resolve, [], [refute_0_40, refute_0_13])).

cnf(refute_0_42, plain,
    (7/200 + exp(skoXC1 * -1) * -1 <
     -256 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    ((7/200 + exp(skoXC1 * -1) * -1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     -256 |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0 |
     -256 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <=
     7/200 + exp(skoXC1 * -1) * -1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_44, plain,
    ((7/200 + exp(skoXC1 * -1) * -1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     -256 |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_43, refute_0_42])).

cnf(refute_0_45, plain,
    (256 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200))))))) <
     exp(skoXC1 * -1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_47, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) <= 0 |
     exp(1) = 0), inference(split, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) <= 0),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) <= 0 | exp(1) = 0),
    inference(split, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) <= 0),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_52, plain, (0 < skoXC1 * (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [], [refute_0_14, refute_0_51])).

cnf(refute_0_53, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(resolve, [], [refute_0_50, refute_0_52])).

cnf(refute_0_54, plain,
    (256 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200))))))) <
     exp(skoXC1 * -1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))),
    inference(resolve, [], [refute_0_45, refute_0_53])).

cnf(refute_0_55, plain,
    ((256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     exp(skoXC1 * -1) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    inference(split, [], [refute_0_54])).

cnf(refute_0_56, plain,
    ((256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_55, refute_0_53])).

cnf(refute_0_57, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_58, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_59, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (0 < skoXC1 * -1 |
     (256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     2304 /
     (2304 +
      skoXC1 * -1 *
      (-2304 +
       skoXC1 * -1 *
       (1152 +
        skoXC1 * -1 *
        (-384 +
         skoXC1 * -1 * (84 + skoXC1 * -1 * (-12 + skoXC1 * -1)))))) |
     exp(skoXC1 * -1) <=
     (256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))),
    inference(subst, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (0 < skoXC1 * -1 |
     (256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     2304 /
     (2304 +
      skoXC1 * -1 *
      (-2304 +
       skoXC1 * -1 *
       (1152 +
        skoXC1 * -1 *
        (-384 +
         skoXC1 * -1 * (84 + skoXC1 * -1 * (-12 + skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_61, refute_0_56])).

cnf(refute_0_63, plain,
    (skoXC1 < 0 |
     (256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_14, refute_0_51])).

cnf(refute_0_65, plain,
    ((256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))))),
    inference(resolve, [], [refute_0_64, refute_0_63])).

cnf(refute_0_66, plain,
    ((256 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 7/200)))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) *
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <
     2304 |
     2304 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_12, refute_0_65])).

cnf(refute_0_67, plain,
    (-2304 <
     (-589824 +
      skoXC1 *
      (-589824 +
       skoXC1 *
       (-294912 +
        skoXC1 *
        (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-2016/25 +
               skoXC1 *
               (-1008/25 +
                skoXC1 *
                (-336/25 +
                 skoXC1 *
                 (-147/50 +
                  skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) |
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (-2304 <
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))),
    inference(decision, [], [refute_0_14, refute_0_51])).

cnf(refute_0_69, plain,
    (-2304 <
     (-589824 +
      skoXC1 *
      (-589824 +
       skoXC1 *
       (-294912 +
        skoXC1 *
        (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-2016/25 +
               skoXC1 *
               (-1008/25 +
                skoXC1 *
                (-336/25 +
                 skoXC1 *
                 (-147/50 +
                  skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (-2304 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))) <
     -589824 +
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_11, refute_0_69])).

cnf(refute_0_71, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <= 0),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))))),
    inference(resolve, [], [refute_0_71, refute_0_53])).

cnf(refute_0_73, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) < exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) <=
     0), inference(split, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) <
     0 |
     589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))),
    inference(split, [], [refute_0_72])).

cnf(refute_0_75, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) |
     589824 <
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256)))))),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) <=
     589824), inference(decision, [], [refute_0_51, refute_0_14])).

cnf(refute_0_77, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))),
    inference(resolve, [], [refute_0_76, refute_0_75])).

cnf(refute_0_78, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     0 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     0 <=
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_79, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_80, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_81, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_80, refute_0_79])).

cnf(refute_0_82, plain,
    (0 <
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) <=
     exp(1)), inference(subst, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (exp(1) <
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     0 <=
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))),
    inference(split, [], [refute_0_72])).

cnf(refute_0_84, plain,
    (exp(1) <
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (0 <
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(resolve, [], [refute_0_82, refute_0_84])).

cnf(refute_0_86, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     0 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     0 <=
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(resolve, [], [refute_0_78, refute_0_85])).

cnf(refute_0_87, plain,
    (589824 <
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_89, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-55584/25 +
          skoXC1 *
          (2016/25 +
           skoXC1 *
           (1008/25 +
            skoXC1 *
            (336/25 +
             skoXC1 *
             (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(split, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-55584/25 +
          skoXC1 *
          (2016/25 +
           skoXC1 *
           (1008/25 +
            skoXC1 *
            (336/25 +
             skoXC1 *
             (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(split, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))) <= 0),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) <
     0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(split, [], [refute_0_89])).

cnf(refute_0_93, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(resolve, [], [refute_0_91, refute_0_93])).

cnf(refute_0_95, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0 |
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256))))) <=
     589824),
    inference(decision, [],
              [refute_0_14, refute_0_51, refute_0_91, refute_0_94])).

cnf(refute_0_96, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-55584/25 +
            skoXC1 *
            (2016/25 +
             skoXC1 *
             (1008/25 +
              skoXC1 *
              (336/25 +
               skoXC1 *
               (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200)))))))))))) <=
     0), inference(resolve, [], [refute_0_95, refute_0_87])).

cnf(refute_0_97, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))),
    inference(resolve, [], [refute_0_96, refute_0_77])).

cnf(refute_0_98, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) < exp(1)),
    inference(resolve, [], [refute_0_73, refute_0_97])).

cnf(refute_0_99, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1)))))))))))) |
     exp(1) <=
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_100, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1))))))))))))),
    inference(resolve, [], [refute_0_99, refute_0_98])).

cnf(refute_0_101, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) <
     21743271936/7992538801), inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (55584/25 +
             skoXC1 *
             (-2016/25 +
              skoXC1 *
              (-1008/25 +
               skoXC1 *
               (-336/25 +
                skoXC1 *
                (-147/50 +
                 skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) <=
     0), inference(resolve, [], [refute_0_10, refute_0_101])).

cnf(refute_0_103, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (55584/25 +
            skoXC1 *
            (-2016/25 +
             skoXC1 *
             (-1008/25 +
              skoXC1 *
              (-336/25 +
               skoXC1 *
               (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025))))))))))))),
    inference(resolve, [], [refute_0_103, refute_0_97])).

cnf(refute_0_105, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) < exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) <= 0),
    inference(split, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) < 0 |
     589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025)))))))))))),
    inference(split, [], [refute_0_104])).

cnf(refute_0_107, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) |
     589824 <
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256)))))),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025)))))))))))),
    inference(resolve, [], [refute_0_76, refute_0_107])).

cnf(refute_0_109, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_110, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-1208578027290624/199813470025 +
         skoXC1 *
         (43834436222976/199813470025 +
          skoXC1 *
          (21917218111488/199813470025 +
           skoXC1 *
           (7305739370496/199813470025 +
            skoXC1 *
            (1598130487296/199813470025 +
             skoXC1 *
             (228304355328/199813470025 +
              skoXC1 * 19025362944/199813470025))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) <= 0),
    inference(split, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (55584/25 +
          skoXC1 *
          (-2016/25 +
           skoXC1 *
           (-1008/25 +
            skoXC1 *
            (-336/25 +
             skoXC1 *
             (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))))),
    inference(split, [], [refute_0_97])).

cnf(refute_0_112, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200))))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))),
    inference(split, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))),
    inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (55584/25 +
          skoXC1 *
          (-2016/25 +
           skoXC1 *
           (-1008/25 +
            skoXC1 *
            (-336/25 +
             skoXC1 *
             (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))),
    inference(split, [], [refute_0_111])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-55584/25 +
         skoXC1 *
         (2016/25 +
          skoXC1 *
          (1008/25 +
           skoXC1 *
           (336/25 +
            skoXC1 * (147/50 + skoXC1 * (21/50 + skoXC1 * 7/200))))))))) <=
     0), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (55584/25 +
         skoXC1 *
         (-2016/25 +
          skoXC1 *
          (-1008/25 +
           skoXC1 *
           (-336/25 +
            skoXC1 *
            (-147/50 + skoXC1 * (-21/50 + skoXC1 * -7/200)))))))))),
    inference(resolve, [], [refute_0_115, refute_0_113])).

cnf(refute_0_117, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-1208578027290624/199813470025 +
         skoXC1 *
         (43834436222976/199813470025 +
          skoXC1 *
          (21917218111488/199813470025 +
           skoXC1 *
           (7305739370496/199813470025 +
            skoXC1 *
            (1598130487296/199813470025 +
             skoXC1 *
             (228304355328/199813470025 +
              skoXC1 * 19025362944/199813470025))))))))) <= 0),
    inference(decision, [],
              [refute_0_116, refute_0_115, refute_0_51, refute_0_14])).

cnf(refute_0_118, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-1208578027290624/199813470025 +
           skoXC1 *
           (43834436222976/199813470025 +
            skoXC1 *
            (21917218111488/199813470025 +
             skoXC1 *
             (7305739370496/199813470025 +
              skoXC1 *
              (1598130487296/199813470025 +
               skoXC1 *
               (228304355328/199813470025 +
                skoXC1 * 19025362944/199813470025))))))))))) <= 0),
    inference(resolve, [], [refute_0_117, refute_0_110])).

cnf(refute_0_119, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025)))))))))))),
    inference(resolve, [], [refute_0_118, refute_0_108])).

cnf(refute_0_120, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) < exp(1)),
    inference(resolve, [], [refute_0_105, refute_0_119])).

cnf(refute_0_121, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1))))))))))))),
    inference(resolve, [], [refute_0_9, refute_0_120])).

cnf(refute_0_122, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) <
     21743271936/7992538801), inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (1208578027290624/199813470025 +
            skoXC1 *
            (-43834436222976/199813470025 +
             skoXC1 *
             (-21917218111488/199813470025 +
              skoXC1 *
              (-7305739370496/199813470025 +
               skoXC1 *
               (-1598130487296/199813470025 +
                skoXC1 *
                (-228304355328/199813470025 +
                 skoXC1 * -19025362944/199813470025)))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) <= 0),
    inference(resolve, [], [refute_0_8, refute_0_122])).

cnf(refute_0_124, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (1208578027290624/199813470025 +
           skoXC1 *
           (-43834436222976/199813470025 +
            skoXC1 *
            (-21917218111488/199813470025 +
             skoXC1 *
             (-7305739370496/199813470025 +
              skoXC1 *
              (-1598130487296/199813470025 +
               skoXC1 *
               (-228304355328/199813470025 +
                skoXC1 * -19025362944/199813470025))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025)))))))))))),
    inference(resolve, [], [refute_0_124, refute_0_119])).

cnf(refute_0_126, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) <
     exp(1) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(split, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) <
     0 |
     589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) < 0 |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025))))))))))),
    inference(split, [], [refute_0_125])).

cnf(refute_0_128, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) |
     589824 <
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 + skoXC1 * (-21504 + skoXC1 * (-3072 + skoXC1 * -256)))))),
    inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025))))))))))),
    inference(resolve, [], [refute_0_76, refute_0_128])).

cnf(refute_0_130, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_131, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(split, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (-953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (-476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (-158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (-34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (-4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              -413673640172489539584/1597016912137262940025))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(split, [], [refute_0_130])).

cnf(refute_0_134, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (-953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (-476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (-158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (-34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (-4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              -413673640172489539584/1597016912137262940025))))))))) <= 0),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(resolve, [], [refute_0_134, refute_0_132])).

cnf(refute_0_136, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-26278440703254466935128064/1597016912137262940025 +
         skoXC1 *
         (953104066957415899201536/1597016912137262940025 +
          skoXC1 *
          (476552033478707949600768/1597016912137262940025 +
           skoXC1 *
           (158850677826235983200256/1597016912137262940025 +
            skoXC1 *
            (34748585774489121325056/1597016912137262940025 +
             skoXC1 *
             (4964083682069874475008/1597016912137262940025 +
              skoXC1 *
              413673640172489539584/1597016912137262940025))))))))) <= 0),
    inference(decision, [],
              [refute_0_14, refute_0_51, refute_0_115, refute_0_116])).

cnf(refute_0_137, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(resolve, [], [refute_0_136, refute_0_135])).

cnf(refute_0_138, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025))))))))))),
    inference(resolve, [], [refute_0_137, refute_0_129])).

cnf(refute_0_139, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) <
     exp(1)), inference(resolve, [], [refute_0_126, refute_0_138])).

cnf(refute_0_140, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1))))))))))))),
    inference(resolve, [], [refute_0_7, refute_0_139])).

cnf(refute_0_141, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) <
     21743271936/7992538801), inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (26278440703254466935128064/1597016912137262940025 +
           skoXC1 *
           (-953104066957415899201536/1597016912137262940025 +
            skoXC1 *
            (-476552033478707949600768/1597016912137262940025 +
             skoXC1 *
             (-158850677826235983200256/1597016912137262940025 +
              skoXC1 *
              (-34748585774489121325056/1597016912137262940025 +
               skoXC1 *
               (-4964083682069874475008/1597016912137262940025 +
                skoXC1 *
                -413673640172489539584/1597016912137262940025))))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(resolve, [], [refute_0_6, refute_0_141])).

cnf(refute_0_143, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (26278440703254466935128064/1597016912137262940025 +
          skoXC1 *
          (-953104066957415899201536/1597016912137262940025 +
           skoXC1 *
           (-476552033478707949600768/1597016912137262940025 +
            skoXC1 *
            (-158850677826235983200256/1597016912137262940025 +
             skoXC1 *
             (-34748585774489121325056/1597016912137262940025 +
              skoXC1 *
              (-4964083682069874475008/1597016912137262940025 +
               skoXC1 *
               -413673640172489539584/1597016912137262940025)))))))))) <=
     0), inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025))))))))))),
    inference(resolve, [], [refute_0_143, refute_0_138])).

cnf(refute_0_145, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) < exp(1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0),
    inference(split, [], [refute_0_144])).

cnf(refute_0_146, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025)))))))))),
    inference(decision, [],
              [refute_0_116, refute_0_115, refute_0_51, refute_0_14])).

cnf(refute_0_147, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) < exp(1)),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    (3 < 1 |
     (589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) <
     21743271936 /
     (21743271936 +
      1 *
      (-21743271936 +
       1 *
       (10871635968 +
        1 *
        (-3623878656 +
         1 *
         (891813888 +
          1 *
          (-169869312 +
           1 *
           (25657344 +
            1 *
            (-3096576 +
             1 *
             (297216 + 1 * (-22272 + 1 * (1248 + 1 * (-48 + 1))))))))))))),
    inference(resolve, [], [refute_0_5, refute_0_147])).

cnf(refute_0_149, plain,
    ((589824 +
      skoXC1 *
      (589824 +
       skoXC1 *
       (294912 +
        skoXC1 *
        (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) <
     21743271936/7992538801), inference(arithmetic, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (589824 +
     skoXC1 *
     (589824 +
      skoXC1 *
      (294912 +
       skoXC1 *
       (98304 + skoXC1 * (21504 + skoXC1 * (3072 + skoXC1 * 256))))) <
     21743271936/7992538801 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (571379282264912954777209966537211904/
          12764219636110281886089148410025 +
          skoXC1 *
          (-20723600911162646028188962516893696/
           12764219636110281886089148410025 +
           skoXC1 *
           (-10361800455581323014094481258446848/
            12764219636110281886089148410025 +
            skoXC1 *
            (-3453933485193774338031493752815616/
             12764219636110281886089148410025 +
             skoXC1 *
             (-755547949886138136444389258428416/
              12764219636110281886089148410025 +
              skoXC1 *
              (-107935421412305448063484179775488/
               12764219636110281886089148410025 +
               skoXC1 *
               -8994618451025454005290348314624/
               12764219636110281886089148410025)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_149])).

cnf(refute_0_151, plain,
    (589824 <
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 +
        skoXC1 *
        (10229848843618583110148666050403106200428268544/
         102018520706097528689614980812172269880025 +
         skoXC1 *
         (-763999785713678398650720154653523839203552256/
          102018520706097528689614980812172269880025 +
          skoXC1 *
          (-251416186353034362602652901887181414155344128/
           102018520706097528689614980812172269880025 +
           skoXC1 *
           (-75099815017424465086037155599755104355352576/
            102018520706097528689614980812172269880025 +
            skoXC1 *
            (-16428084535061601737570627787446429077733376/
             102018520706097528689614980812172269880025 +
             skoXC1 *
             (-2346869219294514533938661112492347011104768/
              102018520706097528689614980812172269880025 +
              skoXC1 *
              -195572434941209544494888426041028917592064/
              102018520706097528689614980812172269880025))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0 |
     skoXC1 *
     (-589824 +
      skoXC1 *
      (-294912 +
       skoXC1 *
       (-98304 +
        skoXC1 *
        (10229848843618583110148666050403106200428268544/
         102018520706097528689614980812172269880025 +
         skoXC1 *
         (-763999785713678398650720154653523839203552256/
          102018520706097528689614980812172269880025 +
          skoXC1 *
          (-251416186353034362602652901887181414155344128/
           102018520706097528689614980812172269880025 +
           skoXC1 *
           (-75099815017424465086037155599755104355352576/
            102018520706097528689614980812172269880025 +
            skoXC1 *
            (-16428084535061601737570627787446429077733376/
             102018520706097528689614980812172269880025 +
             skoXC1 *
             (-2346869219294514533938661112492347011104768/
              102018520706097528689614980812172269880025 +
              skoXC1 *
              -195572434941209544494888426041028917592064/
              102018520706097528689614980812172269880025))))))))) <=
     589824),
    inference(decision, [],
              [refute_0_116, refute_0_115, refute_0_51, refute_0_14])).

cnf(refute_0_153, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (571379282264912954777209966537211904/
         12764219636110281886089148410025 +
         skoXC1 *
         (-20723600911162646028188962516893696/
          12764219636110281886089148410025 +
          skoXC1 *
          (-10361800455581323014094481258446848/
           12764219636110281886089148410025 +
           skoXC1 *
           (-3453933485193774338031493752815616/
            12764219636110281886089148410025 +
            skoXC1 *
            (-755547949886138136444389258428416/
             12764219636110281886089148410025 +
             skoXC1 *
             (-107935421412305448063484179775488/
              12764219636110281886089148410025 +
              skoXC1 *
              -8994618451025454005290348314624/
              12764219636110281886089148410025))))))))) <= 0),
    inference(resolve, [], [refute_0_152, refute_0_151])).

cnf(refute_0_154, plain, ($false),
    inference(resolve, [], [refute_0_153, refute_0_146])).
% SZS output end CNFRefutation for exp-361-neg-4.tptp
