% SZS output start CNFRefutation for exp-problem-10-2-sqrt-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf1, axiom,
    (X < 0 | 3/2 < X | ~ lgen(R, -(X + 2) / (X - 2), Y) |
     lgen(R, exp(X), Y))).

cnf(sqrt_upper_bound_1, axiom,
    (X < 0 | ~ lgen(R, 1/4 * (X ^ 2 + 6 * X + 1) / (X + 1), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_1, axiom,
    (X < 0 | ~ lgen(R, Y, 4 * X * (X + 1) / (X ^ 2 + 6 * X + 1)) |
     lgen(R, Y, sqrt(X)))).

fof(exp_problem_10_2, conjecture,
    (! [X] :
       (2 < X =>
        -2/3 <=
        exp(1 / sqrt(X - 1)) - 2 * exp(1 / sqrt(X)) +
        exp(1 / sqrt(X + 1))))).

fof(subgoal_0, plain,
    (! [X] :
       (2 < X =>
        -2/3 <=
        exp(1 / sqrt(X - 1)) - 2 * exp(1 / sqrt(X)) +
        exp(1 / sqrt(X + 1)))), inference(strip, [], [exp_problem_10_2])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (2 < X =>
          -2/3 <=
          exp(1 / sqrt(X - 1)) - 2 * exp(1 / sqrt(X)) +
          exp(1 / sqrt(X + 1)))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (2 < X &
        exp(1 / sqrt(X - 1)) - 2 * exp(1 / sqrt(X)) +
        exp(1 / sqrt(X + 1)) < -2/3)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (exp(1 / sqrt(skoXC1 - 1)) - 2 * exp(1 / sqrt(skoXC1)) +
     exp(1 / sqrt(skoXC1 + 1)) < -2/3 & 2 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (exp(1 / sqrt(skoXC1 - 1)) - 2 * exp(1 / sqrt(skoXC1)) +
     exp(1 / sqrt(skoXC1 + 1)) < -2/3),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (2 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) * (-4 + skoXC1 * (4 + skoXC1)) <
     skoXC1 *
     (352/3 + skoXC1 * (128/3 + skoXC1 * (-340/3 + skoXC1 * -140/3))) |
     -4 + skoXC1 * (4 + skoXC1) <= 0 |
     skoXC1 *
     (352/3 + skoXC1 * (128/3 + skoXC1 * (-340/3 + skoXC1 * -140/3))) /
     (-4 + skoXC1 * (4 + skoXC1)) <= 8 + skoXC1 * (12 + skoXC1 * 4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) *
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     0 <= -88/3 + skoXC1 * (-40 + skoXC1 * -35/3) |
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) <=
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000062, sqrt(X_000061)) | X_000062 <= sqrt(X_000061)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062,
            4 * X_000061 * (X_000061 + 1) /
            (X_000061 ^ 2 + 6 * X_000061 + 1)) |
     lgen(0, X_000062, sqrt(X_000061))),
    inference(subst, [], [sqrt_lower_bound_1])).

cnf(refute_0_4, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062,
            4 * X_000061 * (X_000061 + 1) /
            (X_000061 ^ 2 + 6 * X_000061 + 1)) |
     X_000062 <= sqrt(X_000061)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000061 * (4 + X_000061 * 4) / (1 + X_000061 * (6 + X_000061)) <
     X_000062 | X_000061 < 0 | X_000062 <= sqrt(X_000061)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) <= sqrt(-1 + skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (sqrt(-1 + skoXC1) * (8 + skoXC1 * (8 + skoXC1)) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) * (-1 + sqrt(-1 + skoXC1) * -8/3) |
     0 <= -1 + sqrt(-1 + skoXC1) * -8/3 |
     sqrt(-1 + skoXC1) * (8 + skoXC1 * (8 + skoXC1)) /
     (-1 + sqrt(-1 + skoXC1) * -8/3) <= 8 + skoXC1 * (12 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_8, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) *
     (8 + skoXC1 * (8 + skoXC1)) | 8 + skoXC1 * (8 + skoXC1) <= 0 |
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) <=
     (8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_9, plain,
    (1 + skoXC1 < 0 |
     (1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) |
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) <=
     sqrt(1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_10, plain,
    ((1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) *
     sqrt(-1 + skoXC1) < sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -1) |
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -1) /
     sqrt(-1 + skoXC1) <=
     1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) |
     sqrt(-1 + skoXC1) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_11, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     (1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1) *
     (sqrt(1 + skoXC1) * -1) | 0 <= sqrt(1 + skoXC1) * -1 |
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) <=
     (1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_12, plain,
    (~ lgen(0, X_000048, exp(X_000047)) | X_000048 <= exp(X_000047)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (X_000047 < -1 | ~ lgen(0, X_000048, 1 + X_000047) |
     lgen(0, X_000048, exp(X_000047))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_14, plain,
    (X_000047 < -1 | ~ lgen(0, X_000048, 1 + X_000047) |
     X_000048 <= exp(X_000047)),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (1 + X_000047 < X_000048 | X_000047 < -1 | X_000048 <= exp(X_000047)),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (1 / sqrt(-1 + skoXC1) < -1 |
     1 + 1 / sqrt(-1 + skoXC1) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) |
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) <= exp(1 / sqrt(-1 + skoXC1))),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (1 <
     (-5/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1) *
     sqrt(1 + skoXC1) |
     -5/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1 <=
     1 / sqrt(1 + skoXC1) | sqrt(1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_18, plain,
    (exp(1 / sqrt(skoXC1 - 1)) - 2 * exp(1 / sqrt(skoXC1)) +
     exp(1 / sqrt(skoXC1 + 1)) < -2/3),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_19, plain,
    (exp(1 / sqrt(1 + skoXC1)) <
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (1 / sqrt(1 + skoXC1) < -1 |
     1 + 1 / sqrt(1 + skoXC1) <
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1 |
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1 <=
     exp(1 / sqrt(1 + skoXC1))), inference(subst, [], [refute_0_15])).

cnf(refute_0_21, plain,
    (1 / sqrt(1 + skoXC1) < -1 |
     1 + 1 / sqrt(1 + skoXC1) <
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (1 / sqrt(1 + skoXC1) < -1 |
     1 / sqrt(1 + skoXC1) <
     -5/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) < -1 * (8 + skoXC1 * (8 + skoXC1)) |
     -1 <= (8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) |
     8 + skoXC1 * (8 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_24, plain,
    ((1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) < -1 | 1 + skoXC1 < 0 |
     -1 <= sqrt(1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_25, plain,
    (1 < -1 * sqrt(1 + skoXC1) | -1 <= 1 / sqrt(1 + skoXC1) |
     sqrt(1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_26, plain,
    (sqrt(1 + skoXC1) < -1 | -1 <= 1 / sqrt(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain, (2 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_28, plain, (-1 < skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (sqrt(1 + skoXC1) < -1 | -1 <= 1 / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_26, refute_0_28])).

cnf(refute_0_30, plain,
    ((1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) < -1 | 1 + skoXC1 < 0 |
     -1 <= 1 / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_24, refute_0_29])).

cnf(refute_0_31, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) < -1 |
     skoXC1 < -1 | -1 <= 1 / sqrt(1 + skoXC1)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain, (-1 <= skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_33, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) < -1 |
     -1 <= 1 / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) < -1 * (8 + skoXC1 * (8 + skoXC1)) |
     -1 <= 1 / sqrt(1 + skoXC1) | 8 + skoXC1 * (8 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_23, refute_0_33])).

cnf(refute_0_35, plain,
    (16 < skoXC1 * (-20 + skoXC1 * -5) | -1 <= 1 / sqrt(1 + skoXC1) |
     skoXC1 * (8 + skoXC1) <= -8),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 * (-20 + skoXC1 * -5) <= 16 | skoXC1 * (8 + skoXC1) <= -8),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_37, plain,
    (-1 <= 1 / sqrt(1 + skoXC1) | skoXC1 * (8 + skoXC1) <= -8),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain, (-8 < skoXC1 * (8 + skoXC1)),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_39, plain, (-1 <= 1 / sqrt(1 + skoXC1)),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (1 / sqrt(1 + skoXC1) <
     -5/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1),
    inference(resolve, [], [refute_0_39, refute_0_22])).

cnf(refute_0_41, plain,
    (1 <
     (-5/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1) *
     sqrt(1 + skoXC1) | sqrt(1 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_17, refute_0_40])).

cnf(refute_0_42, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     exp(1 / sqrt(-1 + skoXC1)) * (sqrt(1 + skoXC1) * -1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     exp(1 / sqrt(-1 + skoXC1)) * (sqrt(1 + skoXC1) * -1)),
    inference(resolve, [], [refute_0_42, refute_0_28])).

cnf(refute_0_44, plain,
    (exp(1 / sqrt(-1 + skoXC1)) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) | 0 <= sqrt(1 + skoXC1) * -1),
    inference(split, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (exp(1 / sqrt(-1 + skoXC1)) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (exp(1 / sqrt(-1 + skoXC1)) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1)),
    inference(resolve, [], [refute_0_45, refute_0_28])).

cnf(refute_0_47, plain,
    (1 / sqrt(-1 + skoXC1) < -1 |
     1 + 1 / sqrt(-1 + skoXC1) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1)),
    inference(resolve, [], [refute_0_16, refute_0_46])).

cnf(refute_0_48, plain,
    (1 / sqrt(-1 + skoXC1) < -1 |
     (1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1) | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain, (1 < skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_50, plain,
    (1 / sqrt(-1 + skoXC1) < -1 |
     (1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1)),
    inference(resolve, [], [refute_0_48, refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 * (-4 + skoXC1 * 4) < -1 * (-4 + skoXC1 * (4 + skoXC1)) |
     -1 <= skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_52, plain,
    (1 < -1 * sqrt(-1 + skoXC1) | -1 <= 1 / sqrt(-1 + skoXC1) |
     sqrt(-1 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_53, plain,
    (sqrt(-1 + skoXC1) < -1 | -1 <= 1 / sqrt(-1 + skoXC1) | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (sqrt(-1 + skoXC1) < -1 | -1 <= 1 / sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_53, refute_0_49])).

cnf(refute_0_55, plain,
    ((-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) < -1 | -1 + skoXC1 < 0 |
     -1 <= sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_56, plain,
    ((-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) < -1 | -1 + skoXC1 < 0 |
     -1 <= 1 / sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_55, refute_0_54])).

cnf(refute_0_57, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) < -1 |
     skoXC1 < 1 | -1 <= 1 / sqrt(-1 + skoXC1)),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain, (1 <= skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_59, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) < -1 |
     -1 <= 1 / sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (skoXC1 * (-4 + skoXC1 * 4) < -1 * (-4 + skoXC1 * (4 + skoXC1)) |
     -1 <= 1 / sqrt(-1 + skoXC1) | -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_51, refute_0_59])).

cnf(refute_0_61, plain,
    (-4 < skoXC1 * (skoXC1 * -5) | -1 <= 1 / sqrt(-1 + skoXC1) |
     skoXC1 * (4 + skoXC1) <= 4),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 * (skoXC1 * -5) <= -4 | skoXC1 * (4 + skoXC1) <= 4),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_63, plain,
    (-1 <= 1 / sqrt(-1 + skoXC1) | skoXC1 * (4 + skoXC1) <= 4),
    inference(resolve, [], [refute_0_62, refute_0_61])).

cnf(refute_0_64, plain, (4 < skoXC1 * (4 + skoXC1)),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_65, plain, (-1 <= 1 / sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    ((1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1) <
     (1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) /
     (sqrt(1 + skoXC1) * -1)),
    inference(resolve, [], [refute_0_65, refute_0_50])).

cnf(refute_0_67, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     (1 + sqrt(-1 + skoXC1)) / sqrt(-1 + skoXC1) *
     (sqrt(1 + skoXC1) * -1) | 0 <= sqrt(1 + skoXC1) * -1),
    inference(resolve, [], [refute_0_11, refute_0_66])).

cnf(refute_0_68, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -1) / sqrt(-1 + skoXC1) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (1 + sqrt(1 + skoXC1) * 5/3 +
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2) <
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -1) / sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_68, refute_0_28])).

cnf(refute_0_70, plain,
    ((1 + sqrt(1 + skoXC1) * 5/3 +
      exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * -2)) *
     sqrt(-1 + skoXC1) < sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -1) |
     sqrt(-1 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_69])).

cnf(refute_0_71, plain,
    (sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     exp(1 / sqrt(skoXC1)) * (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2))),
    inference(resolve, [], [refute_0_71, refute_0_49])).

cnf(refute_0_73, plain,
    (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) < 0 |
     sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     0 | 0 < sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)),
    inference(split, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (0 < sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) |
     0 < sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) |
     sqrt(-1 + skoXC1) <
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (0 < sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_76, plain,
    (sqrt(1 + skoXC1) < 0 / (sqrt(-1 + skoXC1) * -2) |
     0 <= sqrt(-1 + skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0),
    inference(split, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (0 < 0 | sqrt(-1 + skoXC1) * -2 < 0 | 0 < sqrt(-1 + skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0),
    inference(split, [], [refute_0_75])).

cnf(refute_0_79, plain,
    (1 < skoXC1 | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain, (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * -2) <= 0),
    inference(resolve, [], [refute_0_77, refute_0_79])).

cnf(refute_0_81, plain,
    (0 < sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) |
     sqrt(-1 + skoXC1) <
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_80, refute_0_74])).

cnf(refute_0_82, plain,
    ((4 + skoXC1 * (44/3 + skoXC1 * -59/3)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (32 + skoXC1 * (64/3 + skoXC1 * (-40 + skoXC1 * -88/3))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 * (32 + skoXC1 * (64/3 + skoXC1 * (-40 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) <= 4 + skoXC1 * (44/3 + skoXC1 * -59/3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_83, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     0 <= 8 + skoXC1 * (-8/3 + skoXC1 * -22/3) |
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_84, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) <= sqrt(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_85, plain,
    (4 + skoXC1 * (-4 + skoXC1 * -1) +
     sqrt(skoXC1) * (-8 + skoXC1 * (8 + skoXC1 * 2)) <
     skoXC1 * (-4 + skoXC1 * 4) * (14/3 + sqrt(skoXC1) * -4/3) |
     0 <= 14/3 + sqrt(skoXC1) * -4/3 |
     (4 + skoXC1 * (-4 + skoXC1 * -1) +
      sqrt(skoXC1) * (-8 + skoXC1 * (8 + skoXC1 * 2))) /
     (14/3 + sqrt(skoXC1) * -4/3) <= skoXC1 * (-4 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_86, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) *
     (-4 + skoXC1 * (4 + skoXC1)) | -4 + skoXC1 * (4 + skoXC1) <= 0 |
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <=
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_87, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <=
     sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_88, plain,
    (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * (-2 + sqrt(skoXC1) * -4)) <
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) *
     (1 + sqrt(skoXC1) * -2) | 0 <= 1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * (-2 + sqrt(skoXC1) * -4)) /
     (1 + sqrt(skoXC1) * -2) <=
     sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_89, plain,
    (sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) *
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) <=
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_90, plain,
    ((sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) < exp(1 / sqrt(skoXC1)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_72])).

cnf(refute_0_91, plain,
    (~ lgen(0, exp(X_000103), X_000104) | exp(X_000103) <= X_000104),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_92, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, -(X_000103 + 2) / (X_000103 - 2), X_000104) |
     lgen(0, exp(X_000103), X_000104)),
    inference(subst, [], [exp_upper_bound_cf1])).

cnf(refute_0_93, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, -(X_000103 + 2) / (X_000103 - 2), X_000104) |
     exp(X_000103) <= X_000104),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     X_000104 < (-2 + X_000103 * -1) / (-2 + X_000103) |
     exp(X_000103) <= X_000104), inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-2 + 1 / sqrt(skoXC1) * -1) / (-2 + 1 / sqrt(skoXC1)) |
     exp(1 / sqrt(skoXC1)) <=
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2))),
    inference(subst, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-2 + 1 / sqrt(skoXC1) * -1) / (-2 + 1 / sqrt(skoXC1)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_95, refute_0_90])).

cnf(refute_0_97, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain, (0 < skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_99, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (1 < 0 * sqrt(skoXC1) | 0 <= 1 / sqrt(skoXC1) | sqrt(skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_101, plain, (0 <= 1 / sqrt(skoXC1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain, (0 <= 1 / sqrt(skoXC1)),
    inference(resolve, [], [refute_0_101, refute_0_98])).

cnf(refute_0_103, plain,
    (3/2 < 1 / sqrt(skoXC1) |
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_102, refute_0_99])).

cnf(refute_0_104, plain,
    (skoXC1 * (4 + skoXC1 * 4) < 2/3 * (1 + skoXC1 * (6 + skoXC1)) |
     2/3 <= skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_105, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 2/3 |
     2/3 <= sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_106, plain,
    (3/2 * sqrt(skoXC1) < 1 | 1 / sqrt(skoXC1) <= 3/2 | sqrt(skoXC1) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_107, plain,
    (sqrt(skoXC1) < 2/3 | 1 / sqrt(skoXC1) <= 3/2 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain, (sqrt(skoXC1) < 2/3 | 1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_107, refute_0_98])).

cnf(refute_0_109, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 2/3 |
     1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_105, refute_0_108])).

cnf(refute_0_110, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_111, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 2/3 |
     1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_110, refute_0_109])).

cnf(refute_0_112, plain,
    (skoXC1 * (4 + skoXC1 * 4) < 2/3 * (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0 | 1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_104, refute_0_111])).

cnf(refute_0_113, plain,
    (-2/3 < skoXC1 * (skoXC1 * -10/3) | skoXC1 * (6 + skoXC1) <= -1 |
     1 / sqrt(skoXC1) <= 3/2), inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 * (skoXC1 * -10/3) <= -2/3 | skoXC1 * (6 + skoXC1) <= -1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_115, plain,
    (skoXC1 * (6 + skoXC1) <= -1 | 1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_114, refute_0_113])).

cnf(refute_0_116, plain, (-1 < skoXC1 * (6 + skoXC1)),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_117, plain, (1 / sqrt(skoXC1) <= 3/2),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    ((sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) /
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_117, refute_0_103])).

cnf(refute_0_119, plain,
    (sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) *
     (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_89, refute_0_118])).

cnf(refute_0_120, plain,
    (sqrt(-1 + skoXC1) + sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3) <
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * (-2 + sqrt(skoXC1) * -4)) /
     (1 + sqrt(skoXC1) * -2) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * (-2 + sqrt(skoXC1) * -4)) <
     (sqrt(-1 + skoXC1) +
      sqrt(1 + skoXC1) * (1 + sqrt(-1 + skoXC1) * 8/3)) *
     (1 + sqrt(skoXC1) * -2) | 0 <= 1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_120])).

cnf(refute_0_122, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) <
     sqrt(1 + skoXC1) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) <= 1/2), inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (1/2 < skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) |
     skoXC1 * (4 + skoXC1 * 4) <= 1/2 * (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_124, plain,
    (X_000062 < sqrt(X_000061) | ~ lgen(1, X_000062, sqrt(X_000061))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_125, plain,
    (X_000061 < 0 |
     ~ lgen(1, X_000062,
            4 * X_000061 * (X_000061 + 1) /
            (X_000061 ^ 2 + 6 * X_000061 + 1)) |
     lgen(1, X_000062, sqrt(X_000061))),
    inference(subst, [], [sqrt_lower_bound_1])).

cnf(refute_0_126, plain,
    (X_000061 < 0 | X_000062 < sqrt(X_000061) |
     ~ lgen(1, X_000062,
            4 * X_000061 * (X_000061 + 1) /
            (X_000061 ^ 2 + 6 * X_000061 + 1))),
    inference(resolve, [], [refute_0_125, refute_0_124])).

cnf(refute_0_127, plain,
    (X_000061 < 0 | X_000062 < sqrt(X_000061) |
     X_000061 * (4 + X_000061 * 4) / (1 + X_000061 * (6 + X_000061)) <=
     X_000062), inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    (skoXC1 < 0 | 1/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <= 1/2),
    inference(subst, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (1/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <= 1/2),
    inference(resolve, [], [refute_0_110, refute_0_128])).

cnf(refute_0_130, plain,
    (1/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) <= 1/2 * (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_123])).

cnf(refute_0_131, plain,
    (1/2 < sqrt(skoXC1) | skoXC1 * (1 + skoXC1 * 7/2) <= 1/2 |
     skoXC1 * (6 + skoXC1) <= -1),
    inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (1/2 < skoXC1 * (1 + skoXC1 * 7/2) | skoXC1 * (6 + skoXC1) <= -1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_133, plain,
    (1/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (1/2 < sqrt(skoXC1)),
    inference(resolve, [], [refute_0_133, refute_0_116])).

cnf(refute_0_135, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) <
     sqrt(1 + skoXC1) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_122, refute_0_134])).

cnf(refute_0_136, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) < 0 |
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) < 0 |
     0 <
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_135])).

cnf(refute_0_137, plain,
    (1 + sqrt(skoXC1) * -2 <
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) |
     0 < sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) |
     -1 + sqrt(skoXC1) * 2 <
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain,
    (0 < sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) |
     sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_139, plain,
    (sqrt(-1 + skoXC1) < 0 / (1 + sqrt(skoXC1) * -2) |
     0 <= 1 + sqrt(skoXC1) * -2 |
     sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) <= 0),
    inference(split, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) <= 0 |
     sqrt(skoXC1) <= 1/2), inference(arithmetic, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (sqrt(-1 + skoXC1) * (1 + sqrt(skoXC1) * -2) <= 0),
    inference(resolve, [], [refute_0_140, refute_0_134])).

cnf(refute_0_142, plain,
    (1 + sqrt(skoXC1) * -2 <
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) |
     -1 + sqrt(skoXC1) * 2 <
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_141, refute_0_137])).

cnf(refute_0_143, plain,
    ((32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_144, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_145, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) <=
     sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_146, plain,
    (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
     sqrt(skoXC1) *
     (-64 + skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <=
     skoXC1 * (-4 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_147, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) | -4 + skoXC1 * (4 + skoXC1) <= 0 |
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <=
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_148, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <=
     sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_149, plain,
    (sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     0 <=
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <=
     8 + skoXC1 * (12 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_150, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) *
     (8 + skoXC1 * (8 + skoXC1)) | 8 + skoXC1 * (8 + skoXC1) <= 0 |
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <=
     (8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_151, plain,
    (1 + skoXC1 < 0 |
     (1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <=
     sqrt(1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_152, plain,
    (sqrt(1 + skoXC1) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     0 <=
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_135])).

cnf(refute_0_153, plain,
    (sqrt(1 + skoXC1) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (1 + skoXC1 < 0 |
     (1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_151, refute_0_153])).

cnf(refute_0_155, plain,
    (skoXC1 < -1 |
     (8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_32, refute_0_155])).

cnf(refute_0_157, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) *
     (8 + skoXC1 * (8 + skoXC1)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     8 + skoXC1 * (8 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_150, refute_0_156])).

cnf(refute_0_158, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    (sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     0 <=
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_149, refute_0_158])).

cnf(refute_0_160, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) <
     sqrt(-1 + skoXC1) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (7/2 * (1 + skoXC1) < 1/4 + skoXC1 * (3/2 + skoXC1 * 1/4) |
     1 + skoXC1 <= 0 |
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) <= 7/2),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_164, plain,
    (~ lgen(0, sqrt(X_000054), X_000055) | sqrt(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_165, plain,
    (X_000054 < 0 |
     ~ lgen(0, 1/4 * (X_000054 ^ 2 + 6 * X_000054 + 1) / (X_000054 + 1),
            X_000055) | lgen(0, sqrt(X_000054), X_000055)),
    inference(subst, [], [sqrt_upper_bound_1])).

cnf(refute_0_166, plain,
    (X_000054 < 0 |
     ~ lgen(0, 1/4 * (X_000054 ^ 2 + 6 * X_000054 + 1) / (X_000054 + 1),
            X_000055) | sqrt(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_165, refute_0_164])).

cnf(refute_0_167, plain,
    (X_000054 < 0 |
     X_000055 <
     (1/4 + X_000054 * (3/2 + X_000054 * 1/4)) / (1 + X_000054) |
     sqrt(X_000054) <= X_000055),
    inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain,
    (skoXC1 < 0 |
     7/2 < (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(skoXC1) <= 7/2), inference(subst, [], [refute_0_167])).

cnf(refute_0_169, plain,
    (skoXC1 * (4 + skoXC1 * 4) < 7/2 * (1 + skoXC1 * (6 + skoXC1)) |
     7/2 <= skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_170, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 7/2 |
     7/2 <= sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_171, plain,
    ((1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) < 7/2 |
     7/2 * (1 + skoXC1) <= 1/4 + skoXC1 * (3/2 + skoXC1 * 1/4) |
     1 + skoXC1 <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_172, plain,
    (sqrt(X_000054) < X_000055 | ~ lgen(1, sqrt(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_173, plain,
    (X_000054 < 0 |
     ~ lgen(1, 1/4 * (X_000054 ^ 2 + 6 * X_000054 + 1) / (X_000054 + 1),
            X_000055) | lgen(1, sqrt(X_000054), X_000055)),
    inference(subst, [], [sqrt_upper_bound_1])).

cnf(refute_0_174, plain,
    (sqrt(X_000054) < X_000055 | X_000054 < 0 |
     ~ lgen(1, 1/4 * (X_000054 ^ 2 + 6 * X_000054 + 1) / (X_000054 + 1),
            X_000055)),
    inference(resolve, [], [refute_0_173, refute_0_172])).

cnf(refute_0_175, plain,
    (sqrt(X_000054) < X_000055 | X_000054 < 0 |
     X_000055 <=
     (1/4 + X_000054 * (3/2 + X_000054 * 1/4)) / (1 + X_000054)),
    inference(arithmetic, [], [refute_0_174])).

cnf(refute_0_176, plain,
    (skoXC1 < 0 | sqrt(skoXC1) < 7/2 |
     7/2 <= (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1)),
    inference(subst, [], [refute_0_175])).

cnf(refute_0_177, plain,
    (7/2 < skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) |
     skoXC1 * (4 + skoXC1 * 4) <= 7/2 * (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_178, plain,
    (skoXC1 < 0 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <= 7/2),
    inference(subst, [], [refute_0_127])).

cnf(refute_0_179, plain,
    (7/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <= 7/2),
    inference(resolve, [], [refute_0_110, refute_0_178])).

cnf(refute_0_180, plain,
    (7/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 4) <= 7/2 * (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_179, refute_0_177])).

cnf(refute_0_181, plain,
    (7/2 < sqrt(skoXC1) | skoXC1 * (-17 + skoXC1 * 1/2) <= 7/2 |
     skoXC1 * (6 + skoXC1) <= -1),
    inference(arithmetic, [], [refute_0_180])).

cnf(refute_0_182, plain,
    (-1 < skoXC1 * (6 + skoXC1) | skoXC1 * (-17 + skoXC1 * 1/2) <= 7/2),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_183, plain,
    (7/2 < sqrt(skoXC1) | skoXC1 * (-17 + skoXC1 * 1/2) <= 7/2),
    inference(resolve, [], [refute_0_181, refute_0_182])).

cnf(refute_0_184, plain, (7/2 < sqrt(skoXC1) | 0 <= skoXC1),
    inference(decision, [], [refute_0_183, refute_0_27])).

cnf(refute_0_185, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     7/2 <= (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1)),
    inference(resolve, [], [refute_0_184, refute_0_176])).

cnf(refute_0_186, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     7/2 * (1 + skoXC1) <= 1/4 + skoXC1 * (3/2 + skoXC1 * 1/4) |
     1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_185, refute_0_171])).

cnf(refute_0_187, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (2 + skoXC1 * -1/4) <= -13/4 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_186])).

cnf(refute_0_188, plain,
    (-1 < skoXC1 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (2 + skoXC1 * -1/4) <= -13/4),
    inference(decision, [], [refute_0_183, refute_0_27])).

cnf(refute_0_189, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (2 + skoXC1 * -1/4) <= -13/4),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (sqrt(skoXC1) < 7/2 | -8 < skoXC1 * (8 + skoXC1) | 7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_189, refute_0_183, refute_0_27])).

cnf(refute_0_191, plain,
    (sqrt(skoXC1) < 7/2 |
     8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_158, refute_0_190])).

cnf(refute_0_192, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(-1 + skoXC1) *
     (-8 + skoXC1 * (-8 + skoXC1 * -1) +
      sqrt(skoXC1) * (16 + skoXC1 * (16 + skoXC1 * 2))) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     7/2 < sqrt(skoXC1) |
     0 <=
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_149, refute_0_191])).

cnf(refute_0_193, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) <
     sqrt(-1 + skoXC1) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_194])).

cnf(refute_0_196, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) < 0 |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) < 0 |
     sqrt(skoXC1) < 7/2 |
     0 <
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_193])).

cnf(refute_0_197, plain,
    (sqrt(skoXC1) < 7/2 |
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_196])).

cnf(refute_0_198, plain,
    ((-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_199, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     0 <= -80/3 + skoXC1 * (-32 + skoXC1 * -22/3) |
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_200, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <= sqrt(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_201, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    introduced(tautology, [assume])).

cnf(refute_0_202, plain,
    (sqrt(skoXC1) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     0 <= -80/3 + skoXC1 * (-32 + skoXC1 * -22/3) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(split, [], [refute_0_201])).

cnf(refute_0_203, plain,
    (sqrt(skoXC1) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (sqrt(skoXC1) < 7/2 | -80/3 < skoXC1 * (32 + skoXC1 * 22/3) |
     7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_205, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(skoXC1) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_203, refute_0_204])).

cnf(refute_0_206, plain,
    (skoXC1 < 0 | sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_200, refute_0_205])).

cnf(refute_0_207, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) | 0 <= skoXC1),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_208, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_207, refute_0_206])).

cnf(refute_0_209, plain,
    (sqrt(skoXC1) < 7/2 |
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     0 <= -80/3 + skoXC1 * (-32 + skoXC1 * -22/3) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_199, refute_0_208])).

cnf(refute_0_210, plain,
    (sqrt(skoXC1) < 7/2 |
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain,
    (sqrt(skoXC1) < 7/2 |
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_210, refute_0_204])).

cnf(refute_0_212, plain,
    (sqrt(skoXC1) < 7/2 |
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_198, refute_0_211])).

cnf(refute_0_213, plain,
    (sqrt(skoXC1) < 7/2 |
     -136/3 <
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_212])).

cnf(refute_0_214, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1 |
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) <=
     -136/3),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_215, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_214, refute_0_213])).

cnf(refute_0_216, plain,
    (sqrt(skoXC1) < 7/2 | -1 < skoXC1 * (6 + skoXC1) | 7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_217, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_215, refute_0_216])).

cnf(refute_0_218, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_217, refute_0_197])).

cnf(refute_0_219, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) <= -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_220, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     0 <= -16 + skoXC1 * (-24 + skoXC1 * -8) |
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_221, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) <= sqrt(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_222, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)), introduced(tautology, [assume])).

cnf(refute_0_223, plain,
    (sqrt(skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     0 <= -16 + skoXC1 * (-24 + skoXC1 * -8) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(split, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (sqrt(skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     skoXC1 * (24 + skoXC1 * 8) <= -16 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (sqrt(skoXC1) < 7/2 | -16 < skoXC1 * (24 + skoXC1 * 8) |
     7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_189, refute_0_183, refute_0_27])).

cnf(refute_0_226, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_224, refute_0_225])).

cnf(refute_0_227, plain,
    (skoXC1 < 0 | sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_221, refute_0_226])).

cnf(refute_0_228, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_207, refute_0_227])).

cnf(refute_0_229, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     0 <= -16 + skoXC1 * (-24 + skoXC1 * -8) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_220, refute_0_228])).

cnf(refute_0_230, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (24 + skoXC1 * 8) <= -16 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_230, refute_0_225])).

cnf(refute_0_232, plain,
    (sqrt(skoXC1) < 7/2 |
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) |
     7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_219, refute_0_231])).

cnf(refute_0_233, plain,
    (sqrt(skoXC1) < 7/2 |
     -8 < skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_232])).

cnf(refute_0_234, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) <= -8 |
     skoXC1 * (6 + skoXC1) <= -1),
    inference(decision, [], [refute_0_189, refute_0_183, refute_0_27])).

cnf(refute_0_235, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_234, refute_0_233])).

cnf(refute_0_236, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_235, refute_0_216])).

cnf(refute_0_237, plain,
    (sqrt(skoXC1) < 7/2 |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_236, refute_0_218])).

cnf(refute_0_238, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_195, refute_0_237])).

cnf(refute_0_239, plain,
    (-1 + skoXC1 < 0 | sqrt(skoXC1) < 7/2 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_148, refute_0_238])).

cnf(refute_0_240, plain,
    (skoXC1 < 1 | sqrt(skoXC1) < 7/2 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_239])).

cnf(refute_0_241, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) | 1 <= skoXC1),
    inference(decision, [], [refute_0_189, refute_0_183, refute_0_27])).

cnf(refute_0_242, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_241, refute_0_240])).

cnf(refute_0_243, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (-4 + skoXC1 * 4) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_147, refute_0_242])).

cnf(refute_0_244, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (sqrt(skoXC1) < 7/2 | 4 < skoXC1 * (4 + skoXC1) | 7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_189, refute_0_183, refute_0_27])).

cnf(refute_0_246, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (sqrt(skoXC1) < 7/2 |
     32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
     sqrt(skoXC1) *
     (-64 + skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_246])).

cnf(refute_0_248, plain,
    (sqrt(skoXC1) < 7/2 |
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (sqrt(skoXC1) < 7/2 |
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_248, refute_0_237])).

cnf(refute_0_250, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (sqrt(skoXC1) < 7/2 |
     64 <
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) |
     7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_253, plain,
    (sqrt(skoXC1) < 7/2 |
     sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_251, refute_0_252])).

cnf(refute_0_254, plain,
    (skoXC1 < 0 | sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_145, refute_0_253])).

cnf(refute_0_255, plain,
    (sqrt(skoXC1) < 7/2 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_207, refute_0_254])).

cnf(refute_0_256, plain,
    (sqrt(skoXC1) < 7/2 |
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_144, refute_0_255])).

cnf(refute_0_257, plain,
    (sqrt(skoXC1) < 7/2 |
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (sqrt(skoXC1) < 7/2 |
     32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_257, refute_0_252])).

cnf(refute_0_259, plain,
    (sqrt(skoXC1) < 7/2 |
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_143, refute_0_258])).

cnf(refute_0_260, plain,
    (sqrt(skoXC1) < 7/2 |
     32 <
     skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) <=
     32 | skoXC1 * (6 + skoXC1) <= -1),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_189])).

cnf(refute_0_262, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_261, refute_0_260])).

cnf(refute_0_263, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_262, refute_0_216])).

cnf(refute_0_264, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 7/2 |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_170, refute_0_263])).

cnf(refute_0_265, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 7/2 |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_184, refute_0_264])).

cnf(refute_0_266, plain,
    (skoXC1 * (4 + skoXC1 * 4) < 7/2 * (1 + skoXC1 * (6 + skoXC1)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_169, refute_0_265])).

cnf(refute_0_267, plain,
    (-7/2 < skoXC1 * (17 + skoXC1 * -1/2) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_266])).

cnf(refute_0_268, plain,
    (-1 < skoXC1 * (6 + skoXC1) | -7/2 < skoXC1 * (17 + skoXC1 * -1/2) |
     7/2 < sqrt(skoXC1)),
    inference(decision, [], [refute_0_183, refute_0_27])).

cnf(refute_0_269, plain,
    (-7/2 < skoXC1 * (17 + skoXC1 * -1/2) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_267, refute_0_268])).

cnf(refute_0_270, plain,
    ((-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) * (1 + skoXC1) <
     -2 +
     skoXC1 * (-34/3 + skoXC1 * (23/6 + skoXC1 * (35/3 + skoXC1 * 11/6))) |
     1 + skoXC1 <= 0 |
     (-2 +
      skoXC1 *
      (-34/3 + skoXC1 * (23/6 + skoXC1 * (35/3 + skoXC1 * 11/6)))) /
     (1 + skoXC1) <= -4 + skoXC1 * (-44/3 + skoXC1 * 59/3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_271, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) |
     -8 + skoXC1 * (8/3 + skoXC1 * 22/3) <= 0 |
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) <=
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_272, plain,
    (skoXC1 < 0 |
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(skoXC1) <=
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_273, plain,
    (-4 + skoXC1 * (4 + skoXC1) +
     sqrt(skoXC1) * (8 + skoXC1 * (-8 + skoXC1 * -2)) <
     skoXC1 * (-4 + skoXC1 * 4) * (-14/3 + sqrt(skoXC1) * 4/3) |
     0 <= -14/3 + sqrt(skoXC1) * 4/3 |
     (-4 + skoXC1 * (4 + skoXC1) +
      sqrt(skoXC1) * (8 + skoXC1 * (-8 + skoXC1 * -2))) /
     (-14/3 + sqrt(skoXC1) * 4/3) <= skoXC1 * (-4 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_274, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) *
     (-4 + skoXC1 * (4 + skoXC1)) | -4 + skoXC1 * (4 + skoXC1) <= 0 |
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) <=
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_275, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) <=
     sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_276, plain,
    (1 + sqrt(skoXC1) * -2 <
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2), introduced(tautology, [assume])).

cnf(refute_0_277, plain,
    (sqrt(-1 + skoXC1) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     0 <= -14/3 + sqrt(skoXC1) * 4/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2), inference(split, [], [refute_0_276])).

cnf(refute_0_278, plain,
    (sqrt(-1 + skoXC1) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 <= sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2), inference(arithmetic, [], [refute_0_277])).

cnf(refute_0_279, plain,
    (sqrt(-1 + skoXC1) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_278, refute_0_263])).

cnf(refute_0_280, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_275, refute_0_279])).

cnf(refute_0_281, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_280])).

cnf(refute_0_282, plain, (7/2 < sqrt(skoXC1) | 1 <= skoXC1),
    inference(decision, [], [refute_0_183, refute_0_27])).

cnf(refute_0_283, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_282, refute_0_281])).

cnf(refute_0_284, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (1 + sqrt(skoXC1) * -2) / (-14/3 + sqrt(skoXC1) * 4/3) *
     (-4 + skoXC1 * (4 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_274, refute_0_283])).

cnf(refute_0_285, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-4 + skoXC1 * (4 + skoXC1) +
      sqrt(skoXC1) * (8 + skoXC1 * (-8 + skoXC1 * -2))) /
     (-14/3 + sqrt(skoXC1) * 4/3) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (4 < skoXC1 * (4 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_269, refute_0_183, refute_0_27])).

cnf(refute_0_287, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-4 + skoXC1 * (4 + skoXC1) +
      sqrt(skoXC1) * (8 + skoXC1 * (-8 + skoXC1 * -2))) /
     (-14/3 + sqrt(skoXC1) * 4/3) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_285, refute_0_286])).

cnf(refute_0_288, plain,
    (-4 + skoXC1 * (4 + skoXC1) +
     sqrt(skoXC1) * (8 + skoXC1 * (-8 + skoXC1 * -2)) <
     skoXC1 * (-4 + skoXC1 * 4) * (-14/3 + sqrt(skoXC1) * 4/3) |
     7/2 < sqrt(skoXC1) | 0 <= -14/3 + sqrt(skoXC1) * 4/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_273, refute_0_287])).

cnf(refute_0_289, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_288])).

cnf(refute_0_290, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_289, refute_0_263])).

cnf(refute_0_291, plain,
    (7/2 < sqrt(skoXC1) |
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -8 + skoXC1 * (8/3 + skoXC1 * 22/3) <= 0),
    inference(split, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (7/2 < sqrt(skoXC1) |
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     skoXC1 * (8/3 + skoXC1 * 22/3) <= 8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_291])).

cnf(refute_0_293, plain,
    (8 < skoXC1 * (8/3 + skoXC1 * 22/3) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_294, plain,
    (7/2 < sqrt(skoXC1) |
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_292, refute_0_293])).

cnf(refute_0_295, plain,
    (skoXC1 < 0 |
     (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_272, refute_0_294])).

cnf(refute_0_296, plain,
    (7/2 < sqrt(skoXC1) | 0 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_297, plain,
    ((-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) /
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_296, refute_0_295])).

cnf(refute_0_298, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (-8 + skoXC1 * (8/3 + skoXC1 * 22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -8 + skoXC1 * (8/3 + skoXC1 * 22/3) <= 0),
    inference(resolve, [], [refute_0_271, refute_0_297])).

cnf(refute_0_299, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     (-2 +
      skoXC1 *
      (-34/3 + skoXC1 * (23/6 + skoXC1 * (35/3 + skoXC1 * 11/6)))) /
     (1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (8/3 + skoXC1 * 22/3) <= 8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (-4 + skoXC1 * (-44/3 + skoXC1 * 59/3) <
     (-2 +
      skoXC1 *
      (-34/3 + skoXC1 * (23/6 + skoXC1 * (35/3 + skoXC1 * 11/6)))) /
     (1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_299, refute_0_293])).

cnf(refute_0_301, plain,
    ((-4 + skoXC1 * (-44/3 + skoXC1 * 59/3)) * (1 + skoXC1) <
     -2 +
     skoXC1 * (-34/3 + skoXC1 * (23/6 + skoXC1 * (35/3 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_270, refute_0_300])).

cnf(refute_0_302, plain,
    (-2 <
     skoXC1 * (22/3 + skoXC1 * (-7/6 + skoXC1 * (-8 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_301])).

cnf(refute_0_303, plain,
    (-2 <
     skoXC1 * (22/3 + skoXC1 * (-7/6 + skoXC1 * (-8 + skoXC1 * 11/6))) |
     -1 < skoXC1 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_304, plain,
    (-2 <
     skoXC1 * (22/3 + skoXC1 * (-7/6 + skoXC1 * (-8 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_302, refute_0_303])).

cnf(refute_0_305, plain,
    (-8 < skoXC1 * (8 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304])).

cnf(refute_0_306, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) <
     sqrt(-1 + skoXC1) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_160, refute_0_305])).

cnf(refute_0_307, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_307])).

cnf(refute_0_309, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) < 0 |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) < 0 |
     0 <
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_306])).

cnf(refute_0_310, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_309])).

cnf(refute_0_311, plain,
    ((136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (320/3 +
      skoXC1 *
      (64 +
       skoXC1 *
       (-264 + skoXC1 * (-944/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 *
     (320/3 +
      skoXC1 *
      (64 +
       skoXC1 *
       (-264 + skoXC1 * (-944/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_312, plain,
    (136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     0 <=
     80/3 +
     skoXC1 *
     (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_313, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) <=
     sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_314, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     skoXC1 * (8 + skoXC1 * (12 + skoXC1 * 4)) +
     sqrt(skoXC1) * (skoXC1 * (-16 + skoXC1 * (-24 + skoXC1 * -8))) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <= 0 |
     (skoXC1 * (8 + skoXC1 * (12 + skoXC1 * 4)) +
      sqrt(skoXC1) * (skoXC1 * (-16 + skoXC1 * (-24 + skoXC1 * -8)))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <=
     1 + skoXC1 * (-1 + skoXC1 * -1/4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_315, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) |
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 <=
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     skoXC1 <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_316, plain,
    (-1 + skoXC1 < 0 |
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) |
     sqrt(-1 + skoXC1) <=
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_317, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     sqrt(-1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <= 0),
    inference(split, [], [refute_0_306])).

cnf(refute_0_318, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     sqrt(-1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_317])).

cnf(refute_0_319, plain,
    (-1 + skoXC1 < 0 |
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_316, refute_0_318])).

cnf(refute_0_320, plain,
    (skoXC1 < 1 |
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (-80/3 < skoXC1 * (32 + skoXC1 * 22/3) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_304, refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_322, plain,
    (sqrt(skoXC1) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_203, refute_0_321])).

cnf(refute_0_323, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_200, refute_0_322])).

cnf(refute_0_324, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_296, refute_0_323])).

cnf(refute_0_325, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) | 7/2 < sqrt(skoXC1) |
     0 <= -80/3 + skoXC1 * (-32 + skoXC1 * -22/3) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_199, refute_0_324])).

cnf(refute_0_326, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_325])).

cnf(refute_0_327, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_326, refute_0_321])).

cnf(refute_0_328, plain,
    ((-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_198, refute_0_327])).

cnf(refute_0_329, plain,
    (-136/3 <
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_328])).

cnf(refute_0_330, plain,
    (-1 < skoXC1 * (6 + skoXC1) |
     -136/3 <
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_304, refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_331, plain,
    (-136/3 <
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_329, refute_0_330])).

cnf(refute_0_332, plain,
    (7/2 < sqrt(skoXC1) | 1 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304,
               refute_0_331])).

cnf(refute_0_333, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_332, refute_0_320])).

cnf(refute_0_334, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_315, refute_0_333])).

cnf(refute_0_335, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 * (8 + skoXC1 * (12 + skoXC1 * 4)) +
      sqrt(skoXC1) * (skoXC1 * (-16 + skoXC1 * (-24 + skoXC1 * -8)))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_334])).

cnf(refute_0_336, plain,
    (0 < skoXC1 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304,
               refute_0_331])).

cnf(refute_0_337, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 * (8 + skoXC1 * (12 + skoXC1 * 4)) +
      sqrt(skoXC1) * (skoXC1 * (-16 + skoXC1 * (-24 + skoXC1 * -8)))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_335, refute_0_336])).

cnf(refute_0_338, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) <
     skoXC1 * (8 + skoXC1 * (12 + skoXC1 * 4)) +
     sqrt(skoXC1) * (skoXC1 * (-16 + skoXC1 * (-24 + skoXC1 * -8))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <= 0),
    inference(resolve, [], [refute_0_314, refute_0_337])).

cnf(refute_0_339, plain,
    (136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     sqrt(skoXC1) *
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_338])).

cnf(refute_0_340, plain,
    (sqrt(skoXC1) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     80/3 +
     skoXC1 *
     (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(split, [], [refute_0_339])).

cnf(refute_0_341, plain,
    (sqrt(skoXC1) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     skoXC1 *
     (32/3 + skoXC1 * (166/3 + skoXC1 * (70/3 + skoXC1 * 11/6))) <= 80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_340])).

cnf(refute_0_342, plain,
    (80/3 <
     skoXC1 * (32/3 + skoXC1 * (166/3 + skoXC1 * (70/3 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_331, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_343, plain,
    (sqrt(skoXC1) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_341, refute_0_342])).

cnf(refute_0_344, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_313, refute_0_343])).

cnf(refute_0_345, plain,
    (7/2 < sqrt(skoXC1) | 0 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_331, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_346, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_345, refute_0_344])).

cnf(refute_0_347, plain,
    (136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (80/3 +
      skoXC1 *
      (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     80/3 +
     skoXC1 *
     (-32/3 + skoXC1 * (-166/3 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_312, refute_0_346])).

cnf(refute_0_348, plain,
    (136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 *
     (320/3 +
      skoXC1 *
      (64 +
       skoXC1 *
       (-264 + skoXC1 * (-944/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 *
     (32/3 + skoXC1 * (166/3 + skoXC1 * (70/3 + skoXC1 * 11/6))) <= 80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_347])).

cnf(refute_0_349, plain,
    (136/3 +
     skoXC1 *
     (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 *
     (320/3 +
      skoXC1 *
      (64 +
       skoXC1 *
       (-264 + skoXC1 * (-944/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_348, refute_0_342])).

cnf(refute_0_350, plain,
    ((136/3 +
      skoXC1 *
      (32/3 + skoXC1 * (-203/3 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (320/3 +
      skoXC1 *
      (64 +
       skoXC1 *
       (-264 + skoXC1 * (-944/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_311, refute_0_349])).

cnf(refute_0_351, plain,
    (136/3 <
     skoXC1 *
     (-176 +
      skoXC1 *
      (67/3 +
       skoXC1 *
       (171 + skoXC1 * (-49/12 + skoXC1 * (-63/2 + skoXC1 * -29/12))))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_350])).

cnf(refute_0_352, plain,
    (7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-176 +
      skoXC1 *
      (67/3 +
       skoXC1 *
       (171 + skoXC1 * (-49/12 + skoXC1 * (-63/2 + skoXC1 * -29/12))))) <=
     136/3 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_331, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_353, plain,
    (7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_352, refute_0_351])).

cnf(refute_0_354, plain,
    (-1 < skoXC1 * (6 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(decision, [],
              [refute_0_331, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_355, plain,
    (7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_353, refute_0_354])).

cnf(refute_0_356, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_355, refute_0_310])).

cnf(refute_0_357, plain,
    (-16 < skoXC1 * (24 + skoXC1 * 8) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304])).

cnf(refute_0_358, plain,
    (sqrt(skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_224, refute_0_357])).

cnf(refute_0_359, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_221, refute_0_358])).

cnf(refute_0_360, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_296, refute_0_359])).

cnf(refute_0_361, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) | 7/2 < sqrt(skoXC1) |
     0 <= -16 + skoXC1 * (-24 + skoXC1 * -8) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_220, refute_0_360])).

cnf(refute_0_362, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (24 + skoXC1 * 8) <= -16 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_361])).

cnf(refute_0_363, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_362, refute_0_357])).

cnf(refute_0_364, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_219, refute_0_363])).

cnf(refute_0_365, plain,
    (-8 < skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_364])).

cnf(refute_0_366, plain,
    (7/2 < sqrt(skoXC1) |
     skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) <= -8 |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304])).

cnf(refute_0_367, plain,
    (7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_366, refute_0_365])).

cnf(refute_0_368, plain,
    (-1 < skoXC1 * (6 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304])).

cnf(refute_0_369, plain,
    (7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_367, refute_0_368])).

cnf(refute_0_370, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_369, refute_0_356])).

cnf(refute_0_371, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_308, refute_0_370])).

cnf(refute_0_372, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_148, refute_0_371])).

cnf(refute_0_373, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_372])).

cnf(refute_0_374, plain,
    ((136/3 + skoXC1 * (64 + skoXC1 * 59/3)) * (1 + skoXC1) <
     20/3 +
     skoXC1 * (48 + skoXC1 * (113/2 + skoXC1 * (19 + skoXC1 * 11/6))) |
     1 + skoXC1 <= 0 |
     (20/3 +
      skoXC1 * (48 + skoXC1 * (113/2 + skoXC1 * (19 + skoXC1 * 11/6)))) /
     (1 + skoXC1) <= 136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_375, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     80/3 + skoXC1 * (32 + skoXC1 * 22/3) <= 0 |
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) <=
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_376, plain,
    (skoXC1 < 0 |
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(skoXC1) <=
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_377, plain,
    (7/2 < sqrt(skoXC1) |
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     80/3 + skoXC1 * (32 + skoXC1 * 22/3) <= 0),
    inference(split, [], [refute_0_370])).

cnf(refute_0_378, plain,
    (7/2 < sqrt(skoXC1) |
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_377])).

cnf(refute_0_379, plain,
    (7/2 < sqrt(skoXC1) |
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_378, refute_0_321])).

cnf(refute_0_380, plain,
    (skoXC1 < 0 |
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_376, refute_0_379])).

cnf(refute_0_381, plain,
    ((136/3 + skoXC1 * (64 + skoXC1 * 59/3)) /
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_296, refute_0_380])).

cnf(refute_0_382, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     80/3 + skoXC1 * (32 + skoXC1 * 22/3) <= 0),
    inference(resolve, [], [refute_0_375, refute_0_381])).

cnf(refute_0_383, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     (20/3 +
      skoXC1 * (48 + skoXC1 * (113/2 + skoXC1 * (19 + skoXC1 * 11/6)))) /
     (1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_382])).

cnf(refute_0_384, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     (20/3 +
      skoXC1 * (48 + skoXC1 * (113/2 + skoXC1 * (19 + skoXC1 * 11/6)))) /
     (1 + skoXC1) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_383, refute_0_321])).

cnf(refute_0_385, plain,
    ((136/3 + skoXC1 * (64 + skoXC1 * 59/3)) * (1 + skoXC1) <
     20/3 +
     skoXC1 * (48 + skoXC1 * (113/2 + skoXC1 * (19 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_374, refute_0_384])).

cnf(refute_0_386, plain,
    (116/3 <
     skoXC1 *
     (-184/3 + skoXC1 * (-163/6 + skoXC1 * (-2/3 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_385])).

cnf(refute_0_387, plain,
    (116/3 <
     skoXC1 *
     (-184/3 + skoXC1 * (-163/6 + skoXC1 * (-2/3 + skoXC1 * 11/6))) |
     -1 < skoXC1 | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_304, refute_0_27, refute_0_183, refute_0_269])).

cnf(refute_0_388, plain,
    (116/3 <
     skoXC1 *
     (-184/3 + skoXC1 * (-163/6 + skoXC1 * (-2/3 + skoXC1 * 11/6))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_386, refute_0_387])).

cnf(refute_0_389, plain,
    (7/2 < sqrt(skoXC1) | 1 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_269, refute_0_183, refute_0_27, refute_0_304,
               refute_0_388])).

cnf(refute_0_390, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_389, refute_0_373])).

cnf(refute_0_391, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_147, refute_0_390])).

cnf(refute_0_392, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_392, refute_0_286])).

cnf(refute_0_394, plain,
    (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
     sqrt(skoXC1) *
     (-64 + skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_393])).

cnf(refute_0_395, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_395, refute_0_370])).

cnf(refute_0_397, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_396])).

cnf(refute_0_398, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (64 <
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_388, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_400, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_398, refute_0_399])).

cnf(refute_0_401, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_145, refute_0_400])).

cnf(refute_0_402, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_296, refute_0_401])).

cnf(refute_0_403, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     7/2 < sqrt(skoXC1) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_144, refute_0_402])).

cnf(refute_0_404, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_403])).

cnf(refute_0_405, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_404, refute_0_399])).

cnf(refute_0_406, plain,
    ((32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) |
     7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_143, refute_0_405])).

cnf(refute_0_407, plain,
    (32 <
     skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) |
     7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_406])).

cnf(refute_0_408, plain,
    (7/2 < sqrt(skoXC1) |
     skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) <=
     32 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_388, refute_0_304, refute_0_27, refute_0_183,
               refute_0_269])).

cnf(refute_0_409, plain,
    (7/2 < sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_408, refute_0_407])).

cnf(refute_0_410, plain,
    (7/2 < sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_409, refute_0_368])).

cnf(refute_0_411, plain,
    (skoXC1 < 0 |
     7/2 < (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_168, refute_0_410])).

cnf(refute_0_412, plain,
    (7/2 < (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_110, refute_0_411])).

cnf(refute_0_413, plain,
    (7/2 * (1 + skoXC1) < 1/4 + skoXC1 * (3/2 + skoXC1 * 1/4) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_163, refute_0_412])).

cnf(refute_0_414, plain,
    (13/4 < skoXC1 * (-2 + skoXC1 * 1/4) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_413])).

cnf(refute_0_415, plain,
    (13/4 < skoXC1 * (-2 + skoXC1 * 1/4) | -1 < skoXC1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_416, plain,
    (13/4 < skoXC1 * (-2 + skoXC1 * 1/4) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_414, refute_0_415])).

cnf(refute_0_417, plain,
    (-8 < skoXC1 * (8 + skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_418, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(resolve, [], [refute_0_162, refute_0_417])).

cnf(refute_0_419, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) +
     sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8)) < 0 |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) < 0 |
     0 <
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_160])).

cnf(refute_0_420, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     skoXC1 * (8 + skoXC1) <= -8 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_419])).

cnf(refute_0_421, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     -8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_420, refute_0_417])).

cnf(refute_0_422, plain,
    (-80/3 < skoXC1 * (32 + skoXC1 * 22/3) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_423, plain,
    (sqrt(skoXC1) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_203, refute_0_422])).

cnf(refute_0_424, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_200, refute_0_423])).

cnf(refute_0_425, plain,
    (0 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_426, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) /
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_425, refute_0_424])).

cnf(refute_0_427, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     0 <= -80/3 + skoXC1 * (-32 + skoXC1 * -22/3) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_199, refute_0_426])).

cnf(refute_0_428, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | skoXC1 * (32 + skoXC1 * 22/3) <= -80/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_427])).

cnf(refute_0_429, plain,
    (-136/3 + skoXC1 * (-64 + skoXC1 * -59/3) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_428, refute_0_422])).

cnf(refute_0_430, plain,
    ((-136/3 + skoXC1 * (-64 + skoXC1 * -59/3)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (-320/3 + skoXC1 * (-704/3 + skoXC1 * (-472/3 + skoXC1 * -88/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_198, refute_0_429])).

cnf(refute_0_431, plain,
    (-136/3 <
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(arithmetic, [], [refute_0_430])).

cnf(refute_0_432, plain,
    (skoXC1 * (6 + skoXC1) <= -1 |
     skoXC1 *
     (688/3 + skoXC1 * (643/3 + skoXC1 * (74/3 + skoXC1 * -29/3))) <=
     -136/3 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_433, plain,
    (skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_432, refute_0_431])).

cnf(refute_0_434, plain,
    (-1 < skoXC1 * (6 + skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_435, plain,
    (sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) <=
     -136/3 + skoXC1 * (-64 + skoXC1 * -59/3)),
    inference(resolve, [], [refute_0_433, refute_0_434])).

cnf(refute_0_436, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_435, refute_0_421])).

cnf(refute_0_437, plain,
    (-16 < skoXC1 * (24 + skoXC1 * 8) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_416, refute_0_27])).

cnf(refute_0_438, plain,
    (sqrt(skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_224, refute_0_437])).

cnf(refute_0_439, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_221, refute_0_438])).

cnf(refute_0_440, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4)) /
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_425, refute_0_439])).

cnf(refute_0_441, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-16 + skoXC1 * (-24 + skoXC1 * -8)) |
     0 <= -16 + skoXC1 * (-24 + skoXC1 * -8) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_220, refute_0_440])).

cnf(refute_0_442, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) | skoXC1 * (24 + skoXC1 * 8) <= -16 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_441])).

cnf(refute_0_443, plain,
    (-8 + skoXC1 * (-12 + skoXC1 * -4) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) /
     (1 + skoXC1 * (6 + skoXC1)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_442, refute_0_437])).

cnf(refute_0_444, plain,
    ((-8 + skoXC1 * (-12 + skoXC1 * -4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-64 + skoXC1 * (-160 + skoXC1 * (-128 + skoXC1 * -32))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_219, refute_0_443])).

cnf(refute_0_445, plain,
    (-8 < skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(arithmetic, [], [refute_0_444])).

cnf(refute_0_446, plain,
    (skoXC1 * (-4 + skoXC1 * (-76 + skoXC1 * (-92 + skoXC1 * -28))) <= -8 |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_416, refute_0_27])).

cnf(refute_0_447, plain,
    (skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_446, refute_0_445])).

cnf(refute_0_448, plain,
    (sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-16 + skoXC1 * (-24 + skoXC1 * -8)) <=
     -8 + skoXC1 * (-12 + skoXC1 * -4)),
    inference(resolve, [], [refute_0_447, refute_0_434])).

cnf(refute_0_449, plain,
    (136/3 + skoXC1 * (64 + skoXC1 * 59/3) <
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_448, refute_0_436])).

cnf(refute_0_450, plain,
    (sqrt(-1 + skoXC1) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_418, refute_0_449])).

cnf(refute_0_451, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_148, refute_0_450])).

cnf(refute_0_452, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_451])).

cnf(refute_0_453, plain,
    (1 <= skoXC1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_416, refute_0_27])).

cnf(refute_0_454, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_453, refute_0_452])).

cnf(refute_0_455, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-8 + skoXC1 * (-12 + skoXC1 * -4) +
      sqrt(skoXC1) * (16 + skoXC1 * (24 + skoXC1 * 8))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_147, refute_0_454])).

cnf(refute_0_456, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_455])).

cnf(refute_0_457, plain,
    (4 < skoXC1 * (4 + skoXC1) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_416, refute_0_27])).

cnf(refute_0_458, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
      sqrt(skoXC1) *
      (-64 +
       skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8))))) /
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_456, refute_0_457])).

cnf(refute_0_459, plain,
    (32 + skoXC1 * (16 + skoXC1 * (-40 + skoXC1 * (-28 + skoXC1 * -4))) +
     sqrt(skoXC1) *
     (-64 + skoXC1 * (-32 + skoXC1 * (80 + skoXC1 * (56 + skoXC1 * 8)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
      sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3))) |
     0 <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3) +
     sqrt(skoXC1) * (-80/3 + skoXC1 * (-32 + skoXC1 * -22/3)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_458])).

cnf(refute_0_460, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (80/3 + skoXC1 * (32 + skoXC1 * 22/3)) <=
     136/3 + skoXC1 * (64 + skoXC1 * 59/3)),
    inference(arithmetic, [], [refute_0_459])).

cnf(refute_0_461, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     sqrt(skoXC1) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_460, refute_0_449])).

cnf(refute_0_462, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_461])).

cnf(refute_0_463, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_462])).

cnf(refute_0_464, plain,
    (64 <
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_465, plain,
    (sqrt(skoXC1) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_463, refute_0_464])).

cnf(refute_0_466, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_145, refute_0_465])).

cnf(refute_0_467, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) /
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_425, refute_0_466])).

cnf(refute_0_468, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (64 +
      skoXC1 *
      (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3)))) |
     0 <=
     64 +
     skoXC1 *
     (416/3 + skoXC1 * (-176/3 + skoXC1 * (-464/3 + skoXC1 * -112/3))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_144, refute_0_467])).

cnf(refute_0_469, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) |
     skoXC1 *
     (-416/3 + skoXC1 * (176/3 + skoXC1 * (464/3 + skoXC1 * 112/3))) <=
     64 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_468])).

cnf(refute_0_470, plain,
    (32 +
     skoXC1 *
     (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3))) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_469, refute_0_464])).

cnf(refute_0_471, plain,
    ((32 +
      skoXC1 *
      (592/3 + skoXC1 * (104/3 + skoXC1 * (-616/3 + skoXC1 * -248/3)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (256 +
      skoXC1 *
      (2432/3 +
       skoXC1 *
       (320 + skoXC1 * (-2560/3 + skoXC1 * (-768 + skoXC1 * -448/3))))) |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_143, refute_0_470])).

cnf(refute_0_472, plain,
    (32 <
     skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_471])).

cnf(refute_0_473, plain,
    (skoXC1 *
     (-400/3 +
      skoXC1 *
      (-440 +
       skoXC1 *
       (120 + skoXC1 * (1280/3 + skoXC1 * (-200/3 + skoXC1 * -200/3))))) <=
     32 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_416])).

cnf(refute_0_474, plain,
    (skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_473, refute_0_472])).

cnf(refute_0_475, plain,
    (sqrt(-1 + skoXC1) * (-14/3 + sqrt(skoXC1) * 4/3) <=
     1 + sqrt(skoXC1) * -2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_474, refute_0_434])).

cnf(refute_0_476, plain,
    (-1 + sqrt(skoXC1) * 2 <
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_475, refute_0_142])).

cnf(refute_0_477, plain,
    (sqrt(-1 + skoXC1) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     0 <= 14/3 + sqrt(skoXC1) * -4/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_476])).

cnf(refute_0_478, plain,
    (sqrt(-1 + skoXC1) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) <= 7/2), inference(arithmetic, [], [refute_0_477])).

cnf(refute_0_479, plain,
    (-1 + sqrt(skoXC1) * 2 < 0 | 14/3 + sqrt(skoXC1) * -4/3 < 0 |
     0 < 14/3 + sqrt(skoXC1) * -4/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_476])).

cnf(refute_0_480, plain,
    (sqrt(skoXC1) < 1/2 | sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_479])).

cnf(refute_0_481, plain,
    (1/3 * (1 + sqrt(skoXC1) * -2) < -1 + sqrt(skoXC1) * -2 |
     1 + sqrt(skoXC1) * -2 <= 0 |
     (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) <= 1/3),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_482, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_483, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_484, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_483, refute_0_482])).

cnf(refute_0_485, plain,
    (0 < -2/3 + exp(1 / sqrt(skoXC1)) * 2 |
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 <= exp(1 / sqrt(-1 + skoXC1))),
    inference(subst, [], [refute_0_484])).

cnf(refute_0_486, plain,
    (0 <
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1 |
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1 <=
     exp(1 / sqrt(1 + skoXC1))), inference(subst, [], [refute_0_484])).

cnf(refute_0_487, plain,
    (0 <
     -2/3 + exp(1 / sqrt(skoXC1)) * 2 + exp(1 / sqrt(-1 + skoXC1)) * -1),
    inference(resolve, [], [refute_0_486, refute_0_19])).

cnf(refute_0_488, plain,
    (exp(1 / sqrt(-1 + skoXC1)) < -2/3 + exp(1 / sqrt(skoXC1)) * 2),
    inference(arithmetic, [], [refute_0_487])).

cnf(refute_0_489, plain, (0 < -2/3 + exp(1 / sqrt(skoXC1)) * 2),
    inference(resolve, [], [refute_0_485, refute_0_488])).

cnf(refute_0_490, plain, (1/3 < exp(1 / sqrt(skoXC1))),
    inference(arithmetic, [], [refute_0_489])).

cnf(refute_0_491, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     1/3 < (-2 + 1 / sqrt(skoXC1) * -1) / (-2 + 1 / sqrt(skoXC1)) |
     exp(1 / sqrt(skoXC1)) <= 1/3), inference(subst, [], [refute_0_94])).

cnf(refute_0_492, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     1/3 < (-2 + 1 / sqrt(skoXC1) * -1) / (-2 + 1 / sqrt(skoXC1))),
    inference(resolve, [], [refute_0_491, refute_0_490])).

cnf(refute_0_493, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     1/3 < (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_492])).

cnf(refute_0_494, plain,
    (1 / sqrt(skoXC1) < 0 | 3/2 < 1 / sqrt(skoXC1) |
     1/3 < (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2)),
    inference(resolve, [], [refute_0_493, refute_0_98])).

cnf(refute_0_495, plain,
    (3/2 < 1 / sqrt(skoXC1) |
     1/3 < (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2)),
    inference(resolve, [], [refute_0_102, refute_0_494])).

cnf(refute_0_496, plain,
    (1/3 < (-1 + sqrt(skoXC1) * -2) / (1 + sqrt(skoXC1) * -2)),
    inference(resolve, [], [refute_0_117, refute_0_495])).

cnf(refute_0_497, plain,
    (1/3 * (1 + sqrt(skoXC1) * -2) < -1 + sqrt(skoXC1) * -2 |
     1 + sqrt(skoXC1) * -2 <= 0),
    inference(resolve, [], [refute_0_481, refute_0_496])).

cnf(refute_0_498, plain, (sqrt(skoXC1) < -1 | 1/2 <= sqrt(skoXC1)),
    inference(arithmetic, [], [refute_0_497])).

cnf(refute_0_499, plain,
    (skoXC1 * (4 + skoXC1 * 4) < -1 * (1 + skoXC1 * (6 + skoXC1)) |
     -1 <= skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_500, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < -1 |
     skoXC1 < 0 | -1 <= sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_501, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < -1 |
     -1 <= sqrt(skoXC1)),
    inference(resolve, [], [refute_0_110, refute_0_500])).

cnf(refute_0_502, plain,
    (skoXC1 * (4 + skoXC1 * 4) < -1 * (1 + skoXC1 * (6 + skoXC1)) |
     -1 <= sqrt(skoXC1) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_499, refute_0_501])).

cnf(refute_0_503, plain,
    (1 < skoXC1 * (-10 + skoXC1 * -5) | -1 <= sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1),
    inference(arithmetic, [], [refute_0_502])).

cnf(refute_0_504, plain,
    (skoXC1 * (-10 + skoXC1 * -5) <= 1 | skoXC1 * (6 + skoXC1) <= -1),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_505, plain,
    (-1 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1),
    inference(resolve, [], [refute_0_504, refute_0_503])).

cnf(refute_0_506, plain, (-1 <= sqrt(skoXC1)),
    inference(resolve, [], [refute_0_505, refute_0_116])).

cnf(refute_0_507, plain, (1/2 <= sqrt(skoXC1)),
    inference(resolve, [], [refute_0_506, refute_0_498])).

cnf(refute_0_508, plain,
    (sqrt(skoXC1) < 7/2 | 7/2 < sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_507, refute_0_480])).

cnf(refute_0_509, plain,
    ((8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (64 +
      skoXC1 *
      (512/3 +
       skoXC1 *
       (8/3 + skoXC1 * (-592/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 *
     (64 +
      skoXC1 *
      (512/3 +
       skoXC1 *
       (8/3 + skoXC1 * (-592/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_510, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     0 <=
     16 +
     skoXC1 * (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_511, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) <=
     sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_512, plain,
    (8 + skoXC1 * (skoXC1 * (-9 + skoXC1 * (-3 + skoXC1 * -1/4))) +
     sqrt(skoXC1) *
     (-16 + skoXC1 * (skoXC1 * (18 + skoXC1 * (6 + skoXC1 * 1/2)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     0 <=
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) |
     (8 + skoXC1 * (skoXC1 * (-9 + skoXC1 * (-3 + skoXC1 * -1/4))) +
      sqrt(skoXC1) *
      (-16 + skoXC1 * (skoXC1 * (18 + skoXC1 * (6 + skoXC1 * 1/2))))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <=
     skoXC1 * (-4 + skoXC1 * 4)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_513, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) | -4 + skoXC1 * (4 + skoXC1) <= 0 |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <=
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_514, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <=
     sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_515, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4)) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     sqrt(-1 + skoXC1) * (2 + skoXC1 + sqrt(skoXC1) * (-4 + skoXC1 * -2)) |
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) <= 0 |
     sqrt(-1 + skoXC1) * (2 + skoXC1 + sqrt(skoXC1) * (-4 + skoXC1 * -2)) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_516, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) * (2 + skoXC1) <
     2 + skoXC1 * (2 + skoXC1 * 1/4) | 2 + skoXC1 <= 0 |
     (2 + skoXC1 * (2 + skoXC1 * 1/4)) / (2 + skoXC1) <=
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_517, plain,
    (1 + skoXC1 < 0 |
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     (1/4 + (1 + skoXC1) * (3/2 + (1 + skoXC1) * 1/4)) /
     (1 + (1 + skoXC1)) |
     sqrt(1 + skoXC1) <=
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_518, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     sqrt(1 + skoXC1) | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) <= 0),
    inference(split, [], [refute_0_135])).

cnf(refute_0_519, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     sqrt(1 + skoXC1) |
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) <=
     -1 + sqrt(skoXC1) * 2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_518])).

cnf(refute_0_520, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     sqrt(1 + skoXC1) | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_519, refute_0_476])).

cnf(refute_0_521, plain,
    (1 + skoXC1 < 0 |
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     (1/4 + (1 + skoXC1) * (3/2 + (1 + skoXC1) * 1/4)) /
     (1 + (1 + skoXC1)) | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_517, refute_0_520])).

cnf(refute_0_522, plain,
    (skoXC1 < -1 |
     sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     (2 + skoXC1 * (2 + skoXC1 * 1/4)) / (2 + skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_521])).

cnf(refute_0_523, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) < 7/2 |
     7/2 <= sqrt(skoXC1)),
    inference(resolve, [], [refute_0_110, refute_0_170])).

cnf(refute_0_524, plain,
    (skoXC1 * (4 + skoXC1 * 4) < 7/2 * (1 + skoXC1 * (6 + skoXC1)) |
     7/2 <= sqrt(skoXC1) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_169, refute_0_523])).

cnf(refute_0_525, plain,
    (-7/2 < skoXC1 * (17 + skoXC1 * -1/2) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1) <= -1),
    inference(arithmetic, [], [refute_0_524])).

cnf(refute_0_526, plain,
    (-1 < skoXC1 * (6 + skoXC1) | -7/2 < skoXC1 * (17 + skoXC1 * -1/2)),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_527, plain,
    (-7/2 < skoXC1 * (17 + skoXC1 * -1/2) | 7/2 <= sqrt(skoXC1)),
    inference(resolve, [], [refute_0_525, refute_0_526])).

cnf(refute_0_528, plain,
    ((14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (16/3 + skoXC1 * (-8 + skoXC1 * (-44/3 + skoXC1 * -4/3))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 * (16/3 + skoXC1 * (-8 + skoXC1 * (-44/3 + skoXC1 * -4/3))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_529, plain,
    (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) |
     0 <= 4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3) |
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_530, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) |
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) <= sqrt(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_531, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) * (14/3 + sqrt(skoXC1) * -4/3) <
     skoXC1 + sqrt(skoXC1) * (skoXC1 * -2) |
     14/3 + sqrt(skoXC1) * -4/3 <= 0 |
     (skoXC1 + sqrt(skoXC1) * (skoXC1 * -2)) /
     (14/3 + sqrt(skoXC1) * -4/3) <= 1 + skoXC1 * (-1 + skoXC1 * -1/4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_532, plain,
    ((-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) |
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 <=
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) | skoXC1 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_533, plain,
    (-1 + skoXC1 < 0 |
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) |
     sqrt(-1 + skoXC1) <=
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3)),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_534, plain,
    ((-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     sqrt(-1 + skoXC1) | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     14/3 + sqrt(skoXC1) * -4/3 <= 0),
    inference(split, [], [refute_0_476])).

cnf(refute_0_535, plain,
    ((-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     sqrt(-1 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_534])).

cnf(refute_0_536, plain,
    (-1 + skoXC1 < 0 |
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_533, refute_0_535])).

cnf(refute_0_537, plain,
    (skoXC1 < 1 |
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_536])).

cnf(refute_0_538, plain, (1 <= skoXC1 | 7/2 <= sqrt(skoXC1)),
    inference(decision, [], [refute_0_27, refute_0_527])).

cnf(refute_0_539, plain,
    ((-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_538, refute_0_537])).

cnf(refute_0_540, plain,
    ((-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_532, refute_0_539])).

cnf(refute_0_541, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 + sqrt(skoXC1) * (skoXC1 * -2)) /
     (14/3 + sqrt(skoXC1) * -4/3) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_540])).

cnf(refute_0_542, plain, (0 < skoXC1 | 7/2 <= sqrt(skoXC1)),
    inference(decision, [], [refute_0_27, refute_0_527])).

cnf(refute_0_543, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 + sqrt(skoXC1) * (skoXC1 * -2)) /
     (14/3 + sqrt(skoXC1) * -4/3) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_541, refute_0_542])).

cnf(refute_0_544, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) * (14/3 + sqrt(skoXC1) * -4/3) <
     skoXC1 + sqrt(skoXC1) * (skoXC1 * -2) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     14/3 + sqrt(skoXC1) * -4/3 <= 0),
    inference(resolve, [], [refute_0_531, refute_0_543])).

cnf(refute_0_545, plain,
    (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6) <
     sqrt(skoXC1) * (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_544])).

cnf(refute_0_546, plain,
    (sqrt(skoXC1) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) |
     0 <= 4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_545])).

cnf(refute_0_547, plain,
    (sqrt(skoXC1) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (10/3 + skoXC1 * 1/3) <= 4/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_546])).

cnf(refute_0_548, plain,
    (4/3 < skoXC1 * (10/3 + skoXC1 * 1/3) | 7/2 <= sqrt(skoXC1)),
    inference(decision, [], [refute_0_527, refute_0_27])).

cnf(refute_0_549, plain,
    (sqrt(skoXC1) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_547, refute_0_548])).

cnf(refute_0_550, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_530, refute_0_549])).

cnf(refute_0_551, plain, (0 <= skoXC1 | 7/2 <= sqrt(skoXC1)),
    inference(decision, [], [refute_0_527, refute_0_27])).

cnf(refute_0_552, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) /
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_551, refute_0_550])).

cnf(refute_0_553, plain,
    (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3)) |
     0 <= 4/3 + skoXC1 * (-10/3 + skoXC1 * -1/3) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_529, refute_0_552])).

cnf(refute_0_554, plain,
    (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6) <
     skoXC1 * (16/3 + skoXC1 * (-8 + skoXC1 * (-44/3 + skoXC1 * -4/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (10/3 + skoXC1 * 1/3) <= 4/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_553])).

cnf(refute_0_555, plain,
    (14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6) <
     skoXC1 * (16/3 + skoXC1 * (-8 + skoXC1 * (-44/3 + skoXC1 * -4/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_554, refute_0_548])).

cnf(refute_0_556, plain,
    ((14/3 + skoXC1 * (-17/3 + skoXC1 * -7/6)) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (16/3 + skoXC1 * (-8 + skoXC1 * (-44/3 + skoXC1 * -4/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_528, refute_0_555])).

cnf(refute_0_557, plain,
    (14/3 <
     skoXC1 * (-17 + skoXC1 * (45/2 + skoXC1 * (-2 + skoXC1 * -1/6))) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_556])).

cnf(refute_0_558, plain,
    (14/3 <
     skoXC1 * (-17 + skoXC1 * (45/2 + skoXC1 * (-2 + skoXC1 * -1/6))) |
     -1 < skoXC1 * (6 + skoXC1) | 7/2 <= sqrt(skoXC1)),
    inference(decision, [], [refute_0_527, refute_0_27])).

cnf(refute_0_559, plain,
    (14/3 <
     skoXC1 * (-17 + skoXC1 * (45/2 + skoXC1 * (-2 + skoXC1 * -1/6))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_557, refute_0_558])).

cnf(refute_0_560, plain,
    (-1 <= skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_561, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     (2 + skoXC1 * (2 + skoXC1 * 1/4)) / (2 + skoXC1) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_560, refute_0_522])).

cnf(refute_0_562, plain,
    (sqrt(-1 + skoXC1) * (-1 + sqrt(skoXC1) * 2) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) * (2 + skoXC1) <
     2 + skoXC1 * (2 + skoXC1 * 1/4) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 2 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_516, refute_0_561])).

cnf(refute_0_563, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     sqrt(-1 + skoXC1) * (2 + skoXC1 + sqrt(skoXC1) * (-4 + skoXC1 * -2)) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_562])).

cnf(refute_0_564, plain,
    (-2 < skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_565, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     sqrt(-1 + skoXC1) * (2 + skoXC1 + sqrt(skoXC1) * (-4 + skoXC1 * -2)) /
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_563, refute_0_564])).

cnf(refute_0_566, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4)) *
     (1 + sqrt(skoXC1) * -2 +
      sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3)) <
     sqrt(-1 + skoXC1) * (2 + skoXC1 + sqrt(skoXC1) * (-4 + skoXC1 * -2)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + sqrt(skoXC1) * -2 +
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) <= 0),
    inference(resolve, [], [refute_0_515, refute_0_565])).

cnf(refute_0_567, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
     sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2)) <
     sqrt(-1 + skoXC1) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(-1 + skoXC1) * (14/3 + sqrt(skoXC1) * -4/3) <=
     -1 + sqrt(skoXC1) * 2 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_566])).

cnf(refute_0_568, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
     sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2)) <
     sqrt(-1 + skoXC1) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_567, refute_0_476])).

cnf(refute_0_569, plain,
    (sqrt(-1 + skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     0 <=
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_568])).

cnf(refute_0_570, plain,
    (sqrt(-1 + skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) <=
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)),
    inference(arithmetic, [], [refute_0_569])).

cnf(refute_0_571, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
     sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2)) < 0 |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) < 0 |
     0 <
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_568])).

cnf(refute_0_572, plain,
    (-34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6) <
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) |
     -2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     sqrt(skoXC1) * (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_571])).

cnf(refute_0_573, plain,
    ((34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (80/3 +
      skoXC1 *
      (8/3 +
       skoXC1 *
       (-64 + skoXC1 * (-48 + skoXC1 * (-25/3 + skoXC1 * -1/3))))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 *
     (80/3 +
      skoXC1 *
      (8/3 +
       skoXC1 *
       (-64 + skoXC1 * (-48 + skoXC1 * (-25/3 + skoXC1 * -1/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) <=
     34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_574, plain,
    (34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     0 <=
     20/3 +
     skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12))) |
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_575, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) <=
     sqrt(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_576, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     skoXC1 * (2 + skoXC1 * (2 + skoXC1 * 1/4)) +
     sqrt(skoXC1) * (skoXC1 * (-4 + skoXC1 * (-4 + skoXC1 * -1/2))) |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <= 0 |
     (skoXC1 * (2 + skoXC1 * (2 + skoXC1 * 1/4)) +
      sqrt(skoXC1) * (skoXC1 * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <=
     1 + skoXC1 * (-1 + skoXC1 * -1/4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_577, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) |
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 <=
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     skoXC1 <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_578, plain,
    (-1 + skoXC1 < 0 |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) |
     sqrt(-1 + skoXC1) <=
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_579, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     sqrt(-1 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <= 0),
    inference(split, [], [refute_0_568])).

cnf(refute_0_580, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     sqrt(-1 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_579])).

cnf(refute_0_581, plain,
    (-1 + skoXC1 < 0 |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     (1/4 + (-1 + skoXC1) * (3/2 + (-1 + skoXC1) * 1/4)) /
     (1 + (-1 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_578, refute_0_580])).

cnf(refute_0_582, plain,
    (skoXC1 < 1 |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_581])).

cnf(refute_0_583, plain,
    (1 <= skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_584, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     (-1 + skoXC1 * (1 + skoXC1 * 1/4)) / skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_583, refute_0_582])).

cnf(refute_0_585, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) * skoXC1 <
     -1 + skoXC1 * (1 + skoXC1 * 1/4) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_577, refute_0_584])).

cnf(refute_0_586, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 * (2 + skoXC1 * (2 + skoXC1 * 1/4)) +
      sqrt(skoXC1) * (skoXC1 * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_585])).

cnf(refute_0_587, plain,
    (0 < skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_588, plain,
    (1 + skoXC1 * (-1 + skoXC1 * -1/4) <
     (skoXC1 * (2 + skoXC1 * (2 + skoXC1 * 1/4)) +
      sqrt(skoXC1) * (skoXC1 * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_586, refute_0_587])).

cnf(refute_0_589, plain,
    ((1 + skoXC1 * (-1 + skoXC1 * -1/4)) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) <
     skoXC1 * (2 + skoXC1 * (2 + skoXC1 * 1/4)) +
     sqrt(skoXC1) * (skoXC1 * (-4 + skoXC1 * (-4 + skoXC1 * -1/2))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6) |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <= 0),
    inference(resolve, [], [refute_0_576, refute_0_588])).

cnf(refute_0_590, plain,
    (34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24))) <
     sqrt(skoXC1) *
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_589])).

cnf(refute_0_591, plain,
    (sqrt(skoXC1) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     0 <=
     20/3 +
     skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(split, [], [refute_0_590])).

cnf(refute_0_592, plain,
    (sqrt(skoXC1) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     7/2 <= sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1 * (10 + skoXC1 * (2 + skoXC1 * 1/12))) <= 20/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_591])).

cnf(refute_0_593, plain,
    (20/3 < skoXC1 * (6 + skoXC1 * (10 + skoXC1 * (2 + skoXC1 * 1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_594, plain,
    (sqrt(skoXC1) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_592, refute_0_593])).

cnf(refute_0_595, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_575, refute_0_594])).

cnf(refute_0_596, plain,
    (0 <= skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_597, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) /
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_596, refute_0_595])).

cnf(refute_0_598, plain,
    (34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (20/3 +
      skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12)))) |
     0 <=
     20/3 +
     skoXC1 * (-6 + skoXC1 * (-10 + skoXC1 * (-2 + skoXC1 * -1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_574, refute_0_597])).

cnf(refute_0_599, plain,
    (34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24))) <
     skoXC1 *
     (80/3 +
      skoXC1 *
      (8/3 +
       skoXC1 *
       (-64 + skoXC1 * (-48 + skoXC1 * (-25/3 + skoXC1 * -1/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (6 + skoXC1 * (10 + skoXC1 * (2 + skoXC1 * 1/12))) <= 20/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_598])).

cnf(refute_0_600, plain,
    (34/3 +
     skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24))) <
     skoXC1 *
     (80/3 +
      skoXC1 *
      (8/3 +
       skoXC1 *
       (-64 + skoXC1 * (-48 + skoXC1 * (-25/3 + skoXC1 * -1/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_599, refute_0_593])).

cnf(refute_0_601, plain,
    ((34/3 +
      skoXC1 * (-3 + skoXC1 * (-14 + skoXC1 * (-4 + skoXC1 * -7/24)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (80/3 +
      skoXC1 *
      (8/3 +
       skoXC1 *
       (-64 + skoXC1 * (-48 + skoXC1 * (-25/3 + skoXC1 * -1/3))))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6) |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_573, refute_0_600])).

cnf(refute_0_602, plain,
    (34/3 <
     skoXC1 *
     (-115/3 +
      skoXC1 *
      (70/3 +
       skoXC1 *
       (27 + skoXC1 * (-233/24 + skoXC1 * (-31/12 + skoXC1 * -1/24))))) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(arithmetic, [], [refute_0_601])).

cnf(refute_0_603, plain,
    (7/2 <= sqrt(skoXC1) |
     skoXC1 *
     (-115/3 +
      skoXC1 *
      (70/3 +
       skoXC1 *
       (27 + skoXC1 * (-233/24 + skoXC1 * (-31/12 + skoXC1 * -1/24))))) <=
     34/3 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_604, plain,
    (7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_603, refute_0_602])).

cnf(refute_0_605, plain,
    (-1 < skoXC1 * (6 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_606, plain,
    (7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) <=
     -34/3 + skoXC1 * (-31/3 + skoXC1 * -7/6)),
    inference(resolve, [], [refute_0_604, refute_0_605])).

cnf(refute_0_607, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     sqrt(skoXC1) * (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_606, refute_0_572])).

cnf(refute_0_608, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-16 + skoXC1 * (-32 + skoXC1 * (-18 + skoXC1 * -2))) |
     1 + skoXC1 * (6 + skoXC1) <= 0 |
     skoXC1 * (-16 + skoXC1 * (-32 + skoXC1 * (-18 + skoXC1 * -2))) /
     (1 + skoXC1 * (6 + skoXC1)) <= -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_609, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     0 <= -4 + skoXC1 * (-4 + skoXC1 * -1/2) |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_610, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <= sqrt(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_611, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)), introduced(tautology, [assume])).

cnf(refute_0_612, plain,
    (sqrt(skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     0 <= -4 + skoXC1 * (-4 + skoXC1 * -1/2) |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(split, [], [refute_0_611])).

cnf(refute_0_613, plain,
    (sqrt(skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     skoXC1 * (4 + skoXC1 * 1/2) <= -4 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(arithmetic, [], [refute_0_612])).

cnf(refute_0_614, plain,
    (-4 < skoXC1 * (4 + skoXC1 * 1/2) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_615, plain,
    (sqrt(skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_613, refute_0_614])).

cnf(refute_0_616, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_610, refute_0_615])).

cnf(refute_0_617, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4)) /
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_596, refute_0_616])).

cnf(refute_0_618, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) |
     0 <= -4 + skoXC1 * (-4 + skoXC1 * -1/2) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_609, refute_0_617])).

cnf(refute_0_619, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     skoXC1 * (-16 + skoXC1 * (-32 + skoXC1 * (-18 + skoXC1 * -2))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (4 + skoXC1 * 1/2) <= -4 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(arithmetic, [], [refute_0_618])).

cnf(refute_0_620, plain,
    (-2 + skoXC1 * (-2 + skoXC1 * -1/4) <
     skoXC1 * (-16 + skoXC1 * (-32 + skoXC1 * (-18 + skoXC1 * -2))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_619, refute_0_614])).

cnf(refute_0_621, plain,
    ((-2 + skoXC1 * (-2 + skoXC1 * -1/4)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (-16 + skoXC1 * (-32 + skoXC1 * (-18 + skoXC1 * -2))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4) | 1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_608, refute_0_620])).

cnf(refute_0_622, plain,
    (-2 <
     skoXC1 * (-2 + skoXC1 * (-71/4 + skoXC1 * (-29/2 + skoXC1 * -7/4))) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(arithmetic, [], [refute_0_621])).

cnf(refute_0_623, plain,
    (7/2 <= sqrt(skoXC1) |
     skoXC1 * (-2 + skoXC1 * (-71/4 + skoXC1 * (-29/2 + skoXC1 * -7/4))) <=
     -2 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_527, refute_0_559])).

cnf(refute_0_624, plain,
    (7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_623, refute_0_622])).

cnf(refute_0_625, plain,
    (7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (-4 + skoXC1 * (-4 + skoXC1 * -1/2)) <=
     -2 + skoXC1 * (-2 + skoXC1 * -1/4)),
    inference(resolve, [], [refute_0_624, refute_0_605])).

cnf(refute_0_626, plain,
    (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     sqrt(skoXC1) * (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_625, refute_0_607])).

cnf(refute_0_627, plain,
    (sqrt(-1 + skoXC1) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_570, refute_0_626])).

cnf(refute_0_628, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_514, refute_0_627])).

cnf(refute_0_629, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_628])).

cnf(refute_0_630, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_583, refute_0_629])).

cnf(refute_0_631, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-2 + skoXC1 * (-2 + skoXC1 * -1/4) +
      sqrt(skoXC1) * (4 + skoXC1 * (4 + skoXC1 * 1/2))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) *
     (-4 + skoXC1 * (4 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_513, refute_0_630])).

cnf(refute_0_632, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (8 + skoXC1 * (skoXC1 * (-9 + skoXC1 * (-3 + skoXC1 * -1/4))) +
      sqrt(skoXC1) *
      (-16 + skoXC1 * (skoXC1 * (18 + skoXC1 * (6 + skoXC1 * 1/2))))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_631])).

cnf(refute_0_633, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) * (1 + skoXC1) <
     5/3 +
     skoXC1 * (67/6 + skoXC1 * (35/4 + skoXC1 * (5/3 + skoXC1 * 1/12))) |
     1 + skoXC1 <= 0 |
     (5/3 +
      skoXC1 * (67/6 + skoXC1 * (35/4 + skoXC1 * (5/3 + skoXC1 * 1/12)))) /
     (1 + skoXC1) <= 34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_634, plain,
    (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) |
     20/3 + skoXC1 * (14/3 + skoXC1 * 1/3) <= 0 |
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) <=
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_635, plain,
    (skoXC1 < 0 |
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(skoXC1) <=
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3))),
    inference(subst, [], [refute_0_167])).

cnf(refute_0_636, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) < sqrt(skoXC1) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     20/3 + skoXC1 * (14/3 + skoXC1 * 1/3) <= 0),
    inference(split, [], [refute_0_626])).

cnf(refute_0_637, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) < sqrt(skoXC1) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (14/3 + skoXC1 * 1/3) <= -20/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_636])).

cnf(refute_0_638, plain,
    (-20/3 < skoXC1 * (14/3 + skoXC1 * 1/3) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_639, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) < sqrt(skoXC1) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_637, refute_0_638])).

cnf(refute_0_640, plain,
    (skoXC1 < 0 |
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_635, refute_0_639])).

cnf(refute_0_641, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) /
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_596, refute_0_640])).

cnf(refute_0_642, plain,
    (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) *
     (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     20/3 + skoXC1 * (14/3 + skoXC1 * 1/3) <= 0),
    inference(resolve, [], [refute_0_634, refute_0_641])).

cnf(refute_0_643, plain,
    (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     (5/3 +
      skoXC1 * (67/6 + skoXC1 * (35/4 + skoXC1 * (5/3 + skoXC1 * 1/12)))) /
     (1 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (14/3 + skoXC1 * 1/3) <= -20/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_642])).

cnf(refute_0_644, plain,
    (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) <
     (5/3 +
      skoXC1 * (67/6 + skoXC1 * (35/4 + skoXC1 * (5/3 + skoXC1 * 1/12)))) /
     (1 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_643, refute_0_638])).

cnf(refute_0_645, plain,
    ((34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)) * (1 + skoXC1) <
     5/3 +
     skoXC1 * (67/6 + skoXC1 * (35/4 + skoXC1 * (5/3 + skoXC1 * 1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_633, refute_0_644])).

cnf(refute_0_646, plain,
    (29/3 <
     skoXC1 * (-21/2 + skoXC1 * (-11/4 + skoXC1 * (1/2 + skoXC1 * 1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_645])).

cnf(refute_0_647, plain,
    (29/3 <
     skoXC1 * (-21/2 + skoXC1 * (-11/4 + skoXC1 * (1/2 + skoXC1 * 1/12))) |
     -1 < skoXC1 | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_648, plain,
    (29/3 <
     skoXC1 * (-21/2 + skoXC1 * (-11/4 + skoXC1 * (1/2 + skoXC1 * 1/12))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_646, refute_0_647])).

cnf(refute_0_649, plain,
    (4 < skoXC1 * (4 + skoXC1) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_27, refute_0_527, refute_0_559, refute_0_648])).

cnf(refute_0_650, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (8 + skoXC1 * (skoXC1 * (-9 + skoXC1 * (-3 + skoXC1 * -1/4))) +
      sqrt(skoXC1) *
      (-16 + skoXC1 * (skoXC1 * (18 + skoXC1 * (6 + skoXC1 * 1/2))))) /
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_632, refute_0_649])).

cnf(refute_0_651, plain,
    (8 + skoXC1 * (skoXC1 * (-9 + skoXC1 * (-3 + skoXC1 * -1/4))) +
     sqrt(skoXC1) *
     (-16 + skoXC1 * (skoXC1 * (18 + skoXC1 * (6 + skoXC1 * 1/2)))) <
     skoXC1 * (-4 + skoXC1 * 4) *
     (34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
      sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3))) |
     0 <=
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6) +
     sqrt(skoXC1) * (-20/3 + skoXC1 * (-14/3 + skoXC1 * -1/3)) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_512, refute_0_650])).

cnf(refute_0_652, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     sqrt(skoXC1) *
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) * (20/3 + skoXC1 * (14/3 + skoXC1 * 1/3)) <=
     34/3 + skoXC1 * (31/3 + skoXC1 * 7/6)),
    inference(arithmetic, [], [refute_0_651])).

cnf(refute_0_653, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     sqrt(skoXC1) *
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_652, refute_0_626])).

cnf(refute_0_654, plain,
    (sqrt(skoXC1) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     0 <=
     16 +
     skoXC1 * (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_653])).

cnf(refute_0_655, plain,
    (sqrt(skoXC1) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     skoXC1 * (-80/3 + skoXC1 * (26 + skoXC1 * (70/3 + skoXC1 * 11/6))) <=
     16 | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_654])).

cnf(refute_0_656, plain,
    (16 <
     skoXC1 * (-80/3 + skoXC1 * (26 + skoXC1 * (70/3 + skoXC1 * 11/6))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_648, refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_657, plain,
    (sqrt(skoXC1) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_655, refute_0_656])).

cnf(refute_0_658, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_511, refute_0_657])).

cnf(refute_0_659, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) /
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_596, refute_0_658])).

cnf(refute_0_660, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (16 +
      skoXC1 *
      (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6)))) |
     0 <=
     16 +
     skoXC1 * (80/3 + skoXC1 * (-26 + skoXC1 * (-70/3 + skoXC1 * -11/6))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_510, refute_0_659])).

cnf(refute_0_661, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 *
     (64 +
      skoXC1 *
      (512/3 +
       skoXC1 *
       (8/3 + skoXC1 * (-592/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     skoXC1 * (-80/3 + skoXC1 * (26 + skoXC1 * (70/3 + skoXC1 * 11/6))) <=
     16 | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_660])).

cnf(refute_0_662, plain,
    (8 +
     skoXC1 *
     (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12))) <
     skoXC1 *
     (64 +
      skoXC1 *
      (512/3 +
       skoXC1 *
       (8/3 + skoXC1 * (-592/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) /
     (1 + skoXC1 * (6 + skoXC1)) | 7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_661, refute_0_656])).

cnf(refute_0_663, plain,
    ((8 +
      skoXC1 *
      (136/3 + skoXC1 * (-13 + skoXC1 * (-119/3 + skoXC1 * -59/12)))) *
     (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 *
     (64 +
      skoXC1 *
      (512/3 +
       skoXC1 *
       (8/3 + skoXC1 * (-592/3 + skoXC1 * (-302/3 + skoXC1 * -22/3))))) |
     7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_509, refute_0_662])).

cnf(refute_0_664, plain,
    (8 <
     skoXC1 *
     (-88/3 +
      skoXC1 *
      (-289/3 +
       skoXC1 *
       (75 + skoXC1 * (703/12 + skoXC1 * (-63/2 + skoXC1 * -29/12))))) |
     7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_663])).

cnf(refute_0_665, plain,
    (7/2 <= sqrt(skoXC1) |
     skoXC1 *
     (-88/3 +
      skoXC1 *
      (-289/3 +
       skoXC1 *
       (75 + skoXC1 * (703/12 + skoXC1 * (-63/2 + skoXC1 * -29/12))))) <=
     8 | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [],
              [refute_0_648, refute_0_559, refute_0_527, refute_0_27])).

cnf(refute_0_666, plain,
    (7/2 <= sqrt(skoXC1) | skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_665, refute_0_664])).

cnf(refute_0_667, plain,
    (7/2 <= sqrt(skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_666, refute_0_605])).

cnf(refute_0_668, plain,
    (7/2 < sqrt(skoXC1) | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_667, refute_0_508])).

cnf(refute_0_669, plain,
    (sqrt(-1 + skoXC1) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_478, refute_0_668])).

cnf(refute_0_670, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_87, refute_0_669])).

cnf(refute_0_671, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_670])).

cnf(refute_0_672, plain,
    (skoXC1 < 0 |
     7/2 < (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_168, refute_0_668])).

cnf(refute_0_673, plain,
    (7/2 < (1/4 + skoXC1 * (3/2 + skoXC1 * 1/4)) / (1 + skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_110, refute_0_672])).

cnf(refute_0_674, plain,
    (7/2 * (1 + skoXC1) < 1/4 + skoXC1 * (3/2 + skoXC1 * 1/4) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_163, refute_0_673])).

cnf(refute_0_675, plain,
    (13/4 < skoXC1 * (-2 + skoXC1 * 1/4) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_674])).

cnf(refute_0_676, plain,
    (13/4 < skoXC1 * (-2 + skoXC1 * 1/4) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_675, refute_0_415])).

cnf(refute_0_677, plain,
    (1 <= skoXC1 | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_676])).

cnf(refute_0_678, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_677, refute_0_671])).

cnf(refute_0_679, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (-1 + sqrt(skoXC1) * 2) / (14/3 + sqrt(skoXC1) * -4/3) *
     (-4 + skoXC1 * (4 + skoXC1)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_86, refute_0_678])).

cnf(refute_0_680, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (4 + skoXC1 * (-4 + skoXC1 * -1) +
      sqrt(skoXC1) * (-8 + skoXC1 * (8 + skoXC1 * 2))) /
     (14/3 + sqrt(skoXC1) * -4/3) | skoXC1 * (4 + skoXC1) <= 4 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_679])).

cnf(refute_0_681, plain,
    (4 < skoXC1 * (4 + skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_27, refute_0_676])).

cnf(refute_0_682, plain,
    (skoXC1 * (-4 + skoXC1 * 4) <
     (4 + skoXC1 * (-4 + skoXC1 * -1) +
      sqrt(skoXC1) * (-8 + skoXC1 * (8 + skoXC1 * 2))) /
     (14/3 + sqrt(skoXC1) * -4/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_680, refute_0_681])).

cnf(refute_0_683, plain,
    (4 + skoXC1 * (-4 + skoXC1 * -1) +
     sqrt(skoXC1) * (-8 + skoXC1 * (8 + skoXC1 * 2)) <
     skoXC1 * (-4 + skoXC1 * 4) * (14/3 + sqrt(skoXC1) * -4/3) |
     0 <= 14/3 + sqrt(skoXC1) * -4/3 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_85, refute_0_682])).

cnf(refute_0_684, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     sqrt(skoXC1) <= 7/2), inference(arithmetic, [], [refute_0_683])).

cnf(refute_0_685, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     sqrt(skoXC1) * (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_684, refute_0_668])).

cnf(refute_0_686, plain,
    (sqrt(skoXC1) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     0 <= 8 + skoXC1 * (-8/3 + skoXC1 * -22/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(split, [], [refute_0_685])).

cnf(refute_0_687, plain,
    (sqrt(skoXC1) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     skoXC1 * (8/3 + skoXC1 * 22/3) <= 8 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_686])).

cnf(refute_0_688, plain,
    (8 < skoXC1 * (8/3 + skoXC1 * 22/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_676, refute_0_27])).

cnf(refute_0_689, plain,
    (sqrt(skoXC1) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_687, refute_0_688])).

cnf(refute_0_690, plain,
    (skoXC1 < 0 |
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_84, refute_0_689])).

cnf(refute_0_691, plain,
    (0 <= skoXC1 | sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_676, refute_0_27])).

cnf(refute_0_692, plain,
    (skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) <
     (4 + skoXC1 * (44/3 + skoXC1 * -59/3)) /
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_691, refute_0_690])).

cnf(refute_0_693, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     skoXC1 * (4 + skoXC1 * 4) / (1 + skoXC1 * (6 + skoXC1)) *
     (8 + skoXC1 * (-8/3 + skoXC1 * -22/3)) |
     0 <= 8 + skoXC1 * (-8/3 + skoXC1 * -22/3) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_83, refute_0_692])).

cnf(refute_0_694, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     skoXC1 * (32 + skoXC1 * (64/3 + skoXC1 * (-40 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) | skoXC1 * (8/3 + skoXC1 * 22/3) <= 8 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_693])).

cnf(refute_0_695, plain,
    (4 + skoXC1 * (44/3 + skoXC1 * -59/3) <
     skoXC1 * (32 + skoXC1 * (64/3 + skoXC1 * (-40 + skoXC1 * -88/3))) /
     (1 + skoXC1 * (6 + skoXC1)) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_694, refute_0_688])).

cnf(refute_0_696, plain,
    ((4 + skoXC1 * (44/3 + skoXC1 * -59/3)) * (1 + skoXC1 * (6 + skoXC1)) <
     skoXC1 * (32 + skoXC1 * (64/3 + skoXC1 * (-40 + skoXC1 * -88/3))) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0 |
     1 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_82, refute_0_695])).

cnf(refute_0_697, plain,
    (4 <
     skoXC1 *
     (-20/3 + skoXC1 * (-51 + skoXC1 * (190/3 + skoXC1 * -29/3))) |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(arithmetic, [], [refute_0_696])).

cnf(refute_0_698, plain,
    (skoXC1 *
     (-20/3 + skoXC1 * (-51 + skoXC1 * (190/3 + skoXC1 * -29/3))) <= 4 |
     skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_676, refute_0_27])).

cnf(refute_0_699, plain,
    (skoXC1 * (6 + skoXC1) <= -1 |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_698, refute_0_697])).

cnf(refute_0_700, plain,
    (-1 < skoXC1 * (6 + skoXC1) |
     sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(decision, [], [refute_0_676, refute_0_27])).

cnf(refute_0_701, plain, (sqrt(1 + skoXC1) * (sqrt(-1 + skoXC1) * 2) <= 0),
    inference(resolve, [], [refute_0_699, refute_0_700])).

cnf(refute_0_702, plain,
    (sqrt(-1 + skoXC1) <
     sqrt(1 + skoXC1) * (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_701, refute_0_81])).

cnf(refute_0_703, plain,
    (sqrt(1 + skoXC1) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) |
     0 <= -1 + sqrt(-1 + skoXC1) * -8/3),
    inference(split, [], [refute_0_702])).

cnf(refute_0_704, plain,
    (sqrt(1 + skoXC1) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) |
     sqrt(-1 + skoXC1) <= -3/8),
    inference(arithmetic, [], [refute_0_703])).

cnf(refute_0_705, plain,
    (-1 + sqrt(-1 + skoXC1) * -8/3 < 0 | sqrt(-1 + skoXC1) < 0 |
     0 < -1 + sqrt(-1 + skoXC1) * -8/3),
    inference(split, [], [refute_0_702])).

cnf(refute_0_706, plain,
    (sqrt(-1 + skoXC1) < -3/8 | -3/8 < sqrt(-1 + skoXC1)),
    inference(arithmetic, [], [refute_0_705])).

cnf(refute_0_707, plain,
    (skoXC1 * (-4 + skoXC1 * 4) < -3/8 * (-4 + skoXC1 * (4 + skoXC1)) |
     -3/8 <= skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_708, plain,
    ((-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) < -3/8 | -1 + skoXC1 < 0 |
     -3/8 <= sqrt(-1 + skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_709, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) < -3/8 |
     skoXC1 < 1 | -3/8 <= sqrt(-1 + skoXC1)),
    inference(arithmetic, [], [refute_0_708])).

cnf(refute_0_710, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) < -3/8 |
     -3/8 <= sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_58, refute_0_709])).

cnf(refute_0_711, plain,
    (skoXC1 * (-4 + skoXC1 * 4) < -3/8 * (-4 + skoXC1 * (4 + skoXC1)) |
     -3/8 <= sqrt(-1 + skoXC1) | -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_707, refute_0_710])).

cnf(refute_0_712, plain,
    (-3/2 < skoXC1 * (5/2 + skoXC1 * -35/8) | -3/8 <= sqrt(-1 + skoXC1) |
     skoXC1 * (4 + skoXC1) <= 4),
    inference(arithmetic, [], [refute_0_711])).

cnf(refute_0_713, plain,
    (skoXC1 * (5/2 + skoXC1 * -35/8) <= -3/2 | skoXC1 * (4 + skoXC1) <= 4),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_714, plain,
    (-3/8 <= sqrt(-1 + skoXC1) | skoXC1 * (4 + skoXC1) <= 4),
    inference(resolve, [], [refute_0_713, refute_0_712])).

cnf(refute_0_715, plain, (-3/8 <= sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_714, refute_0_64])).

cnf(refute_0_716, plain, (-3/8 < sqrt(-1 + skoXC1)),
    inference(resolve, [], [refute_0_715, refute_0_706])).

cnf(refute_0_717, plain,
    (sqrt(1 + skoXC1) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_704, refute_0_716])).

cnf(refute_0_718, plain,
    (1 + skoXC1 < 0 |
     (1 + skoXC1) * (4 + (1 + skoXC1) * 4) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1))) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_9, refute_0_717])).

cnf(refute_0_719, plain,
    (skoXC1 < -1 |
     (8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(arithmetic, [], [refute_0_718])).

cnf(refute_0_720, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) / (8 + skoXC1 * (8 + skoXC1)) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_32, refute_0_719])).

cnf(refute_0_721, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) / (-1 + sqrt(-1 + skoXC1) * -8/3) *
     (8 + skoXC1 * (8 + skoXC1)) | 8 + skoXC1 * (8 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_8, refute_0_720])).

cnf(refute_0_722, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (8 + skoXC1 * (8 + skoXC1)) /
     (-1 + sqrt(-1 + skoXC1) * -8/3) | skoXC1 * (8 + skoXC1) <= -8),
    inference(arithmetic, [], [refute_0_721])).

cnf(refute_0_723, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (8 + skoXC1 * (8 + skoXC1)) /
     (-1 + sqrt(-1 + skoXC1) * -8/3)),
    inference(resolve, [], [refute_0_722, refute_0_38])).

cnf(refute_0_724, plain,
    (sqrt(-1 + skoXC1) * (8 + skoXC1 * (8 + skoXC1)) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) * (-1 + sqrt(-1 + skoXC1) * -8/3) |
     0 <= -1 + sqrt(-1 + skoXC1) * -8/3),
    inference(resolve, [], [refute_0_7, refute_0_723])).

cnf(refute_0_725, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     sqrt(-1 + skoXC1) <= -3/8),
    inference(arithmetic, [], [refute_0_724])).

cnf(refute_0_726, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     sqrt(-1 + skoXC1) * (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3))),
    inference(resolve, [], [refute_0_725, refute_0_716])).

cnf(refute_0_727, plain,
    (sqrt(-1 + skoXC1) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     0 <= -88/3 + skoXC1 * (-40 + skoXC1 * -35/3)),
    inference(split, [], [refute_0_726])).

cnf(refute_0_728, plain,
    (sqrt(-1 + skoXC1) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     skoXC1 * (40 + skoXC1 * 35/3) <= -88/3),
    inference(arithmetic, [], [refute_0_727])).

cnf(refute_0_729, plain, (-88/3 < skoXC1 * (40 + skoXC1 * 35/3)),
    inference(decision, [], [refute_0_27])).

cnf(refute_0_730, plain,
    (sqrt(-1 + skoXC1) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3))),
    inference(resolve, [], [refute_0_728, refute_0_729])).

cnf(refute_0_731, plain,
    (-1 + skoXC1 < 0 |
     (-1 + skoXC1) * (4 + (-1 + skoXC1) * 4) /
     (1 + (-1 + skoXC1) * (6 + (-1 + skoXC1))) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3))),
    inference(resolve, [], [refute_0_6, refute_0_730])).

cnf(refute_0_732, plain,
    (skoXC1 < 1 |
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3))),
    inference(arithmetic, [], [refute_0_731])).

cnf(refute_0_733, plain,
    (skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) <
     (8 + skoXC1 * (12 + skoXC1 * 4)) /
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3))),
    inference(resolve, [], [refute_0_58, refute_0_732])).

cnf(refute_0_734, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     skoXC1 * (-4 + skoXC1 * 4) / (-4 + skoXC1 * (4 + skoXC1)) *
     (-88/3 + skoXC1 * (-40 + skoXC1 * -35/3)) |
     0 <= -88/3 + skoXC1 * (-40 + skoXC1 * -35/3)),
    inference(resolve, [], [refute_0_1, refute_0_733])).

cnf(refute_0_735, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     skoXC1 *
     (352/3 + skoXC1 * (128/3 + skoXC1 * (-340/3 + skoXC1 * -140/3))) /
     (-4 + skoXC1 * (4 + skoXC1)) |
     skoXC1 * (40 + skoXC1 * 35/3) <= -88/3),
    inference(arithmetic, [], [refute_0_734])).

cnf(refute_0_736, plain,
    (8 + skoXC1 * (12 + skoXC1 * 4) <
     skoXC1 *
     (352/3 + skoXC1 * (128/3 + skoXC1 * (-340/3 + skoXC1 * -140/3))) /
     (-4 + skoXC1 * (4 + skoXC1))),
    inference(resolve, [], [refute_0_735, refute_0_729])).

cnf(refute_0_737, plain,
    ((8 + skoXC1 * (12 + skoXC1 * 4)) * (-4 + skoXC1 * (4 + skoXC1)) <
     skoXC1 *
     (352/3 + skoXC1 * (128/3 + skoXC1 * (-340/3 + skoXC1 * -140/3))) |
     -4 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_736])).

cnf(refute_0_738, plain,
    (-32 <
     skoXC1 *
     (400/3 + skoXC1 * (8/3 + skoXC1 * (-424/3 + skoXC1 * -152/3))) |
     skoXC1 * (4 + skoXC1) <= 4),
    inference(arithmetic, [], [refute_0_737])).

cnf(refute_0_739, plain,
    (skoXC1 * (4 + skoXC1) <= 4 |
     skoXC1 *
     (400/3 + skoXC1 * (8/3 + skoXC1 * (-424/3 + skoXC1 * -152/3))) <=
     -32), inference(decision, [], [refute_0_27])).

cnf(refute_0_740, plain, (skoXC1 * (4 + skoXC1) <= 4),
    inference(resolve, [], [refute_0_739, refute_0_738])).

cnf(refute_0_741, plain, ($false),
    inference(resolve, [], [refute_0_740, refute_0_64])).
% SZS output end CNFRefutation for exp-problem-10-2-sqrt-weak.tptp
