% SZS output start CNFRefutation for exp-error-analysis-7.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf3, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf1, axiom,
    (X < 0 | 3/2 < X | ~ lgen(R, -(X + 2) / (X - 2), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_lower_bound_cf5, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf4, axiom,
    (0 <= X |
     ~ lgen(R,
            (X ^ 4 + 20 * X ^ 3 + 180 * X ^ 2 + 840 * X + 1680) /
            (X ^ 4 - 20 * X ^ 3 + 180 * X ^ 2 - 840 * X + 1680), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf6, axiom,
    (0 <= X |
     ~ lgen(R,
            (X ^ 6 + 42 * X ^ 5 + 840 * X ^ 4 + 10080 * X ^ 3 +
             75600 * X ^ 2 + 332640 * X + 665280) /
            (X ^ 6 - 42 * X ^ 5 + 840 * X ^ 4 - 10080 * X ^ 3 +
             75600 * X ^ 2 - 332640 * X + 665280), Y) |
     lgen(R, exp(X), Y))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(exp_error_analysis_7, conjecture,
    (! [X] :
       ((0 <= X & X <= 693/1000) =>
        abs(exp(-X) -
            (1 - 1999999999/2000000000 * X +
             2499999603/5000000000 * X ^ 2 -
             1666653019/10000000000 * X ^ 3 + 16662939/400000000 * X ^ 4 -
             41506799/5000000000 * X ^ 5 + 664941/500000000 * X ^ 6 -
             1413161/10000000000 * X ^ 7)) <= 1/5000000000))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 693/1000) =>
        abs(exp(-X) -
            (1 - 1999999999/2000000000 * X +
             2499999603/5000000000 * X ^ 2 -
             1666653019/10000000000 * X ^ 3 + 16662939/400000000 * X ^ 4 -
             41506799/5000000000 * X ^ 5 + 664941/500000000 * X ^ 6 -
             1413161/10000000000 * X ^ 7)) <= 1/5000000000)),
    inference(strip, [], [exp_error_analysis_7])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 693/1000) =>
          abs(exp(-X) -
              (1 - 1999999999/2000000000 * X +
               2499999603/5000000000 * X ^ 2 -
               1666653019/10000000000 * X ^ 3 +
               16662939/400000000 * X ^ 4 - 41506799/5000000000 * X ^ 5 +
               664941/500000000 * X ^ 6 - 1413161/10000000000 * X ^ 7)) <=
          1/5000000000)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (1/5000000000 <
        abs(exp(-X) -
            (1 - 1999999999/2000000000 * X +
             2499999603/5000000000 * X ^ 2 -
             1666653019/10000000000 * X ^ 3 + 16662939/400000000 * X ^ 4 -
             41506799/5000000000 * X ^ 5 + 664941/500000000 * X ^ 6 -
             1413161/10000000000 * X ^ 7)) & X <= 693/1000 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (1/5000000000 <
     abs(exp(-skoXC1) -
         (1 - 1999999999/2000000000 * skoXC1 +
          2499999603/5000000000 * skoXC1 ^ 2 -
          1666653019/10000000000 * skoXC1 ^ 3 +
          16662939/400000000 * skoXC1 ^ 4 -
          41506799/5000000000 * skoXC1 ^ 5 +
          664941/500000000 * skoXC1 ^ 6 -
          1413161/10000000000 * skoXC1 ^ 7)) & 0 <= skoXC1 &
     skoXC1 <= 693/1000), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/5000000000 <
     abs(exp(-skoXC1) -
         (1 - 1999999999/2000000000 * skoXC1 +
          2499999603/5000000000 * skoXC1 ^ 2 -
          1666653019/10000000000 * skoXC1 ^ 3 +
          16662939/400000000 * skoXC1 ^ 4 -
          41506799/5000000000 * skoXC1 ^ 5 +
          664941/500000000 * skoXC1 ^ 6 -
          1413161/10000000000 * skoXC1 ^ 7))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 693/1000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-2 + skoXC1 <
     (5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-2 + skoXC1 * -1) | 0 <= -2 + skoXC1 * -1 |
     (-2 + skoXC1) / (-2 + skoXC1 * -1) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (1/5000000000 <
     abs(exp(-skoXC1) -
         (1 - 1999999999/2000000000 * skoXC1 +
          2499999603/5000000000 * skoXC1 ^ 2 -
          1666653019/10000000000 * skoXC1 ^ 3 +
          16662939/400000000 * skoXC1 ^ 4 -
          41506799/5000000000 * skoXC1 ^ 5 +
          664941/500000000 * skoXC1 ^ 6 -
          1413161/10000000000 * skoXC1 ^ 7))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain,
    (1/5000000000 <
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1))), inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (-1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) < 0 |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) =
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1)), inference(subst, [], [abs_nonnegative])).

cnf(refute_0_4, plain,
    (1/5000000000 <
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) !=
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) <= 1/5000000000),
    introduced(tautology, [equality])).

cnf(refute_0_5, plain,
    (-1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) < 0 |
     1/5000000000 <
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) <= 1/5000000000),
    inference(resolve, [], [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain,
    (-1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) < 0 |
     1/5000000000 <
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1)), inference(resolve, [], [refute_0_5, refute_0_2])).

cnf(refute_0_7, plain,
    (exp(skoXC1 * -1) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (0 <=
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) =
     -(-1 +
      skoXC1 *
      (1999999999/2000000000 +
       skoXC1 *
       (-2499999603/5000000000 +
        skoXC1 *
        (1666653019/10000000000 +
         skoXC1 *
         (-16662939/400000000 +
          skoXC1 *
          (41506799/5000000000 +
           skoXC1 *
           (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
      exp(skoXC1 * -1))), inference(subst, [], [abs_negative])).

cnf(refute_0_9, plain,
    (1/5000000000 <
     -(-1 +
      skoXC1 *
      (1999999999/2000000000 +
       skoXC1 *
       (-2499999603/5000000000 +
        skoXC1 *
        (1666653019/10000000000 +
         skoXC1 *
         (-16662939/400000000 +
          skoXC1 *
          (41506799/5000000000 +
           skoXC1 *
           (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
      exp(skoXC1 * -1)) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) !=
     -(-1 +
      skoXC1 *
      (1999999999/2000000000 +
       skoXC1 *
       (-2499999603/5000000000 +
        skoXC1 *
        (1666653019/10000000000 +
         skoXC1 *
         (-16662939/400000000 +
          skoXC1 *
          (41506799/5000000000 +
           skoXC1 *
           (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
      exp(skoXC1 * -1)) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) <= 1/5000000000),
    introduced(tautology, [equality])).

cnf(refute_0_10, plain,
    (1/5000000000 <
     -(-1 +
      skoXC1 *
      (1999999999/2000000000 +
       skoXC1 *
       (-2499999603/5000000000 +
        skoXC1 *
        (1666653019/10000000000 +
         skoXC1 *
         (-16662939/400000000 +
          skoXC1 *
          (41506799/5000000000 +
           skoXC1 *
           (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
      exp(skoXC1 * -1)) |
     0 <=
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1) |
     abs(-1 +
         skoXC1 *
         (1999999999/2000000000 +
          skoXC1 *
          (-2499999603/5000000000 +
           skoXC1 *
           (1666653019/10000000000 +
            skoXC1 *
            (-16662939/400000000 +
             skoXC1 *
             (41506799/5000000000 +
              skoXC1 *
              (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
         exp(skoXC1 * -1)) <= 1/5000000000),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (1/5000000000 <
     -(-1 +
      skoXC1 *
      (1999999999/2000000000 +
       skoXC1 *
       (-2499999603/5000000000 +
        skoXC1 *
        (1666653019/10000000000 +
         skoXC1 *
         (-16662939/400000000 +
          skoXC1 *
          (41506799/5000000000 +
           skoXC1 *
           (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
      exp(skoXC1 * -1)) |
     0 <=
     -1 +
     skoXC1 *
     (1999999999/2000000000 +
      skoXC1 *
      (-2499999603/5000000000 +
       skoXC1 *
       (1666653019/10000000000 +
        skoXC1 *
        (-16662939/400000000 +
         skoXC1 *
         (41506799/5000000000 +
          skoXC1 *
          (-664941/500000000 + skoXC1 * 1413161/10000000000)))))) +
     exp(skoXC1 * -1)), inference(resolve, [], [refute_0_10, refute_0_2])).

cnf(refute_0_12, plain,
    (exp(skoXC1 * -1) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((4999999999/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     0 <=
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_14, plain,
    (~ lgen(0, X_000106, exp(X_000105)) | X_000106 <= exp(X_000105)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_15, plain,
    (~ lgen(0, X_000106,
            -(X_000105 ^ 5 + 30 * X_000105 ^ 4 + 420 * X_000105 ^ 3 +
             3360 * X_000105 ^ 2 + 15120 * X_000105 + 30240) /
            (X_000105 ^ 5 - 30 * X_000105 ^ 4 + 420 * X_000105 ^ 3 -
             3360 * X_000105 ^ 2 + 15120 * X_000105 - 30240)) |
     0 <= X_000105 | lgen(0, X_000106, exp(X_000105))),
    inference(subst, [], [exp_lower_bound_cf5])).

cnf(refute_0_16, plain,
    (~ lgen(0, X_000106,
            -(X_000105 ^ 5 + 30 * X_000105 ^ 4 + 420 * X_000105 ^ 3 +
             3360 * X_000105 ^ 2 + 15120 * X_000105 + 30240) /
            (X_000105 ^ 5 - 30 * X_000105 ^ 4 + 420 * X_000105 ^ 3 -
             3360 * X_000105 ^ 2 + 15120 * X_000105 - 30240)) |
     X_000106 <= exp(X_000105) | 0 <= X_000105),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((-30240 +
      X_000105 *
      (-15120 +
       X_000105 *
       (-3360 + X_000105 * (-420 + X_000105 * (-30 + X_000105 * -1))))) /
     (-30240 +
      X_000105 *
      (15120 +
       X_000105 *
       (-3360 + X_000105 * (420 + X_000105 * (-30 + X_000105))))) <
     X_000106 | X_000106 <= exp(X_000105) | 0 <= X_000105),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((-30240 +
      skoXC1 * -1 *
      (-15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (-420 + skoXC1 * -1 * (-30 + skoXC1 * -1 * -1))))) /
     (-30240 +
      skoXC1 * -1 *
      (15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (420 + skoXC1 * -1 * (-30 + skoXC1 * -1))))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     0 <= skoXC1 * -1 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_17])).

cnf(refute_0_19, plain,
    ((-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000050, exp(X_000049)) | X_000050 <= exp(X_000049)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_21, plain,
    (X_000049 < -1 | ~ lgen(0, X_000050, 1 + X_000049) |
     lgen(0, X_000050, exp(X_000049))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_22, plain,
    (X_000049 < -1 | ~ lgen(0, X_000050, 1 + X_000049) |
     X_000050 <= exp(X_000049)),
    inference(resolve, [], [refute_0_21, refute_0_20])).

cnf(refute_0_23, plain,
    (1 + X_000049 < X_000050 | X_000049 < -1 | X_000050 <= exp(X_000049)),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * -1 < -1 |
     1 + skoXC1 * -1 <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (1/5000000000 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 < skoXC1 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 <= 693/1000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_28, plain,
    (1/5000000000 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     skoXC1 <= 1), inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (1/5000000000 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_28, refute_0_25])).

cnf(refute_0_30, plain,
    (skoXC1 * -1 < -1 |
     1 + skoXC1 * -1 <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_23])).

cnf(refute_0_31, plain,
    (0 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 < skoXC1 |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (0 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     skoXC1 <= 1), inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_33, plain,
    (0 <
     skoXC1 *
     (1/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (0 < skoXC1 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(decision, [],
              [refute_0_29, refute_0_33, refute_0_26, refute_0_27])).

cnf(refute_0_35, plain,
    ((-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_19, refute_0_34])).

cnf(refute_0_36, plain,
    ((4999999999/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     0 <=
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_13, refute_0_35])).

cnf(refute_0_37, plain,
    (189/31250000 <
     skoXC1 *
     (189/15625000 +
      skoXC1 *
      (-299271/125000000 +
       skoXC1 *
       (10017903/250000000 +
        skoXC1 *
        (-65360727/250000000 +
         skoXC1 *
         (2076315327/2500000000 +
          skoXC1 *
          (-2663340151/2000000000 +
           skoXC1 *
           (1172468577/1250000000 +
            skoXC1 *
            (-718119049/10000000000 +
             skoXC1 *
             (-49462041/400000000 +
              skoXC1 *
              (-138788309/5000000000 +
               skoXC1 *
               (-2909601/1000000000 +
                skoXC1 * -1413161/10000000000))))))))))) |
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <=
     -30240 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    ((4999999999/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_39, plain,
    (~ lgen(0, X_000085, exp(X_000084)) | X_000085 <= exp(X_000084)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_40, plain,
    (~ lgen(0, X_000085,
            -(X_000084 ^ 3 + 12 * X_000084 ^ 2 + 60 * X_000084 + 120) /
            (X_000084 ^ 3 - 12 * X_000084 ^ 2 + 60 * X_000084 - 120)) |
     0 <= X_000084 | lgen(0, X_000085, exp(X_000084))),
    inference(subst, [], [exp_lower_bound_cf3])).

cnf(refute_0_41, plain,
    (~ lgen(0, X_000085,
            -(X_000084 ^ 3 + 12 * X_000084 ^ 2 + 60 * X_000084 + 120) /
            (X_000084 ^ 3 - 12 * X_000084 ^ 2 + 60 * X_000084 - 120)) |
     X_000085 <= exp(X_000084) | 0 <= X_000084),
    inference(resolve, [], [refute_0_40, refute_0_39])).

cnf(refute_0_42, plain,
    ((-120 + X_000084 * (-60 + X_000084 * (-12 + X_000084 * -1))) /
     (-120 + X_000084 * (60 + X_000084 * (-12 + X_000084))) < X_000085 |
     X_000085 <= exp(X_000084) | 0 <= X_000084),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    ((-120 +
      skoXC1 * -1 * (-60 + skoXC1 * -1 * (-12 + skoXC1 * -1 * -1))) /
     (-120 + skoXC1 * -1 * (60 + skoXC1 * -1 * (-12 + skoXC1 * -1))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     0 <= skoXC1 * -1 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_42])).

cnf(refute_0_44, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_44, refute_0_34])).

cnf(refute_0_46, plain,
    ((4999999999/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_38, refute_0_45])).

cnf(refute_0_47, plain,
    (3/125000000 <
     skoXC1 *
     (3/62500000 +
      skoXC1 *
      (-23751/2500000000 +
       skoXC1 *
       (795069/5000000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    ((1 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_49, plain,
    ((-120 +
      skoXC1 * -1 * (-60 + skoXC1 * -1 * (-12 + skoXC1 * -1 * -1))) /
     (-120 + skoXC1 * -1 * (60 + skoXC1 * -1 * (-12 + skoXC1 * -1))) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     0 <= skoXC1 * -1 |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(subst, [], [refute_0_42])).

cnf(refute_0_50, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_50, refute_0_34])).

cnf(refute_0_52, plain,
    ((1 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <
     -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_48, refute_0_51])).

cnf(refute_0_53, plain,
    (0 <
     skoXC1 *
     (3/50000000 +
      skoXC1 *
      (-4749/500000000 +
       skoXC1 *
       (79507/500000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)), inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (0 <
     skoXC1 *
     (3/50000000 +
      skoXC1 *
      (-4749/500000000 +
       skoXC1 *
       (79507/500000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     -120 < skoXC1 * (60 + skoXC1 * (12 + skoXC1)) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(decision, [],
              [refute_0_29, refute_0_33, refute_0_26, refute_0_27])).

cnf(refute_0_55, plain,
    (0 <
     skoXC1 *
     (3/50000000 +
      skoXC1 *
      (-4749/500000000 +
       skoXC1 *
       (79507/500000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (3/125000000 <
     skoXC1 *
     (3/62500000 +
      skoXC1 *
      (-23751/2500000000 +
       skoXC1 *
       (795069/5000000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     -120 < skoXC1 * (60 + skoXC1 * (12 + skoXC1)) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(decision, [],
              [refute_0_55, refute_0_29, refute_0_33, refute_0_26,
               refute_0_27])).

cnf(refute_0_57, plain,
    (3/125000000 <
     skoXC1 *
     (3/62500000 +
      skoXC1 *
      (-23751/2500000000 +
       skoXC1 *
       (795069/5000000000 +
        skoXC1 *
        (-10373663/10000000000 +
         skoXC1 *
         (16469859/5000000000 +
          skoXC1 *
          (-52728799/10000000000 +
           skoXC1 *
           (48760179/10000000000 +
            skoXC1 *
            (-4108709/5000000000 +
             skoXC1 *
             (-457389/1250000000 + skoXC1 * -1413161/10000000000))))))))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_47, refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 *
     (189/15625000 +
      skoXC1 *
      (-299271/125000000 +
       skoXC1 *
       (10017903/250000000 +
        skoXC1 *
        (-65360727/250000000 +
         skoXC1 *
         (2076315327/2500000000 +
          skoXC1 *
          (-2663340151/2000000000 +
           skoXC1 *
           (1172468577/1250000000 +
            skoXC1 *
            (-718119049/10000000000 +
             skoXC1 *
             (-49462041/400000000 +
              skoXC1 *
              (-138788309/5000000000 +
               skoXC1 *
               (-2909601/1000000000 +
                skoXC1 * -1413161/10000000000))))))))))) <= 189/31250000 |
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <=
     -30240 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(decision, [],
              [refute_0_57, refute_0_55, refute_0_29, refute_0_33,
               refute_0_26, refute_0_27])).

cnf(refute_0_59, plain,
    (skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <=
     -30240 |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_58, refute_0_37])).

cnf(refute_0_60, plain,
    (-30240 <
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) |
     4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(decision, [],
              [refute_0_57, refute_0_55, refute_0_29, refute_0_33,
               refute_0_26, refute_0_27])).

cnf(refute_0_61, plain,
    (4999999999/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1) |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_61, refute_0_12])).

cnf(refute_0_63, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     exp(skoXC1 * -1)), inference(resolve, [], [refute_0_62, refute_0_7])).

cnf(refute_0_64, plain,
    (~ lgen(0, exp(X_000133), X_000134) | exp(X_000133) <= X_000134),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_65, plain,
    (3/2 < X_000133 | X_000133 < 0 |
     ~ lgen(0, -(X_000133 + 2) / (X_000133 - 2), X_000134) |
     lgen(0, exp(X_000133), X_000134)),
    inference(subst, [], [exp_upper_bound_cf1])).

cnf(refute_0_66, plain,
    (3/2 < X_000133 | X_000133 < 0 |
     ~ lgen(0, -(X_000133 + 2) / (X_000133 - 2), X_000134) |
     exp(X_000133) <= X_000134),
    inference(resolve, [], [refute_0_65, refute_0_64])).

cnf(refute_0_67, plain,
    (3/2 < X_000133 | X_000133 < 0 |
     X_000134 < (-2 + X_000133 * -1) / (-2 + X_000133) |
     exp(X_000133) <= X_000134), inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * -1 < 0 | 3/2 < skoXC1 * -1 |
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (-2 + skoXC1 * -1 * -1) / (-2 + skoXC1 * -1) |
     exp(skoXC1 * -1) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (skoXC1 * -1 < 0 | 3/2 < skoXC1 * -1 |
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (-2 + skoXC1 * -1 * -1) / (-2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_68, refute_0_63])).

cnf(refute_0_70, plain,
    (skoXC1 < -3/2 |
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (-2 + skoXC1) / (-2 + skoXC1 * -1) | 0 < skoXC1),
    inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    ((5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) <
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) |
     665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <= 0 |
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_72, plain,
    (~ lgen(0, exp(X_000119), X_000120) | exp(X_000119) <= X_000120),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_73, plain,
    (~ lgen(0,
            (X_000119 ^ 6 + 42 * X_000119 ^ 5 + 840 * X_000119 ^ 4 +
             10080 * X_000119 ^ 3 + 75600 * X_000119 ^ 2 +
             332640 * X_000119 + 665280) /
            (X_000119 ^ 6 - 42 * X_000119 ^ 5 + 840 * X_000119 ^ 4 -
             10080 * X_000119 ^ 3 + 75600 * X_000119 ^ 2 -
             332640 * X_000119 + 665280), X_000120) | 0 <= X_000119 |
     lgen(0, exp(X_000119), X_000120)),
    inference(subst, [], [exp_upper_bound_cf6])).

cnf(refute_0_74, plain,
    (~ lgen(0,
            (X_000119 ^ 6 + 42 * X_000119 ^ 5 + 840 * X_000119 ^ 4 +
             10080 * X_000119 ^ 3 + 75600 * X_000119 ^ 2 +
             332640 * X_000119 + 665280) /
            (X_000119 ^ 6 - 42 * X_000119 ^ 5 + 840 * X_000119 ^ 4 -
             10080 * X_000119 ^ 3 + 75600 * X_000119 ^ 2 -
             332640 * X_000119 + 665280), X_000120) | 0 <= X_000119 |
     exp(X_000119) <= X_000120),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    (X_000120 <
     (665280 +
      X_000119 *
      (332640 +
       X_000119 *
       (75600 +
        X_000119 *
        (10080 + X_000119 * (840 + X_000119 * (42 + X_000119)))))) /
     (665280 +
      X_000119 *
      (-332640 +
       X_000119 *
       (75600 +
        X_000119 *
        (-10080 + X_000119 * (840 + X_000119 * (-42 + X_000119)))))) |
     0 <= X_000119 | exp(X_000119) <= X_000120),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (665280 +
      skoXC1 * -1 *
      (332640 +
       skoXC1 * -1 *
       (75600 +
        skoXC1 * -1 *
        (10080 +
         skoXC1 * -1 * (840 + skoXC1 * -1 * (42 + skoXC1 * -1)))))) /
     (665280 +
      skoXC1 * -1 *
      (-332640 +
       skoXC1 * -1 *
       (75600 +
        skoXC1 * -1 *
        (-10080 +
         skoXC1 * -1 * (840 + skoXC1 * -1 * (-42 + skoXC1 * -1)))))) |
     0 <= skoXC1 * -1 |
     exp(skoXC1 * -1) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (665280 +
      skoXC1 * -1 *
      (332640 +
       skoXC1 * -1 *
       (75600 +
        skoXC1 * -1 *
        (10080 +
         skoXC1 * -1 * (840 + skoXC1 * -1 * (42 + skoXC1 * -1)))))) /
     (665280 +
      skoXC1 * -1 *
      (-332640 +
       skoXC1 * -1 *
       (75600 +
        skoXC1 * -1 *
        (-10080 +
         skoXC1 * -1 * (840 + skoXC1 * -1 * (-42 + skoXC1 * -1)))))) |
     0 <= skoXC1 * -1),
    inference(resolve, [], [refute_0_76, refute_0_63])).

cnf(refute_0_78, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    ((5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) <
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) |
     665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <= 0 |
     skoXC1 <= 0), inference(resolve, [], [refute_0_71, refute_0_78])).

cnf(refute_0_80, plain,
    (2079/15625000 <
     skoXC1 *
     (-6237/15625000 +
      skoXC1 *
      (102816/1953125 +
       skoXC1 *
       (-110197521/125000000 +
        skoXC1 *
        (718984581/125000000 +
         skoXC1 *
         (-45684502731/2500000000 +
          skoXC1 *
          (73278164347/2500000000 +
           skoXC1 *
           (-206815784617/10000000000 +
            skoXC1 *
            (1033616217/625000000 +
             skoXC1 *
             (26684855389/10000000000 +
              skoXC1 *
              (6143651721/10000000000 +
               skoXC1 *
               (355759199/5000000000 +
                skoXC1 *
                (23026971/5000000000 +
                 skoXC1 * 1413161/10000000000)))))))))))) |
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <=
     -665280 | skoXC1 <= 0), inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    ((1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     (1 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <=
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) |
     1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_82, plain,
    (exp(X_000098) < X_000099 | ~ lgen(1, exp(X_000098), X_000099)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_83, plain,
    (~ lgen(1,
            (X_000098 ^ 4 + 20 * X_000098 ^ 3 + 180 * X_000098 ^ 2 +
             840 * X_000098 + 1680) /
            (X_000098 ^ 4 - 20 * X_000098 ^ 3 + 180 * X_000098 ^ 2 -
             840 * X_000098 + 1680), X_000099) | 0 <= X_000098 |
     lgen(1, exp(X_000098), X_000099)),
    inference(subst, [], [exp_upper_bound_cf4])).

cnf(refute_0_84, plain,
    (exp(X_000098) < X_000099 |
     ~ lgen(1,
            (X_000098 ^ 4 + 20 * X_000098 ^ 3 + 180 * X_000098 ^ 2 +
             840 * X_000098 + 1680) /
            (X_000098 ^ 4 - 20 * X_000098 ^ 3 + 180 * X_000098 ^ 2 -
             840 * X_000098 + 1680), X_000099) | 0 <= X_000098),
    inference(resolve, [], [refute_0_83, refute_0_82])).

cnf(refute_0_85, plain,
    (exp(X_000098) < X_000099 |
     X_000099 <=
     (1680 +
      X_000098 * (840 + X_000098 * (180 + X_000098 * (20 + X_000098)))) /
     (1680 +
      X_000098 * (-840 + X_000098 * (180 + X_000098 * (-20 + X_000098)))) |
     0 <= X_000098), inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (exp(skoXC1 * -1) <
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) |
     0 <= skoXC1 * -1 |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (1680 +
      skoXC1 * -1 *
      (840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (20 + skoXC1 * -1)))) /
     (1680 +
      skoXC1 * -1 *
      (-840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (-20 + skoXC1 * -1))))),
    inference(subst, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 <= skoXC1 * -1 |
     1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (1680 +
      skoXC1 * -1 *
      (840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (20 + skoXC1 * -1)))) /
     (1680 +
      skoXC1 * -1 *
      (-840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (-20 + skoXC1 * -1))))),
    inference(resolve, [], [refute_0_62, refute_0_86])).

cnf(refute_0_88, plain,
    (1 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <=
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    ((1 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <=
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) |
     1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= 0 |
     skoXC1 <= 0), inference(resolve, [], [refute_0_88, refute_0_81])).

cnf(refute_0_90, plain,
    (skoXC1 *
     (21/25000000 +
      skoXC1 *
      (-33243/250000000 +
       skoXC1 *
       (1113101/500000000 +
        skoXC1 *
        (-7262039/500000000 +
         skoXC1 *
         (92263013/2000000000 +
          skoXC1 *
          (-369620197/5000000000 +
           skoXC1 *
           (519267161/10000000000 +
            skoXC1 *
            (-295729/80000000 +
             skoXC1 *
             (-35703089/5000000000 +
              skoXC1 *
              (-37411/25000000 + skoXC1 * -1413161/10000000000)))))))))) <=
     0 |
     skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= -1680 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    ((5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (12 + skoXC1 * (6 + skoXC1)) < 12 + skoXC1 * (-6 + skoXC1) |
     12 + skoXC1 * (6 + skoXC1) <= 0 |
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_92, plain,
    (~ lgen(0, exp(X_000073), X_000074) | exp(X_000073) <= X_000074),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_93, plain,
    (~ lgen(0,
            (X_000073 ^ 2 + 6 * X_000073 + 12) /
            (X_000073 ^ 2 - 6 * X_000073 + 12), X_000074) | 0 <= X_000073 |
     lgen(0, exp(X_000073), X_000074)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_94, plain,
    (~ lgen(0,
            (X_000073 ^ 2 + 6 * X_000073 + 12) /
            (X_000073 ^ 2 - 6 * X_000073 + 12), X_000074) | 0 <= X_000073 |
     exp(X_000073) <= X_000074),
    inference(resolve, [], [refute_0_93, refute_0_92])).

cnf(refute_0_95, plain,
    (X_000074 <
     (12 + X_000073 * (6 + X_000073)) / (12 + X_000073 * (-6 + X_000073)) |
     0 <= X_000073 | exp(X_000073) <= X_000074),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (12 + skoXC1 * -1 * (6 + skoXC1 * -1)) /
     (12 + skoXC1 * -1 * (-6 + skoXC1 * -1)) | 0 <= skoXC1 * -1 |
     exp(skoXC1 * -1) <=
     5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000))))))),
    inference(subst, [], [refute_0_95])).

cnf(refute_0_97, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (12 + skoXC1 * -1 * (6 + skoXC1 * -1)) /
     (12 + skoXC1 * -1 * (-6 + skoXC1 * -1)) | 0 <= skoXC1 * -1),
    inference(resolve, [], [refute_0_96, refute_0_63])).

cnf(refute_0_98, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    ((5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (12 + skoXC1 * (6 + skoXC1)) < 12 + skoXC1 * (-6 + skoXC1) |
     12 + skoXC1 * (6 + skoXC1) <= 0 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_91, refute_0_98])).

cnf(refute_0_100, plain,
    (3/1250000000 <
     skoXC1 *
     (-9/1250000000 +
      skoXC1 *
      (1187/1250000000 +
       skoXC1 *
       (-159013/10000000000 +
        skoXC1 *
        (129651/1250000000 +
         skoXC1 *
         (32675069/2000000000 +
          skoXC1 *
          (-78077727/10000000000 +
           skoXC1 *
           (2017861/1000000000 +
            skoXC1 *
            (-2409927/5000000000 + skoXC1 * 1413161/10000000000)))))))) |
     skoXC1 * (6 + skoXC1) <= -12 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (3/1250000000 <
     skoXC1 *
     (-9/1250000000 +
      skoXC1 *
      (1187/1250000000 +
       skoXC1 *
       (-159013/10000000000 +
        skoXC1 *
        (129651/1250000000 +
         skoXC1 *
         (32675069/2000000000 +
          skoXC1 *
          (-78077727/10000000000 +
           skoXC1 *
           (2017861/1000000000 +
            skoXC1 *
            (-2409927/5000000000 + skoXC1 * 1413161/10000000000)))))))) |
     -12 < skoXC1 * (6 + skoXC1) | skoXC1 <= 0),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_102, plain,
    (3/1250000000 <
     skoXC1 *
     (-9/1250000000 +
      skoXC1 *
      (1187/1250000000 +
       skoXC1 *
       (-159013/10000000000 +
        skoXC1 *
        (129651/1250000000 +
         skoXC1 *
         (32675069/2000000000 +
          skoXC1 *
          (-78077727/10000000000 +
           skoXC1 *
           (2017861/1000000000 +
            skoXC1 *
            (-2409927/5000000000 + skoXC1 * 1413161/10000000000)))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (-1680 < skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) |
     skoXC1 *
     (21/25000000 +
      skoXC1 *
      (-33243/250000000 +
       skoXC1 *
       (1113101/500000000 +
        skoXC1 *
        (-7262039/500000000 +
         skoXC1 *
         (92263013/2000000000 +
          skoXC1 *
          (-369620197/5000000000 +
           skoXC1 *
           (519267161/10000000000 +
            skoXC1 *
            (-295729/80000000 +
             skoXC1 *
             (-35703089/5000000000 +
              skoXC1 *
              (-37411/25000000 + skoXC1 * -1413161/10000000000)))))))))) <=
     0 | skoXC1 <= 0),
    inference(decision, [], [refute_0_27, refute_0_26, refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 *
     (21/25000000 +
      skoXC1 *
      (-33243/250000000 +
       skoXC1 *
       (1113101/500000000 +
        skoXC1 *
        (-7262039/500000000 +
         skoXC1 *
         (92263013/2000000000 +
          skoXC1 *
          (-369620197/5000000000 +
           skoXC1 *
           (519267161/10000000000 +
            skoXC1 *
            (-295729/80000000 +
             skoXC1 *
             (-35703089/5000000000 +
              skoXC1 *
              (-37411/25000000 + skoXC1 * -1413161/10000000000)))))))))) <=
     0 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_90, refute_0_103])).

cnf(refute_0_105, plain,
    (0 < skoXC1 |
     skoXC1 *
     (21/25000000 +
      skoXC1 *
      (-33243/250000000 +
       skoXC1 *
       (1113101/500000000 +
        skoXC1 *
        (-7262039/500000000 +
         skoXC1 *
         (92263013/2000000000 +
          skoXC1 *
          (-369620197/5000000000 +
           skoXC1 *
           (519267161/10000000000 +
            skoXC1 *
            (-295729/80000000 +
             skoXC1 *
             (-35703089/5000000000 +
              skoXC1 *
              (-37411/25000000 + skoXC1 * -1413161/10000000000)))))))))) <=
     0),
    inference(decision, [], [refute_0_27, refute_0_26, refute_0_102])).

cnf(refute_0_106, plain,
    (skoXC1 *
     (21/25000000 +
      skoXC1 *
      (-33243/250000000 +
       skoXC1 *
       (1113101/500000000 +
        skoXC1 *
        (-7262039/500000000 +
         skoXC1 *
         (92263013/2000000000 +
          skoXC1 *
          (-369620197/5000000000 +
           skoXC1 *
           (519267161/10000000000 +
            skoXC1 *
            (-295729/80000000 +
             skoXC1 *
             (-35703089/5000000000 +
              skoXC1 *
              (-37411/25000000 + skoXC1 * -1413161/10000000000)))))))))) <=
     0), inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain,
    (skoXC1 *
     (-6237/15625000 +
      skoXC1 *
      (102816/1953125 +
       skoXC1 *
       (-110197521/125000000 +
        skoXC1 *
        (718984581/125000000 +
         skoXC1 *
         (-45684502731/2500000000 +
          skoXC1 *
          (73278164347/2500000000 +
           skoXC1 *
           (-206815784617/10000000000 +
            skoXC1 *
            (1033616217/625000000 +
             skoXC1 *
             (26684855389/10000000000 +
              skoXC1 *
              (6143651721/10000000000 +
               skoXC1 *
               (355759199/5000000000 +
                skoXC1 *
                (23026971/5000000000 +
                 skoXC1 * 1413161/10000000000)))))))))))) <=
     2079/15625000 |
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <=
     -665280 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_106, refute_0_27, refute_0_26, refute_0_102])).

cnf(refute_0_108, plain,
    (skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <=
     -665280 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_107, refute_0_80])).

cnf(refute_0_109, plain,
    (-665280 <
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) |
     skoXC1 <= 0),
    inference(decision, [],
              [refute_0_106, refute_0_27, refute_0_26, refute_0_102])).

cnf(refute_0_110, plain, (skoXC1 <= 0),
    inference(resolve, [], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain, (0 < skoXC1 | -3/2 <= skoXC1),
    inference(decision, [],
              [refute_0_110, refute_0_106, refute_0_27, refute_0_26,
               refute_0_102])).

cnf(refute_0_112, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (-2 + skoXC1) / (-2 + skoXC1 * -1) | 0 < skoXC1),
    inference(resolve, [], [refute_0_111, refute_0_70])).

cnf(refute_0_113, plain,
    (5000000001/5000000000 +
     skoXC1 *
     (-1999999999/2000000000 +
      skoXC1 *
      (2499999603/5000000000 +
       skoXC1 *
       (-1666653019/10000000000 +
        skoXC1 *
        (16662939/400000000 +
         skoXC1 *
         (-41506799/5000000000 +
          skoXC1 *
          (664941/500000000 + skoXC1 * -1413161/10000000000)))))) <
     (-2 + skoXC1) / (-2 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_110, refute_0_112])).

cnf(refute_0_114, plain,
    (-2 + skoXC1 <
     (5000000001/5000000000 +
      skoXC1 *
      (-1999999999/2000000000 +
       skoXC1 *
       (2499999603/5000000000 +
        skoXC1 *
        (-1666653019/10000000000 +
         skoXC1 *
         (16662939/400000000 +
          skoXC1 *
          (-41506799/5000000000 +
           skoXC1 *
           (664941/500000000 + skoXC1 * -1413161/10000000000))))))) *
     (-2 + skoXC1 * -1) | 0 <= -2 + skoXC1 * -1),
    inference(resolve, [], [refute_0_0, refute_0_113])).

cnf(refute_0_115, plain,
    (1/2500000000 <
     skoXC1 *
     (-3/2500000000 +
      skoXC1 *
      (1583/10000000000 +
       skoXC1 *
       (-104168323/625000000 +
        skoXC1 *
        (833506069/10000000000 +
         skoXC1 *
         (-250546279/10000000000 +
          skoXC1 *
          (28207979/5000000000 +
           skoXC1 *
           (-5236249/5000000000 + skoXC1 * 1413161/10000000000))))))) |
     skoXC1 <= -2), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (skoXC1 *
     (-3/2500000000 +
      skoXC1 *
      (1583/10000000000 +
       skoXC1 *
       (-104168323/625000000 +
        skoXC1 *
        (833506069/10000000000 +
         skoXC1 *
         (-250546279/10000000000 +
          skoXC1 *
          (28207979/5000000000 +
           skoXC1 *
           (-5236249/5000000000 + skoXC1 * 1413161/10000000000))))))) <=
     1/2500000000 | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_110, refute_0_106, refute_0_27, refute_0_26,
               refute_0_102])).

cnf(refute_0_117, plain, (skoXC1 <= -2),
    inference(resolve, [], [refute_0_116, refute_0_115])).

cnf(refute_0_118, plain, (-2 < skoXC1),
    inference(decision, [],
              [refute_0_110, refute_0_106, refute_0_27, refute_0_26,
               refute_0_102])).

cnf(refute_0_119, plain, ($false),
    inference(resolve, [], [refute_0_117, refute_0_118])).
% SZS output end CNFRefutation for exp-error-analysis-7.tptp
