% SZS output start CNFRefutation for exp-lower-14-sqrt-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf2, axiom,
    (X < 0 | ~ lgen(R, Y, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12)) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf4, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 4 + 20 * X ^ 3 + 180 * X ^ 2 + 840 * X + 1680) /
            (X ^ 4 - 20 * X ^ 3 + 180 * X ^ 2 - 840 * X + 1680)) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf6, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 6 + 42 * X ^ 5 + 840 * X ^ 4 + 10080 * X ^ 3 +
             75600 * X ^ 2 + 332640 * X + 665280) /
            (X ^ 6 - 42 * X ^ 5 + 840 * X ^ 4 - 10080 * X ^ 3 +
             75600 * X ^ 2 - 332640 * X + 665280)) | lgen(R, Y, exp(X)))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_1, axiom,
    (X < 0 | ~ lgen(R, 1/4 * (X ^ 2 + 6 * X + 1) / (X + 1), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_4, axiom,
    (X < 0 |
     ~ lgen(R,
            1/32 *
            (225792840 * X ^ 6 + 64512240 * X ^ 5 + 601080390 * X ^ 8 +
             471435600 * X ^ 7 + 496 * X + 1 + 35960 * X ^ 2 +
             906192 * X ^ 3 + 10518300 * X ^ 4 + X ^ 16 + 906192 * X ^ 13 +
             10518300 * X ^ 12 + 496 * X ^ 15 + 35960 * X ^ 14 +
             471435600 * X ^ 9 + 225792840 * X ^ 10 + 64512240 * X ^ 11) /
            ((X + 1) * (X ^ 2 + 6 * X + 1) *
             (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) *
             (1820 * X ^ 6 + 8008 * X ^ 5 + X ^ 8 + 120 * X ^ 7 + 120 * X +
              1 + 1820 * X ^ 2 + 8008 * X ^ 3 + 12870 * X ^ 4)), Y) |
     lgen(R, sqrt(X), Y))).

fof(exp_lower_bound_problem_14, conjecture,
    (! [X] :
       (13/100000 <= X =>
        (6 + 6 * X + X ^ 3 + X ^ 2 * sqrt(9 + X ^ 2)) / 6 <= exp(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (13/100000 <= X =>
        (6 + 6 * X + X ^ 3 + X ^ 2 * sqrt(9 + X ^ 2)) / 6 <= exp(X))),
    inference(strip, [], [exp_lower_bound_problem_14])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (13/100000 <= X =>
          (6 + 6 * X + X ^ 3 + X ^ 2 * sqrt(9 + X ^ 2)) / 6 <= exp(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (exp(X) < (6 + 6 * X + X ^ 3 + X ^ 2 * sqrt(9 + X ^ 2)) / 6 &
        13/100000 <= X)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (exp(skoXC1) <
     (6 + 6 * skoXC1 + skoXC1 ^ 3 + skoXC1 ^ 2 * sqrt(9 + skoXC1 ^ 2)) /
     6 & 13/100000 <= skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (exp(skoXC1) <
     (6 + 6 * skoXC1 + skoXC1 ^ 3 + skoXC1 ^ 2 * sqrt(9 + skoXC1 ^ 2)) /
     6), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (13/100000 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-96076792028200960 +
     skoXC1 *
     (-96076792028200960 +
      skoXC1 *
      (-512409557461368832/3 +
       skoXC1 *
       (-416332765366059008/3 +
        skoXC1 *
        (-409660765977444352/3 +
         skoXC1 *
         (-278889576678817792/3 +
          skoXC1 *
          (-197441761386168320/3 +
           skoXC1 *
           (-114931368063926272/3 +
            skoXC1 *
            (-21527809123418112 +
             skoXC1 *
             (-10862219843076096 +
              skoXC1 *
              (-5085812237795328 +
               skoXC1 *
               (-2244559021539328 +
                skoXC1 *
                (-899640770625536 +
                 skoXC1 *
                 (-349234651267072 +
                  skoXC1 *
                  (-121711614853120 +
                   skoXC1 *
                   (-41671100727296 +
                    skoXC1 *
                    (-12736077383680 +
                     skoXC1 *
                     (-3846151243776 +
                      skoXC1 *
                      (-1034455569408 +
                       skoXC1 *
                       (-274822047744 +
                        skoXC1 *
                        (-64989789952 +
                         skoXC1 *
                         (-15102540032 +
                          skoXC1 *
                          (-3123732352 +
                           skoXC1 *
                           (-628825600 +
                            skoXC1 *
                            (-112565648 +
                             skoXC1 *
                             (-19329232 +
                              skoXC1 *
                              (-2940920 +
                               skoXC1 *
                               (-420272 +
                                skoXC1 *
                                (-52706 +
                                 skoXC1 *
                                 (-18034/3 +
                                  skoXC1 *
                                  (-1763/3 +
                                   skoXC1 *
                                   (-148/3 +
                                    skoXC1 *
                                    (-10/3 +
                                     skoXC1 *
                                     (-1/6 +
                                      skoXC1 *
                                      -1/
                                      192))))))))))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) *
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     0 <=
     -96076792028200960 +
     skoXC1 *
     (skoXC1 *
      (-122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (-72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (-26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (-6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (-1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (-155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (-15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (-1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (-71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (-3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (-102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (-2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (-34328 +
                               skoXC1 *
                               (skoXC1 *
                                (-290 +
                                 skoXC1 *
                                 (skoXC1 *
                                  -1))))))))))))))))))))))))))))) |
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) <=
     1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000018, exp(X_000017)) | X_000018 <= exp(X_000017)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000018,
            (1 + X_000017 / 3 + 1/2 * (X_000017 / 3) ^ 2 +
             1/6 * (X_000017 / 3) ^ 3 + 1/24 * (X_000017 / 3) ^ 4 +
             1/120 * (X_000017 / 3) ^ 5) ^ 3) |
     lgen(0, X_000018, exp(X_000017))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_3, plain,
    (~ lgen(0, X_000018,
            (1 + X_000017 / 3 + 1/2 * (X_000017 / 3) ^ 2 +
             1/6 * (X_000017 / 3) ^ 3 + 1/24 * (X_000017 / 3) ^ 4 +
             1/120 * (X_000017 / 3) ^ 5) ^ 3) | X_000018 <= exp(X_000017)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (1 +
     X_000017 *
     (1 +
      X_000017 *
      (1/2 +
       X_000017 *
       (1/6 +
        X_000017 *
        (1/24 +
         X_000017 *
         (1/120 +
          X_000017 *
          (121/87480 +
           X_000017 *
           (17/87480 +
            X_000017 *
            (49/2099520 +
             X_000017 *
             (409/170061120 +
              X_000017 *
              (361/1700611200 +
               X_000017 *
               (1/62985600 +
                X_000017 *
                (181/183666009600 +
                 X_000017 *
                 (1/20407334400 +
                  X_000017 *
                  (1/550998028800 +
                   X_000017 * 1/24794911296000)))))))))))))) < X_000018 |
     X_000018 <= exp(X_000017)), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) <
     skoXC1 *
     (skoXC1 *
      (144115188109410304/3 +
       skoXC1 *
       (skoXC1 *
        (192153584190619648/3 +
         skoXC1 *
         (skoXC1 *
          (118761590286712832/3 +
           skoXC1 *
           (skoXC1 *
            (15036709896978432 +
             skoXC1 *
             (skoXC1 *
              (3923103087329280 +
               skoXC1 *
               (skoXC1 *
                (744190977769472 +
                 skoXC1 *
                 (skoXC1 *
                  (105948812935168 +
                   skoXC1 *
                   (skoXC1 *
                    (11517059792896 +
                     skoXC1 *
                     (skoXC1 *
                      (962803120128 +
                       skoXC1 *
                       (skoXC1 *
                        (61829324800 +
                         skoXC1 *
                         (skoXC1 *
                          (3021650944 +
                           skoXC1 *
                           (skoXC1 *
                            (110249984 +
                             skoXC1 *
                             (skoXC1 *
                              (2906592 +
                               skoXC1 *
                               (skoXC1 *
                                (52416 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (1760/3 +
                                   skoXC1 *
                                   (skoXC1 *
                                    (10/3 +
                                     skoXC1 *
                                     (skoXC1 *
                                      1/
                                      192))))))))))))))))))))))))))))))))) |
     96076792028200960 +
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (144115188109410304/3 +
       skoXC1 *
       (skoXC1 *
        (192153584190619648/3 +
         skoXC1 *
         (skoXC1 *
          (118761590286712832/3 +
           skoXC1 *
           (skoXC1 *
            (15036709896978432 +
             skoXC1 *
             (skoXC1 *
              (3923103087329280 +
               skoXC1 *
               (skoXC1 *
                (744190977769472 +
                 skoXC1 *
                 (skoXC1 *
                  (105948812935168 +
                   skoXC1 *
                   (skoXC1 *
                    (11517059792896 +
                     skoXC1 *
                     (skoXC1 *
                      (962803120128 +
                       skoXC1 *
                       (skoXC1 *
                        (61829324800 +
                         skoXC1 *
                         (skoXC1 *
                          (3021650944 +
                           skoXC1 *
                           (skoXC1 *
                            (110249984 +
                             skoXC1 *
                             (skoXC1 *
                              (2906592 +
                               skoXC1 *
                               (skoXC1 *
                                (52416 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (1760/3 +
                                   skoXC1 *
                                   (skoXC1 *
                                    (10/3 +
                                     skoXC1 *
                                     (skoXC1 *
                                      1/
                                      192))))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) <=
     -1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_7, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (288230376218820608 +
      skoXC1 *
      (skoXC1 *
       (384307168381239296 +
        skoXC1 *
        (skoXC1 *
         (237523180573425664 +
          skoXC1 *
          (skoXC1 *
           (90220259381870592 +
            skoXC1 *
            (skoXC1 *
             (23538618523975680 +
              skoXC1 *
              (skoXC1 *
               (4465145866616832 +
                skoXC1 *
                (skoXC1 *
                 (635692877611008 +
                  skoXC1 *
                  (skoXC1 *
                   (69102358757376 +
                    skoXC1 *
                    (skoXC1 *
                     (5776818720768 +
                      skoXC1 *
                      (skoXC1 *
                       (370975948800 +
                        skoXC1 *
                        (skoXC1 *
                         (18129905664 +
                          skoXC1 *
                          (skoXC1 *
                           (661499904 +
                            skoXC1 *
                            (skoXC1 *
                             (17439552 +
                              skoXC1 *
                              (skoXC1 *
                               (314496 +
                                skoXC1 *
                                (skoXC1 *
                                 (3520 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (20 +
                                    skoXC1 *
                                    (skoXC1 *
                                     1/32)))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0 |
     (288230376218820608 +
      skoXC1 *
      (skoXC1 *
       (384307168381239296 +
        skoXC1 *
        (skoXC1 *
         (237523180573425664 +
          skoXC1 *
          (skoXC1 *
           (90220259381870592 +
            skoXC1 *
            (skoXC1 *
             (23538618523975680 +
              skoXC1 *
              (skoXC1 *
               (4465145866616832 +
                skoXC1 *
                (skoXC1 *
                 (635692877611008 +
                  skoXC1 *
                  (skoXC1 *
                   (69102358757376 +
                    skoXC1 *
                    (skoXC1 *
                     (5776818720768 +
                      skoXC1 *
                      (skoXC1 *
                       (370975948800 +
                        skoXC1 *
                        (skoXC1 *
                         (18129905664 +
                          skoXC1 *
                          (skoXC1 *
                           (661499904 +
                            skoXC1 *
                            (skoXC1 *
                             (17439552 +
                              skoXC1 *
                              (skoXC1 *
                               (314496 +
                                skoXC1 *
                                (skoXC1 *
                                 (3520 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (20 +
                                    skoXC1 *
                                    (skoXC1 *
                                     1/32)))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_8, plain,
    (exp(skoXC1) <
     (6 + 6 * skoXC1 + skoXC1 ^ 3 + skoXC1 ^ 2 * sqrt(9 + skoXC1 ^ 2)) /
     6), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     sqrt(9 + skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * 1/6))),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < sqrt(9 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * 1/6) <= 0), inference(split, [], [refute_0_9])).

cnf(refute_0_11, plain, (13/100000 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_12, plain, (0 < skoXC1 * (skoXC1 * 1/6)),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < sqrt(9 + skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_10, refute_0_12])).

cnf(refute_0_14, plain,
    (~ lgen(0, sqrt(X_000180), X_000181) | sqrt(X_000180) <= X_000181),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_15, plain,
    (X_000180 < 0 |
     ~ lgen(0,
            1/32 *
            (225792840 * X_000180 ^ 6 + 64512240 * X_000180 ^ 5 +
             601080390 * X_000180 ^ 8 + 471435600 * X_000180 ^ 7 +
             496 * X_000180 + 1 + 35960 * X_000180 ^ 2 +
             906192 * X_000180 ^ 3 + 10518300 * X_000180 ^ 4 +
             X_000180 ^ 16 + 906192 * X_000180 ^ 13 +
             10518300 * X_000180 ^ 12 + 496 * X_000180 ^ 15 +
             35960 * X_000180 ^ 14 + 471435600 * X_000180 ^ 9 +
             225792840 * X_000180 ^ 10 + 64512240 * X_000180 ^ 11) /
            ((X_000180 + 1) * (X_000180 ^ 2 + 6 * X_000180 + 1) *
             (X_000180 ^ 4 + 28 * X_000180 ^ 3 + 70 * X_000180 ^ 2 +
              28 * X_000180 + 1) *
             (1820 * X_000180 ^ 6 + 8008 * X_000180 ^ 5 + X_000180 ^ 8 +
              120 * X_000180 ^ 7 + 120 * X_000180 + 1 +
              1820 * X_000180 ^ 2 + 8008 * X_000180 ^ 3 +
              12870 * X_000180 ^ 4)), X_000181) |
     lgen(0, sqrt(X_000180), X_000181)),
    inference(subst, [], [sqrt_upper_bound_4])).

cnf(refute_0_16, plain,
    (X_000180 < 0 |
     ~ lgen(0,
            1/32 *
            (225792840 * X_000180 ^ 6 + 64512240 * X_000180 ^ 5 +
             601080390 * X_000180 ^ 8 + 471435600 * X_000180 ^ 7 +
             496 * X_000180 + 1 + 35960 * X_000180 ^ 2 +
             906192 * X_000180 ^ 3 + 10518300 * X_000180 ^ 4 +
             X_000180 ^ 16 + 906192 * X_000180 ^ 13 +
             10518300 * X_000180 ^ 12 + 496 * X_000180 ^ 15 +
             35960 * X_000180 ^ 14 + 471435600 * X_000180 ^ 9 +
             225792840 * X_000180 ^ 10 + 64512240 * X_000180 ^ 11) /
            ((X_000180 + 1) * (X_000180 ^ 2 + 6 * X_000180 + 1) *
             (X_000180 ^ 4 + 28 * X_000180 ^ 3 + 70 * X_000180 ^ 2 +
              28 * X_000180 + 1) *
             (1820 * X_000180 ^ 6 + 8008 * X_000180 ^ 5 + X_000180 ^ 8 +
              120 * X_000180 ^ 7 + 120 * X_000180 + 1 +
              1820 * X_000180 ^ 2 + 8008 * X_000180 ^ 3 +
              12870 * X_000180 ^ 4)), X_000181) |
     sqrt(X_000180) <= X_000181),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    (X_000180 < 0 |
     X_000181 <
     (1/32 +
      X_000180 *
      (31/2 +
       X_000180 *
       (4495/4 +
        X_000180 *
        (56637/2 +
         X_000180 *
         (2629575/8 +
          X_000180 *
          (4032015/2 +
           X_000180 *
           (28224105/4 +
            X_000180 *
            (29464725/2 +
             X_000180 *
             (300540195/16 +
              X_000180 *
              (29464725/2 +
               X_000180 *
               (28224105/4 +
                X_000180 *
                (4032015/2 +
                 X_000180 *
                 (2629575/8 +
                  X_000180 *
                  (56637/2 +
                   X_000180 *
                   (4495/4 +
                    X_000180 * (31/2 + X_000180 * 1/32)))))))))))))))) /
     (1 +
      X_000180 *
      (155 +
       X_000180 *
       (6293 +
        X_000180 *
        (105183 +
         X_000180 *
         (876525 +
          X_000180 *
          (4032015 +
           X_000180 *
           (10855425 +
            X_000180 *
            (17678835 +
             X_000180 *
             (17678835 +
              X_000180 *
              (10855425 +
               X_000180 *
               (4032015 +
                X_000180 *
                (876525 +
                 X_000180 *
                 (105183 +
                  X_000180 *
                  (6293 + X_000180 * (155 + X_000180))))))))))))))) |
     sqrt(X_000180) <= X_000181),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/32 +
      (9 + skoXC1 * skoXC1) *
      (31/2 +
       (9 + skoXC1 * skoXC1) *
       (4495/4 +
        (9 + skoXC1 * skoXC1) *
        (56637/2 +
         (9 + skoXC1 * skoXC1) *
         (2629575/8 +
          (9 + skoXC1 * skoXC1) *
          (4032015/2 +
           (9 + skoXC1 * skoXC1) *
           (28224105/4 +
            (9 + skoXC1 * skoXC1) *
            (29464725/2 +
             (9 + skoXC1 * skoXC1) *
             (300540195/16 +
              (9 + skoXC1 * skoXC1) *
              (29464725/2 +
               (9 + skoXC1 * skoXC1) *
               (28224105/4 +
                (9 + skoXC1 * skoXC1) *
                (4032015/2 +
                 (9 + skoXC1 * skoXC1) *
                 (2629575/8 +
                  (9 + skoXC1 * skoXC1) *
                  (56637/2 +
                   (9 + skoXC1 * skoXC1) *
                   (4495/4 +
                    (9 + skoXC1 * skoXC1) *
                    (31/2 + (9 + skoXC1 * skoXC1) * 1/32)))))))))))))))) /
     (1 +
      (9 + skoXC1 * skoXC1) *
      (155 +
       (9 + skoXC1 * skoXC1) *
       (6293 +
        (9 + skoXC1 * skoXC1) *
        (105183 +
         (9 + skoXC1 * skoXC1) *
         (876525 +
          (9 + skoXC1 * skoXC1) *
          (4032015 +
           (9 + skoXC1 * skoXC1) *
           (10855425 +
            (9 + skoXC1 * skoXC1) *
            (17678835 +
             (9 + skoXC1 * skoXC1) *
             (17678835 +
              (9 + skoXC1 * skoXC1) *
              (10855425 +
               (9 + skoXC1 * skoXC1) *
               (4032015 +
                (9 + skoXC1 * skoXC1) *
                (876525 +
                 (9 + skoXC1 * skoXC1) *
                 (105183 +
                  (9 + skoXC1 * skoXC1) *
                  (6293 +
                   (9 + skoXC1 * skoXC1) *
                   (155 + (9 + skoXC1 * skoXC1)))))))))))))))) |
     sqrt(9 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))), inference(subst, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/32 +
      (9 + skoXC1 * skoXC1) *
      (31/2 +
       (9 + skoXC1 * skoXC1) *
       (4495/4 +
        (9 + skoXC1 * skoXC1) *
        (56637/2 +
         (9 + skoXC1 * skoXC1) *
         (2629575/8 +
          (9 + skoXC1 * skoXC1) *
          (4032015/2 +
           (9 + skoXC1 * skoXC1) *
           (28224105/4 +
            (9 + skoXC1 * skoXC1) *
            (29464725/2 +
             (9 + skoXC1 * skoXC1) *
             (300540195/16 +
              (9 + skoXC1 * skoXC1) *
              (29464725/2 +
               (9 + skoXC1 * skoXC1) *
               (28224105/4 +
                (9 + skoXC1 * skoXC1) *
                (4032015/2 +
                 (9 + skoXC1 * skoXC1) *
                 (2629575/8 +
                  (9 + skoXC1 * skoXC1) *
                  (56637/2 +
                   (9 + skoXC1 * skoXC1) *
                   (4495/4 +
                    (9 + skoXC1 * skoXC1) *
                    (31/2 + (9 + skoXC1 * skoXC1) * 1/32)))))))))))))))) /
     (1 +
      (9 + skoXC1 * skoXC1) *
      (155 +
       (9 + skoXC1 * skoXC1) *
       (6293 +
        (9 + skoXC1 * skoXC1) *
        (105183 +
         (9 + skoXC1 * skoXC1) *
         (876525 +
          (9 + skoXC1 * skoXC1) *
          (4032015 +
           (9 + skoXC1 * skoXC1) *
           (10855425 +
            (9 + skoXC1 * skoXC1) *
            (17678835 +
             (9 + skoXC1 * skoXC1) *
             (17678835 +
              (9 + skoXC1 * skoXC1) *
              (10855425 +
               (9 + skoXC1 * skoXC1) *
               (4032015 +
                (9 + skoXC1 * skoXC1) *
                (876525 +
                 (9 + skoXC1 * skoXC1) *
                 (105183 +
                  (9 + skoXC1 * skoXC1) *
                  (6293 +
                   (9 + skoXC1 * skoXC1) *
                   (155 + (9 + skoXC1 * skoXC1))))))))))))))))),
    inference(resolve, [], [refute_0_18, refute_0_13])).

cnf(refute_0_20, plain,
    (9 < skoXC1 * (skoXC1 * -1) |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (288230376218820608 +
      skoXC1 *
      (skoXC1 *
       (384307168381239296 +
        skoXC1 *
        (skoXC1 *
         (237523180573425664 +
          skoXC1 *
          (skoXC1 *
           (90220259381870592 +
            skoXC1 *
            (skoXC1 *
             (23538618523975680 +
              skoXC1 *
              (skoXC1 *
               (4465145866616832 +
                skoXC1 *
                (skoXC1 *
                 (635692877611008 +
                  skoXC1 *
                  (skoXC1 *
                   (69102358757376 +
                    skoXC1 *
                    (skoXC1 *
                     (5776818720768 +
                      skoXC1 *
                      (skoXC1 *
                       (370975948800 +
                        skoXC1 *
                        (skoXC1 *
                         (18129905664 +
                          skoXC1 *
                          (skoXC1 *
                           (661499904 +
                            skoXC1 *
                            (skoXC1 *
                             (17439552 +
                              skoXC1 *
                              (skoXC1 *
                               (314496 +
                                skoXC1 *
                                (skoXC1 *
                                 (3520 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (20 +
                                    skoXC1 *
                                    (skoXC1 *
                                     1/32)))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1)))))))))))))))))))))))))))))),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain, (skoXC1 * (skoXC1 * -1) <= 9),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_22, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (288230376218820608 +
      skoXC1 *
      (skoXC1 *
       (384307168381239296 +
        skoXC1 *
        (skoXC1 *
         (237523180573425664 +
          skoXC1 *
          (skoXC1 *
           (90220259381870592 +
            skoXC1 *
            (skoXC1 *
             (23538618523975680 +
              skoXC1 *
              (skoXC1 *
               (4465145866616832 +
                skoXC1 *
                (skoXC1 *
                 (635692877611008 +
                  skoXC1 *
                  (skoXC1 *
                   (69102358757376 +
                    skoXC1 *
                    (skoXC1 *
                     (5776818720768 +
                      skoXC1 *
                      (skoXC1 *
                       (370975948800 +
                        skoXC1 *
                        (skoXC1 *
                         (18129905664 +
                          skoXC1 *
                          (skoXC1 *
                           (661499904 +
                            skoXC1 *
                            (skoXC1 *
                             (17439552 +
                              skoXC1 *
                              (skoXC1 *
                               (314496 +
                                skoXC1 *
                                (skoXC1 *
                                 (3520 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (20 +
                                    skoXC1 *
                                    (skoXC1 *
                                     1/32)))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1)))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_21, refute_0_20])).

cnf(refute_0_23, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (288230376218820608 +
      skoXC1 *
      (skoXC1 *
       (384307168381239296 +
        skoXC1 *
        (skoXC1 *
         (237523180573425664 +
          skoXC1 *
          (skoXC1 *
           (90220259381870592 +
            skoXC1 *
            (skoXC1 *
             (23538618523975680 +
              skoXC1 *
              (skoXC1 *
               (4465145866616832 +
                skoXC1 *
                (skoXC1 *
                 (635692877611008 +
                  skoXC1 *
                  (skoXC1 *
                   (69102358757376 +
                    skoXC1 *
                    (skoXC1 *
                     (5776818720768 +
                      skoXC1 *
                      (skoXC1 *
                       (370975948800 +
                        skoXC1 *
                        (skoXC1 *
                         (18129905664 +
                          skoXC1 *
                          (skoXC1 *
                           (661499904 +
                            skoXC1 *
                            (skoXC1 *
                             (17439552 +
                              skoXC1 *
                              (skoXC1 *
                               (314496 +
                                skoXC1 *
                                (skoXC1 *
                                 (3520 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (20 +
                                    skoXC1 *
                                    (skoXC1 *
                                     1/32)))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(resolve, [], [refute_0_7, refute_0_22])).

cnf(refute_0_24, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (144115188109410304/3 +
       skoXC1 *
       (skoXC1 *
        (192153584190619648/3 +
         skoXC1 *
         (skoXC1 *
          (118761590286712832/3 +
           skoXC1 *
           (skoXC1 *
            (15036709896978432 +
             skoXC1 *
             (skoXC1 *
              (3923103087329280 +
               skoXC1 *
               (skoXC1 *
                (744190977769472 +
                 skoXC1 *
                 (skoXC1 *
                  (105948812935168 +
                   skoXC1 *
                   (skoXC1 *
                    (11517059792896 +
                     skoXC1 *
                     (skoXC1 *
                      (962803120128 +
                       skoXC1 *
                       (skoXC1 *
                        (61829324800 +
                         skoXC1 *
                         (skoXC1 *
                          (3021650944 +
                           skoXC1 *
                           (skoXC1 *
                            (110249984 +
                             skoXC1 *
                             (skoXC1 *
                              (2906592 +
                               skoXC1 *
                               (skoXC1 *
                                (52416 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (1760/3 +
                                   skoXC1 *
                                   (skoXC1 *
                                    (10/3 +
                                     skoXC1 *
                                     (skoXC1 *
                                      1/
                                      192))))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (144115188109410304/3 +
       skoXC1 *
       (skoXC1 *
        (192153584190619648/3 +
         skoXC1 *
         (skoXC1 *
          (118761590286712832/3 +
           skoXC1 *
           (skoXC1 *
            (15036709896978432 +
             skoXC1 *
             (skoXC1 *
              (3923103087329280 +
               skoXC1 *
               (skoXC1 *
                (744190977769472 +
                 skoXC1 *
                 (skoXC1 *
                  (105948812935168 +
                   skoXC1 *
                   (skoXC1 *
                    (11517059792896 +
                     skoXC1 *
                     (skoXC1 *
                      (962803120128 +
                       skoXC1 *
                       (skoXC1 *
                        (61829324800 +
                         skoXC1 *
                         (skoXC1 *
                          (3021650944 +
                           skoXC1 *
                           (skoXC1 *
                            (110249984 +
                             skoXC1 *
                             (skoXC1 *
                              (2906592 +
                               skoXC1 *
                               (skoXC1 *
                                (52416 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (1760/3 +
                                   skoXC1 *
                                   (skoXC1 *
                                    (10/3 +
                                     skoXC1 *
                                     (skoXC1 *
                                      1/
                                      192))))))))))))))))))))))))))))))))) /
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1)))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_24, refute_0_12])).

cnf(refute_0_26, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (290 +
                                  skoXC1 *
                                  skoXC1))))))))))))))))))))))))))))) <
     skoXC1 *
     (skoXC1 *
      (144115188109410304/3 +
       skoXC1 *
       (skoXC1 *
        (192153584190619648/3 +
         skoXC1 *
         (skoXC1 *
          (118761590286712832/3 +
           skoXC1 *
           (skoXC1 *
            (15036709896978432 +
             skoXC1 *
             (skoXC1 *
              (3923103087329280 +
               skoXC1 *
               (skoXC1 *
                (744190977769472 +
                 skoXC1 *
                 (skoXC1 *
                  (105948812935168 +
                   skoXC1 *
                   (skoXC1 *
                    (11517059792896 +
                     skoXC1 *
                     (skoXC1 *
                      (962803120128 +
                       skoXC1 *
                       (skoXC1 *
                        (61829324800 +
                         skoXC1 *
                         (skoXC1 *
                          (3021650944 +
                           skoXC1 *
                           (skoXC1 *
                            (110249984 +
                             skoXC1 *
                             (skoXC1 *
                              (2906592 +
                               skoXC1 *
                               (skoXC1 *
                                (52416 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (1760/3 +
                                   skoXC1 *
                                   (skoXC1 *
                                    (10/3 +
                                     skoXC1 *
                                     (skoXC1 *
                                      1/
                                      192))))))))))))))))))))))))))))))))) |
     96076792028200960 +
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_25])).

cnf(refute_0_27, plain,
    (-96076792028200960 +
     skoXC1 *
     (-96076792028200960 +
      skoXC1 *
      (-512409557461368832/3 +
       skoXC1 *
       (-416332765366059008/3 +
        skoXC1 *
        (-409660765977444352/3 +
         skoXC1 *
         (-278889576678817792/3 +
          skoXC1 *
          (-197441761386168320/3 +
           skoXC1 *
           (-114931368063926272/3 +
            skoXC1 *
            (-21527809123418112 +
             skoXC1 *
             (-10862219843076096 +
              skoXC1 *
              (-5085812237795328 +
               skoXC1 *
               (-2244559021539328 +
                skoXC1 *
                (-899640770625536 +
                 skoXC1 *
                 (-349234651267072 +
                  skoXC1 *
                  (-121711614853120 +
                   skoXC1 *
                   (-41671100727296 +
                    skoXC1 *
                    (-12736077383680 +
                     skoXC1 *
                     (-3846151243776 +
                      skoXC1 *
                      (-1034455569408 +
                       skoXC1 *
                       (-274822047744 +
                        skoXC1 *
                        (-64989789952 +
                         skoXC1 *
                         (-15102540032 +
                          skoXC1 *
                          (-3123732352 +
                           skoXC1 *
                           (-628825600 +
                            skoXC1 *
                            (-112565648 +
                             skoXC1 *
                             (-19329232 +
                              skoXC1 *
                              (-2940920 +
                               skoXC1 *
                               (-420272 +
                                skoXC1 *
                                (-52706 +
                                 skoXC1 *
                                 (-18034/3 +
                                  skoXC1 *
                                  (-1763/3 +
                                   skoXC1 *
                                   (-148/3 +
                                    skoXC1 *
                                    (-10/3 +
                                     skoXC1 *
                                     (-1/6 +
                                      skoXC1 *
                                      -1/
                                      192))))))))))))))))))))))))))))))))) <
     exp(skoXC1) *
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <=
     -96076792028200960), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (-6652800 +
     skoXC1 *
     (-3326400 +
      skoXC1 *
      (-1864800 +
       skoXC1 *
       (-211680 +
        skoXC1 *
        (-189840 +
         skoXC1 *
         (79380 +
          skoXC1 *
          (-26190 +
           skoXC1 *
           (9370 +
            skoXC1 *
            (-6614/3 +
             skoXC1 *
             (958/3 +
              skoXC1 * (-29 + skoXC1 * (19/12 + skoXC1 * -1/24))))))))))) <
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) *
     (-10 + skoXC1 * (skoXC1 * -1)) | 0 <= -10 + skoXC1 * (skoXC1 * -1) |
     (-6652800 +
      skoXC1 *
      (-3326400 +
       skoXC1 *
       (-1864800 +
        skoXC1 *
        (-211680 +
         skoXC1 *
         (-189840 +
          skoXC1 *
          (79380 +
           skoXC1 *
           (-26190 +
            skoXC1 *
            (9370 +
             skoXC1 *
             (-6614/3 +
              skoXC1 *
              (958/3 +
               skoXC1 *
               (-29 + skoXC1 * (19/12 + skoXC1 * -1/24)))))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <=
     665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_29, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0 |
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <=
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_30, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (10 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (17/3 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/24))))) |
     10 + skoXC1 * skoXC1 <= 0 |
     skoXC1 *
     (skoXC1 *
      (17/3 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/24))))) /
     (10 + skoXC1 * skoXC1) <=
     -1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_31, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (34 + skoXC1 * (skoXC1 * (6 + skoXC1 * (skoXC1 * 1/4)))) /
     (10 + skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * 1/6)) |
     skoXC1 * (skoXC1 * 1/6) <= 0 |
     (34 + skoXC1 * (skoXC1 * (6 + skoXC1 * (skoXC1 * 1/4)))) /
     (10 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_32, plain,
    (~ lgen(0, sqrt(X_000062), X_000063) | sqrt(X_000062) <= X_000063),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_33, plain,
    (X_000062 < 0 |
     ~ lgen(0, 1/4 * (X_000062 ^ 2 + 6 * X_000062 + 1) / (X_000062 + 1),
            X_000063) | lgen(0, sqrt(X_000062), X_000063)),
    inference(subst, [], [sqrt_upper_bound_1])).

cnf(refute_0_34, plain,
    (X_000062 < 0 |
     ~ lgen(0, 1/4 * (X_000062 ^ 2 + 6 * X_000062 + 1) / (X_000062 + 1),
            X_000063) | sqrt(X_000062) <= X_000063),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (X_000062 < 0 |
     X_000063 <
     (1/4 + X_000062 * (3/2 + X_000062 * 1/4)) / (1 + X_000062) |
     sqrt(X_000062) <= X_000063),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/4 + (9 + skoXC1 * skoXC1) * (3/2 + (9 + skoXC1 * skoXC1) * 1/4)) /
     (1 + (9 + skoXC1 * skoXC1)) |
     sqrt(9 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))), inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/4 + (9 + skoXC1 * skoXC1) * (3/2 + (9 + skoXC1 * skoXC1) * 1/4)) /
     (1 + (9 + skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_36, refute_0_13])).

cnf(refute_0_38, plain,
    (9 < skoXC1 * (skoXC1 * -1) |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (34 + skoXC1 * (skoXC1 * (6 + skoXC1 * (skoXC1 * 1/4)))) /
     (10 + skoXC1 * skoXC1)), inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (34 + skoXC1 * (skoXC1 * (6 + skoXC1 * (skoXC1 * 1/4)))) /
     (10 + skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_21, refute_0_38])).

cnf(refute_0_40, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (34 + skoXC1 * (skoXC1 * (6 + skoXC1 * (skoXC1 * 1/4)))) /
     (10 + skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * 1/6)) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(resolve, [], [refute_0_31, refute_0_39])).

cnf(refute_0_41, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (17/3 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/24))))) /
     (10 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (17/3 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/24))))) /
     (10 + skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_41, refute_0_12])).

cnf(refute_0_43, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (10 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (17/3 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * 1/24))))) |
     10 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_30, refute_0_42])).

cnf(refute_0_44, plain,
    (-10 +
     skoXC1 *
     (-10 +
      skoXC1 *
      (-20/3 +
       skoXC1 *
       (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24))))) <
     exp(skoXC1) * (-10 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * skoXC1 <= -10), inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain, (-10 < skoXC1 * skoXC1),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_46, plain,
    (-10 +
     skoXC1 *
     (-10 +
      skoXC1 *
      (-20/3 +
       skoXC1 *
       (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24))))) <
     exp(skoXC1) * (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    (exp(skoXC1) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) | 0 <= -10 + skoXC1 * (skoXC1 * -1)),
    inference(split, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (exp(skoXC1) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) | skoXC1 * skoXC1 <= -10),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (exp(skoXC1) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_48, refute_0_45])).

cnf(refute_0_50, plain,
    (~ lgen(0, X_000126, exp(X_000125)) | X_000126 <= exp(X_000125)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_51, plain,
    (X_000125 < 0 |
     ~ lgen(0, X_000126,
            (X_000125 ^ 6 + 42 * X_000125 ^ 5 + 840 * X_000125 ^ 4 +
             10080 * X_000125 ^ 3 + 75600 * X_000125 ^ 2 +
             332640 * X_000125 + 665280) /
            (X_000125 ^ 6 - 42 * X_000125 ^ 5 + 840 * X_000125 ^ 4 -
             10080 * X_000125 ^ 3 + 75600 * X_000125 ^ 2 -
             332640 * X_000125 + 665280)) |
     lgen(0, X_000126, exp(X_000125))),
    inference(subst, [], [exp_lower_bound_cf6])).

cnf(refute_0_52, plain,
    (X_000125 < 0 |
     ~ lgen(0, X_000126,
            (X_000125 ^ 6 + 42 * X_000125 ^ 5 + 840 * X_000125 ^ 4 +
             10080 * X_000125 ^ 3 + 75600 * X_000125 ^ 2 +
             332640 * X_000125 + 665280) /
            (X_000125 ^ 6 - 42 * X_000125 ^ 5 + 840 * X_000125 ^ 4 -
             10080 * X_000125 ^ 3 + 75600 * X_000125 ^ 2 -
             332640 * X_000125 + 665280)) | X_000126 <= exp(X_000125)),
    inference(resolve, [], [refute_0_51, refute_0_50])).

cnf(refute_0_53, plain,
    ((665280 +
      X_000125 *
      (332640 +
       X_000125 *
       (75600 +
        X_000125 *
        (10080 + X_000125 * (840 + X_000125 * (42 + X_000125)))))) /
     (665280 +
      X_000125 *
      (-332640 +
       X_000125 *
       (75600 +
        X_000125 *
        (-10080 + X_000125 * (840 + X_000125 * (-42 + X_000125)))))) <
     X_000126 | X_000125 < 0 | X_000126 <= exp(X_000125)),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) |
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_54, refute_0_49])).

cnf(refute_0_56, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_57, plain,
    ((665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_56, refute_0_55])).

cnf(refute_0_58, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_29, refute_0_57])).

cnf(refute_0_59, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-6652800 +
      skoXC1 *
      (-3326400 +
       skoXC1 *
       (-1864800 +
        skoXC1 *
        (-211680 +
         skoXC1 *
         (-189840 +
          skoXC1 *
          (79380 +
           skoXC1 *
           (-26190 +
            skoXC1 *
            (9370 +
             skoXC1 *
             (-6614/3 +
              skoXC1 *
              (958/3 +
               skoXC1 *
               (-29 + skoXC1 * (19/12 + skoXC1 * -1/24)))))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     -665280), inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) <=
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_61, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (5 + skoXC1 * (skoXC1 * 1/2)) * (skoXC1 * (skoXC1 * 1/6)) |
     skoXC1 * (skoXC1 * 1/6) <= 0 |
     5 + skoXC1 * (skoXC1 * 1/2) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_62, plain,
    (~ lgen(0, sqrt(X_000055), X_000056) | sqrt(X_000055) <= X_000056),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_63, plain,
    (X_000055 < 0 | ~ lgen(0, (X_000055 + 1) / 2, X_000056) |
     lgen(0, sqrt(X_000055), X_000056)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_64, plain,
    (X_000055 < 0 | ~ lgen(0, (X_000055 + 1) / 2, X_000056) |
     sqrt(X_000055) <= X_000056),
    inference(resolve, [], [refute_0_63, refute_0_62])).

cnf(refute_0_65, plain,
    (X_000055 < 0 | X_000056 < 1/2 + X_000055 * 1/2 |
     sqrt(X_000055) <= X_000056),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < 1/2 + (9 + skoXC1 * skoXC1) * 1/2 |
     sqrt(9 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))), inference(subst, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < 1/2 + (9 + skoXC1 * skoXC1) * 1/2),
    inference(resolve, [], [refute_0_66, refute_0_13])).

cnf(refute_0_68, plain,
    (9 < skoXC1 * (skoXC1 * -1) |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < 5 + skoXC1 * (skoXC1 * 1/2)),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) < 5 + skoXC1 * (skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_21, refute_0_68])).

cnf(refute_0_70, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (5 + skoXC1 * (skoXC1 * 1/2)) * (skoXC1 * (skoXC1 * 1/6)) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(resolve, [], [refute_0_61, refute_0_69])).

cnf(refute_0_71, plain,
    (exp(skoXC1) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (exp(skoXC1) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_71, refute_0_12])).

cnf(refute_0_73, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) |
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_53])).

cnf(refute_0_74, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    ((665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_56, refute_0_74])).

cnf(refute_0_76, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_60, refute_0_75])).

cnf(refute_0_77, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (221760 +
       skoXC1 *
       (-110880 +
        skoXC1 *
        (52920 +
         skoXC1 *
         (-22764 +
          skoXC1 *
          (5278 +
           skoXC1 *
           (-734 +
            skoXC1 * (383/6 + skoXC1 * (-10/3 + skoXC1 * 1/12))))))))) |
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     -665280), inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (-16800 +
     skoXC1 *
     (-8400 +
      skoXC1 *
      (-4600 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-450 +
         skoXC1 *
         (610/3 +
          skoXC1 *
          (-190/3 +
           skoXC1 *
           (67/3 +
            skoXC1 * (-31/6 + skoXC1 * (2/3 + skoXC1 * -1/24))))))))) <
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) *
     (-10 + skoXC1 * (skoXC1 * -1)) | 0 <= -10 + skoXC1 * (skoXC1 * -1) |
     (-16800 +
      skoXC1 *
      (-8400 +
       skoXC1 *
       (-4600 +
        skoXC1 *
        (-480 +
         skoXC1 *
         (-450 +
          skoXC1 *
          (610/3 +
           skoXC1 *
           (-190/3 +
            skoXC1 *
            (67/3 +
             skoXC1 * (-31/6 + skoXC1 * (2/3 + skoXC1 * -1/24)))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <=
     1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_79, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <=
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_80, plain,
    (~ lgen(0, X_000105, exp(X_000104)) | X_000105 <= exp(X_000104)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_81, plain,
    (X_000104 < 0 |
     ~ lgen(0, X_000105,
            (X_000104 ^ 4 + 20 * X_000104 ^ 3 + 180 * X_000104 ^ 2 +
             840 * X_000104 + 1680) /
            (X_000104 ^ 4 - 20 * X_000104 ^ 3 + 180 * X_000104 ^ 2 -
             840 * X_000104 + 1680)) | lgen(0, X_000105, exp(X_000104))),
    inference(subst, [], [exp_lower_bound_cf4])).

cnf(refute_0_82, plain,
    (X_000104 < 0 |
     ~ lgen(0, X_000105,
            (X_000104 ^ 4 + 20 * X_000104 ^ 3 + 180 * X_000104 ^ 2 +
             840 * X_000104 + 1680) /
            (X_000104 ^ 4 - 20 * X_000104 ^ 3 + 180 * X_000104 ^ 2 -
             840 * X_000104 + 1680)) | X_000105 <= exp(X_000104)),
    inference(resolve, [], [refute_0_81, refute_0_80])).

cnf(refute_0_83, plain,
    ((1680 +
      X_000104 * (840 + X_000104 * (180 + X_000104 * (20 + X_000104)))) /
     (1680 +
      X_000104 * (-840 + X_000104 * (180 + X_000104 * (-20 + X_000104)))) <
     X_000105 | X_000104 < 0 | X_000105 <= exp(X_000104)),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) |
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) <= exp(skoXC1)),
    inference(subst, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_84, refute_0_49])).

cnf(refute_0_86, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_56, refute_0_85])).

cnf(refute_0_87, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-10 +
      skoXC1 *
      (-10 +
       skoXC1 *
       (-20/3 +
        skoXC1 *
        (-8/3 + skoXC1 * (-1 + skoXC1 * (-1/6 + skoXC1 * -1/24)))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0), inference(resolve, [], [refute_0_79, refute_0_86])).

cnf(refute_0_88, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-16800 +
      skoXC1 *
      (-8400 +
       skoXC1 *
       (-4600 +
        skoXC1 *
        (-480 +
         skoXC1 *
         (-450 +
          skoXC1 *
          (610/3 +
           skoXC1 *
           (-190/3 +
            skoXC1 *
            (67/3 +
             skoXC1 * (-31/6 + skoXC1 * (2/3 + skoXC1 * -1/24)))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) <=
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_90, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) |
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_83])).

cnf(refute_0_91, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_90, refute_0_72])).

cnf(refute_0_92, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_56, refute_0_91])).

cnf(refute_0_93, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (1 + skoXC1 * (1 + skoXC1 * (5/6 + skoXC1 * (1/6 + skoXC1 * 1/12)))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0), inference(resolve, [], [refute_0_89, refute_0_92])).

cnf(refute_0_94, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (560 +
       skoXC1 *
       (-280 +
        skoXC1 *
        (130 +
         skoXC1 *
         (-167/3 + skoXC1 * (25/2 + skoXC1 * (-3/2 + skoXC1 * 1/12))))))) |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (-360300 +
     skoXC1 *
     (-180150 +
      skoXC1 *
      (-1258193/16 +
       skoXC1 *
       (-777791/32 +
        skoXC1 *
        (-137233/192 +
         skoXC1 *
         (-20017/24 +
          skoXC1 *
          (-1531/48 +
           skoXC1 *
           (-6959/48 +
            skoXC1 *
            (-4331/96 +
             skoXC1 *
             (-265/24 +
              skoXC1 *
              (-19/12 + skoXC1 * (-13/96 + skoXC1 * -1/192))))))))))) <
     (12 + skoXC1 * (6 + skoXC1)) *
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -30025 +
     skoXC1 *
     (skoXC1 *
      (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))) |
     (-360300 +
      skoXC1 *
      (-180150 +
       skoXC1 *
       (-1258193/16 +
        skoXC1 *
        (-777791/32 +
         skoXC1 *
         (-137233/192 +
          skoXC1 *
          (-20017/24 +
           skoXC1 *
           (-1531/48 +
            skoXC1 *
            (-6959/48 +
             skoXC1 *
             (-4331/96 +
              skoXC1 *
              (-265/24 +
               skoXC1 *
               (-19/12 + skoXC1 * (-13/96 + skoXC1 * -1/192)))))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_96, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) *
     (12 + skoXC1 * (-6 + skoXC1)) | 12 + skoXC1 * (-6 + skoXC1) <= 0 |
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) <=
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_97, plain,
    (~ lgen(0, X_000077, exp(X_000076)) | X_000077 <= exp(X_000076)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_98, plain,
    (X_000076 < 0 |
     ~ lgen(0, X_000077,
            (X_000076 ^ 2 + 6 * X_000076 + 12) /
            (X_000076 ^ 2 - 6 * X_000076 + 12)) |
     lgen(0, X_000077, exp(X_000076))),
    inference(subst, [], [exp_lower_bound_cf2])).

cnf(refute_0_99, plain,
    (X_000076 < 0 |
     ~ lgen(0, X_000077,
            (X_000076 ^ 2 + 6 * X_000076 + 12) /
            (X_000076 ^ 2 - 6 * X_000076 + 12)) |
     X_000077 <= exp(X_000076)),
    inference(resolve, [], [refute_0_98, refute_0_97])).

cnf(refute_0_100, plain,
    ((12 + X_000076 * (6 + X_000076)) / (12 + X_000076 * (-6 + X_000076)) <
     X_000077 | X_000076 < 0 | X_000077 <= exp(X_000076)),
    inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_100])).

cnf(refute_0_102, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) <
     skoXC1 *
     (skoXC1 *
      (2882401/192 +
       skoXC1 *
       (skoXC1 *
        (137257/48 +
         skoXC1 *
         (skoXC1 *
          (15251/96 +
           skoXC1 * (skoXC1 * (121/48 + skoXC1 * (skoXC1 * 1/192))))))))) |
     30025 +
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (2882401/192 +
       skoXC1 *
       (skoXC1 *
        (137257/48 +
         skoXC1 *
         (skoXC1 *
          (15251/96 +
           skoXC1 * (skoXC1 * (121/48 + skoXC1 * (skoXC1 * 1/192))))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) <=
     -1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_103, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (2882401/32 +
      skoXC1 *
      (skoXC1 *
       (137257/8 +
        skoXC1 *
        (skoXC1 *
         (15251/16 +
          skoXC1 * (skoXC1 * (121/8 + skoXC1 * (skoXC1 * 1/32)))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0 |
     (2882401/32 +
      skoXC1 *
      (skoXC1 *
       (137257/8 +
        skoXC1 *
        (skoXC1 *
         (15251/16 +
          skoXC1 * (skoXC1 * (121/8 + skoXC1 * (skoXC1 * 1/32)))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_104, plain,
    (~ lgen(0, sqrt(X_000097), X_000098) | sqrt(X_000097) <= X_000098),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_105, plain,
    (X_000097 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000097 + 17920 * X_000097 ^ 2 +
             448 * X_000097 ^ 3 + X_000097 ^ 4) /
            ((X_000097 + 16) * (256 + 96 * X_000097 + X_000097 ^ 2)),
            X_000098) | lgen(0, sqrt(X_000097), X_000098)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_106, plain,
    (X_000097 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000097 + 17920 * X_000097 ^ 2 +
             448 * X_000097 ^ 3 + X_000097 ^ 4) /
            ((X_000097 + 16) * (256 + 96 * X_000097 + X_000097 ^ 2)),
            X_000098) | sqrt(X_000097) <= X_000098),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    (X_000097 < 1 |
     X_000098 <
     (2048 +
      X_000097 *
      (3584 + X_000097 * (560 + X_000097 * (14 + X_000097 * 1/32)))) /
     (4096 + X_000097 * (1792 + X_000097 * (112 + X_000097))) |
     sqrt(X_000097) <= X_000098),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (9 + skoXC1 * skoXC1 < 1 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (2048 +
      (9 + skoXC1 * skoXC1) *
      (3584 +
       (9 + skoXC1 * skoXC1) *
       (560 +
        (9 + skoXC1 * skoXC1) * (14 + (9 + skoXC1 * skoXC1) * 1/32)))) /
     (4096 +
      (9 + skoXC1 * skoXC1) *
      (1792 + (9 + skoXC1 * skoXC1) * (112 + (9 + skoXC1 * skoXC1)))) |
     sqrt(9 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))), inference(subst, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (9 + skoXC1 * skoXC1 < 1 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (2048 +
      (9 + skoXC1 * skoXC1) *
      (3584 +
       (9 + skoXC1 * skoXC1) *
       (560 +
        (9 + skoXC1 * skoXC1) * (14 + (9 + skoXC1 * skoXC1) * 1/32)))) /
     (4096 +
      (9 + skoXC1 * skoXC1) *
      (1792 + (9 + skoXC1 * skoXC1) * (112 + (9 + skoXC1 * skoXC1))))),
    inference(resolve, [], [refute_0_108, refute_0_13])).

cnf(refute_0_110, plain,
    (8 < skoXC1 * (skoXC1 * -1) |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (2882401/32 +
      skoXC1 *
      (skoXC1 *
       (137257/8 +
        skoXC1 *
        (skoXC1 *
         (15251/16 +
          skoXC1 * (skoXC1 * (121/8 + skoXC1 * (skoXC1 * 1/32)))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))))),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (-16320 +
     skoXC1 *
     (-8160 +
      skoXC1 *
      (-5936 +
       skoXC1 *
       (-2224 +
        skoXC1 *
        (-1816/3 +
         skoXC1 *
         (-548/3 +
          skoXC1 *
          (-70/3 +
           skoXC1 *
           (-20/3 +
            skoXC1 *
            (-13/3 +
             skoXC1 *
             (-2/3 +
              skoXC1 *
              (-7/12 + skoXC1 * (-1/24 + skoXC1 * -1/48))))))))))) <
     (12 + skoXC1 * (6 + skoXC1)) *
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -1360 +
     skoXC1 *
     (skoXC1 *
      (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))) |
     (-16320 +
      skoXC1 *
      (-8160 +
       skoXC1 *
       (-5936 +
        skoXC1 *
        (-2224 +
         skoXC1 *
         (-1816/3 +
          skoXC1 *
          (-548/3 +
           skoXC1 *
           (-70/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 *
             (-13/3 +
              skoXC1 *
              (-2/3 +
               skoXC1 *
               (-7/12 + skoXC1 * (-1/24 + skoXC1 * -1/48)))))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) <=
     12 + skoXC1 * (6 + skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_112, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) *
     (12 + skoXC1 * (-6 + skoXC1)) | 12 + skoXC1 * (-6 + skoXC1) <= 0 |
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) <=
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_113, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_100])).

cnf(refute_0_114, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) <
     skoXC1 *
     (skoXC1 *
      (2056/3 +
       skoXC1 *
       (skoXC1 *
        (688/3 +
         skoXC1 *
         (skoXC1 *
          (82/3 +
           skoXC1 * (skoXC1 * (4/3 + skoXC1 * (skoXC1 * 1/48))))))))) |
     1360 +
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (2056/3 +
       skoXC1 *
       (skoXC1 *
        (688/3 +
         skoXC1 *
         (skoXC1 *
          (82/3 +
           skoXC1 * (skoXC1 * (4/3 + skoXC1 * (skoXC1 * 1/48))))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) <=
     -1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_115, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (4112 +
      skoXC1 *
      (skoXC1 *
       (1376 +
        skoXC1 *
        (skoXC1 *
         (164 + skoXC1 * (skoXC1 * (8 + skoXC1 * (skoXC1 * 1/8)))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0 |
     (4112 +
      skoXC1 *
      (skoXC1 *
       (1376 +
        skoXC1 *
        (skoXC1 *
         (164 + skoXC1 * (skoXC1 * (8 + skoXC1 * (skoXC1 * 1/8)))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_116, plain,
    (~ lgen(0, sqrt(X_000090), X_000091) | sqrt(X_000090) <= X_000091),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_117, plain,
    (X_000090 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000090 ^ 4 + 28 * X_000090 ^ 3 + 70 * X_000090 ^ 2 +
             28 * X_000090 + 1) /
            ((X_000090 + 1) * (X_000090 ^ 2 + 6 * X_000090 + 1)),
            X_000091) | lgen(0, sqrt(X_000090), X_000091)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_118, plain,
    (X_000090 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000090 ^ 4 + 28 * X_000090 ^ 3 + 70 * X_000090 ^ 2 +
             28 * X_000090 + 1) /
            ((X_000090 + 1) * (X_000090 ^ 2 + 6 * X_000090 + 1)),
            X_000091) | sqrt(X_000090) <= X_000091),
    inference(resolve, [], [refute_0_117, refute_0_116])).

cnf(refute_0_119, plain,
    (X_000090 < 0 |
     X_000091 <
     (1/8 +
      X_000090 *
      (7/2 + X_000090 * (35/4 + X_000090 * (7/2 + X_000090 * 1/8)))) /
     (1 + X_000090 * (7 + X_000090 * (7 + X_000090))) |
     sqrt(X_000090) <= X_000091),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/8 +
      (9 + skoXC1 * skoXC1) *
      (7/2 +
       (9 + skoXC1 * skoXC1) *
       (35/4 +
        (9 + skoXC1 * skoXC1) * (7/2 + (9 + skoXC1 * skoXC1) * 1/8)))) /
     (1 +
      (9 + skoXC1 * skoXC1) *
      (7 + (9 + skoXC1 * skoXC1) * (7 + (9 + skoXC1 * skoXC1)))) |
     sqrt(9 + skoXC1 * skoXC1) <=
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6))), inference(subst, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (9 + skoXC1 * skoXC1 < 0 |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (1/8 +
      (9 + skoXC1 * skoXC1) *
      (7/2 +
       (9 + skoXC1 * skoXC1) *
       (35/4 +
        (9 + skoXC1 * skoXC1) * (7/2 + (9 + skoXC1 * skoXC1) * 1/8)))) /
     (1 +
      (9 + skoXC1 * skoXC1) *
      (7 + (9 + skoXC1 * skoXC1) * (7 + (9 + skoXC1 * skoXC1))))),
    inference(resolve, [], [refute_0_120, refute_0_13])).

cnf(refute_0_122, plain,
    (9 < skoXC1 * (skoXC1 * -1) |
     (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (4112 +
      skoXC1 *
      (skoXC1 *
       (1376 +
        skoXC1 *
        (skoXC1 *
         (164 + skoXC1 * (skoXC1 * (8 + skoXC1 * (skoXC1 * 1/8)))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (4112 +
      skoXC1 *
      (skoXC1 *
       (1376 +
        skoXC1 *
        (skoXC1 *
         (164 + skoXC1 * (skoXC1 * (8 + skoXC1 * (skoXC1 * 1/8)))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_21, refute_0_122])).

cnf(refute_0_124, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (4112 +
      skoXC1 *
      (skoXC1 *
       (1376 +
        skoXC1 *
        (skoXC1 *
         (164 + skoXC1 * (skoXC1 * (8 + skoXC1 * (skoXC1 * 1/8)))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(resolve, [], [refute_0_115, refute_0_123])).

cnf(refute_0_125, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (2056/3 +
       skoXC1 *
       (skoXC1 *
        (688/3 +
         skoXC1 *
         (skoXC1 *
          (82/3 +
           skoXC1 * (skoXC1 * (4/3 + skoXC1 * (skoXC1 * 1/48))))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (2056/3 +
       skoXC1 *
       (skoXC1 *
        (688/3 +
         skoXC1 *
         (skoXC1 *
          (82/3 +
           skoXC1 * (skoXC1 * (4/3 + skoXC1 * (skoXC1 * 1/48))))))))) /
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_125, refute_0_12])).

cnf(refute_0_127, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (1360 +
      skoXC1 *
      (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))) <
     skoXC1 *
     (skoXC1 *
      (2056/3 +
       skoXC1 *
       (skoXC1 *
        (688/3 +
         skoXC1 *
         (skoXC1 *
          (82/3 +
           skoXC1 * (skoXC1 * (4/3 + skoXC1 * (skoXC1 * 1/48))))))))) |
     1360 +
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))) <= 0),
    inference(resolve, [], [refute_0_114, refute_0_126])).

cnf(refute_0_128, plain,
    (-1360 +
     skoXC1 *
     (-1360 +
      skoXC1 *
      (-3184/3 +
       skoXC1 *
       (-1808/3 +
        skoXC1 *
        (-790/3 +
         skoXC1 *
         (-290/3 +
          skoXC1 *
          (-85/3 +
           skoXC1 *
           (-20/3 +
            skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48))))))))) <
     exp(skoXC1) *
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))) <=
     -1360), inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (-1360 <
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_130, plain,
    (-1360 +
     skoXC1 *
     (-1360 +
      skoXC1 *
      (-3184/3 +
       skoXC1 *
       (-1808/3 +
        skoXC1 *
        (-790/3 +
         skoXC1 *
         (-290/3 +
          skoXC1 *
          (-85/3 +
           skoXC1 *
           (-20/3 +
            skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48))))))))) <
     exp(skoXC1) *
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (exp(skoXC1) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -1360 +
     skoXC1 *
     (skoXC1 *
      (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))),
    inference(split, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (exp(skoXC1) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))) <=
     -1360), inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (exp(skoXC1) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_132, refute_0_129])).

cnf(refute_0_134, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_113, refute_0_133])).

cnf(refute_0_135, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_56, refute_0_134])).

cnf(refute_0_136, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-1360 +
      skoXC1 *
      (-1360 +
       skoXC1 *
       (-3184/3 +
        skoXC1 *
        (-1808/3 +
         skoXC1 *
         (-790/3 +
          skoXC1 *
          (-290/3 +
           skoXC1 *
           (-85/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 * (-4/3 + skoXC1 * (-1/6 + skoXC1 * -1/48)))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) *
     (12 + skoXC1 * (-6 + skoXC1)) | 12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_112, refute_0_135])).

cnf(refute_0_137, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-16320 +
      skoXC1 *
      (-8160 +
       skoXC1 *
       (-5936 +
        skoXC1 *
        (-2224 +
         skoXC1 *
         (-1816/3 +
          skoXC1 *
          (-548/3 +
           skoXC1 *
           (-70/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 *
             (-13/3 +
              skoXC1 *
              (-2/3 +
               skoXC1 *
               (-7/12 + skoXC1 * (-1/24 + skoXC1 * -1/48)))))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain, (-12 < skoXC1 * (-6 + skoXC1)),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_139, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-16320 +
      skoXC1 *
      (-8160 +
       skoXC1 *
       (-5936 +
        skoXC1 *
        (-2224 +
         skoXC1 *
         (-1816/3 +
          skoXC1 *
          (-548/3 +
           skoXC1 *
           (-70/3 +
            skoXC1 *
            (-20/3 +
             skoXC1 *
             (-13/3 +
              skoXC1 *
              (-2/3 +
               skoXC1 *
               (-7/12 + skoXC1 * (-1/24 + skoXC1 * -1/48)))))))))))) /
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain,
    (-16320 +
     skoXC1 *
     (-8160 +
      skoXC1 *
      (-5936 +
       skoXC1 *
       (-2224 +
        skoXC1 *
        (-1816/3 +
         skoXC1 *
         (-548/3 +
          skoXC1 *
          (-70/3 +
           skoXC1 *
           (-20/3 +
            skoXC1 *
            (-13/3 +
             skoXC1 *
             (-2/3 +
              skoXC1 *
              (-7/12 + skoXC1 * (-1/24 + skoXC1 * -1/48))))))))))) <
     (12 + skoXC1 * (6 + skoXC1)) *
     (-1360 +
      skoXC1 *
      (skoXC1 *
       (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -1360 +
     skoXC1 *
     (skoXC1 *
      (-376 + skoXC1 * (skoXC1 * (-34 + skoXC1 * (skoXC1 * -1)))))),
    inference(resolve, [], [refute_0_111, refute_0_139])).

cnf(refute_0_141, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (64 +
       skoXC1 *
       (-32 +
        skoXC1 *
        (-536/3 +
         skoXC1 *
         (-64/3 +
          skoXC1 *
          (-68/3 +
           skoXC1 *
           (2/3 +
            skoXC1 *
            (10/3 +
             skoXC1 *
             (2/3 +
              skoXC1 * (7/12 + skoXC1 * (1/24 + skoXC1 * 1/48))))))))))) |
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1)))) <=
     -1360), inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (64 +
       skoXC1 *
       (-32 +
        skoXC1 *
        (-536/3 +
         skoXC1 *
         (-64/3 +
          skoXC1 *
          (-68/3 +
           skoXC1 *
           (2/3 +
            skoXC1 *
            (10/3 +
             skoXC1 *
             (2/3 +
              skoXC1 * (7/12 + skoXC1 * (1/24 + skoXC1 * 1/48))))))))))) |
     -1360 <
     skoXC1 *
     (skoXC1 * (376 + skoXC1 * (skoXC1 * (34 + skoXC1 * skoXC1))))),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_143, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (64 +
       skoXC1 *
       (-32 +
        skoXC1 *
        (-536/3 +
         skoXC1 *
         (-64/3 +
          skoXC1 *
          (-68/3 +
           skoXC1 *
           (2/3 +
            skoXC1 *
            (10/3 +
             skoXC1 *
             (2/3 +
              skoXC1 * (7/12 + skoXC1 * (1/24 + skoXC1 * 1/48)))))))))))),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain, (skoXC1 * (skoXC1 * -1) <= 8),
    inference(decision, [], [refute_0_143, refute_0_11])).

cnf(refute_0_145, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) /
     (skoXC1 * (skoXC1 * 1/6)) <
     (2882401/32 +
      skoXC1 *
      (skoXC1 *
       (137257/8 +
        skoXC1 *
        (skoXC1 *
         (15251/16 +
          skoXC1 * (skoXC1 * (121/8 + skoXC1 * (skoXC1 * 1/32)))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_144, refute_0_110])).

cnf(refute_0_146, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     (2882401/32 +
      skoXC1 *
      (skoXC1 *
       (137257/8 +
        skoXC1 *
        (skoXC1 *
         (15251/16 +
          skoXC1 * (skoXC1 * (121/8 + skoXC1 * (skoXC1 * 1/32)))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) *
     (skoXC1 * (skoXC1 * 1/6)) | skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(resolve, [], [refute_0_103, refute_0_145])).

cnf(refute_0_147, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (2882401/192 +
       skoXC1 *
       (skoXC1 *
        (137257/48 +
         skoXC1 *
         (skoXC1 *
          (15251/96 +
           skoXC1 * (skoXC1 * (121/48 + skoXC1 * (skoXC1 * 1/192))))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) |
     skoXC1 * (skoXC1 * 1/6) <= 0),
    inference(arithmetic, [], [refute_0_146])).

cnf(refute_0_148, plain,
    (-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (2882401/192 +
       skoXC1 *
       (skoXC1 *
        (137257/48 +
         skoXC1 *
         (skoXC1 *
          (15251/96 +
           skoXC1 * (skoXC1 * (121/48 + skoXC1 * (skoXC1 * 1/192))))))))) /
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))))),
    inference(resolve, [], [refute_0_147, refute_0_12])).

cnf(refute_0_149, plain,
    ((-1 + skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/6)) + exp(skoXC1)) *
     (30025 +
      skoXC1 *
      (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))) <
     skoXC1 *
     (skoXC1 *
      (2882401/192 +
       skoXC1 *
       (skoXC1 *
        (137257/48 +
         skoXC1 *
         (skoXC1 *
          (15251/96 +
           skoXC1 * (skoXC1 * (121/48 + skoXC1 * (skoXC1 * 1/192))))))))) |
     30025 +
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))) <= 0),
    inference(resolve, [], [refute_0_102, refute_0_148])).

cnf(refute_0_150, plain,
    (-30025 +
     skoXC1 *
     (-30025 +
      skoXC1 *
      (-3660193/192 +
       skoXC1 *
       (-54331/6 +
        skoXC1 *
        (-143929/48 +
         skoXC1 *
         (-4885/6 +
          skoXC1 *
          (-15347/96 +
           skoXC1 *
           (-145/6 +
            skoXC1 * (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192))))))))) <
     exp(skoXC1) *
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))) <=
     -30025), inference(arithmetic, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (-30025 <
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))),
    inference(decision, [], [refute_0_143, refute_0_11])).

cnf(refute_0_152, plain,
    (-30025 +
     skoXC1 *
     (-30025 +
      skoXC1 *
      (-3660193/192 +
       skoXC1 *
       (-54331/6 +
        skoXC1 *
        (-143929/48 +
         skoXC1 *
         (-4885/6 +
          skoXC1 *
          (-15347/96 +
           skoXC1 *
           (-145/6 +
            skoXC1 * (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192))))))))) <
     exp(skoXC1) *
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_150, refute_0_151])).

cnf(refute_0_153, plain,
    (exp(skoXC1) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -30025 +
     skoXC1 *
     (skoXC1 *
      (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))),
    inference(split, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (exp(skoXC1) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))) <=
     -30025), inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (exp(skoXC1) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_154, refute_0_151])).

cnf(refute_0_156, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_101, refute_0_155])).

cnf(refute_0_157, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_56, refute_0_156])).

cnf(refute_0_158, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-30025 +
      skoXC1 *
      (-30025 +
       skoXC1 *
       (-3660193/192 +
        skoXC1 *
        (-54331/6 +
         skoXC1 *
         (-143929/48 +
          skoXC1 *
          (-4885/6 +
           skoXC1 *
           (-15347/96 +
            skoXC1 *
            (-145/6 +
             skoXC1 *
             (-121/48 + skoXC1 * (-1/6 + skoXC1 * -1/192)))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) *
     (12 + skoXC1 * (-6 + skoXC1)) | 12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_157])).

cnf(refute_0_159, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-360300 +
      skoXC1 *
      (-180150 +
       skoXC1 *
       (-1258193/16 +
        skoXC1 *
        (-777791/32 +
         skoXC1 *
         (-137233/192 +
          skoXC1 *
          (-20017/24 +
           skoXC1 *
           (-1531/48 +
            skoXC1 *
            (-6959/48 +
             skoXC1 *
             (-4331/96 +
              skoXC1 *
              (-265/24 +
               skoXC1 *
               (-19/12 + skoXC1 * (-13/96 + skoXC1 * -1/192)))))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-360300 +
      skoXC1 *
      (-180150 +
       skoXC1 *
       (-1258193/16 +
        skoXC1 *
        (-777791/32 +
         skoXC1 *
         (-137233/192 +
          skoXC1 *
          (-20017/24 +
           skoXC1 *
           (-1531/48 +
            skoXC1 *
            (-6959/48 +
             skoXC1 *
             (-4331/96 +
              skoXC1 *
              (-265/24 +
               skoXC1 *
               (-19/12 + skoXC1 * (-13/96 + skoXC1 * -1/192)))))))))))) /
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_159, refute_0_138])).

cnf(refute_0_161, plain,
    (-360300 +
     skoXC1 *
     (-180150 +
      skoXC1 *
      (-1258193/16 +
       skoXC1 *
       (-777791/32 +
        skoXC1 *
        (-137233/192 +
         skoXC1 *
         (-20017/24 +
          skoXC1 *
          (-1531/48 +
           skoXC1 *
           (-6959/48 +
            skoXC1 *
            (-4331/96 +
             skoXC1 *
             (-265/24 +
              skoXC1 *
              (-19/12 + skoXC1 * (-13/96 + skoXC1 * -1/192))))))))))) <
     (12 + skoXC1 * (6 + skoXC1)) *
     (-30025 +
      skoXC1 *
      (skoXC1 *
       (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))) |
     0 <=
     -30025 +
     skoXC1 *
     (skoXC1 *
      (-4051 + skoXC1 * (skoXC1 * (-139 + skoXC1 * (skoXC1 * -1)))))),
    inference(resolve, [], [refute_0_95, refute_0_160])).

cnf(refute_0_162, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/16 +
       skoXC1 *
       (-1/32 +
        skoXC1 *
        (-960815/192 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (-5717/48 +
           skoXC1 *
           (6671/48 +
            skoXC1 *
            (4235/96 +
             skoXC1 *
             (265/24 +
              skoXC1 *
              (19/12 + skoXC1 * (13/96 + skoXC1 * 1/192))))))))))) |
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1)))) <=
     -30025), inference(arithmetic, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/16 +
       skoXC1 *
       (-1/32 +
        skoXC1 *
        (-960815/192 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (-5717/48 +
           skoXC1 *
           (6671/48 +
            skoXC1 *
            (4235/96 +
             skoXC1 *
             (265/24 +
              skoXC1 *
              (19/12 + skoXC1 * (13/96 + skoXC1 * 1/192))))))))))) |
     -30025 <
     skoXC1 *
     (skoXC1 * (4051 + skoXC1 * (skoXC1 * (139 + skoXC1 * skoXC1))))),
    inference(decision, [], [refute_0_143, refute_0_11])).

cnf(refute_0_164, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/16 +
       skoXC1 *
       (-1/32 +
        skoXC1 *
        (-960815/192 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (-5717/48 +
           skoXC1 *
           (6671/48 +
            skoXC1 *
            (4235/96 +
             skoXC1 *
             (265/24 +
              skoXC1 *
              (19/12 + skoXC1 * (13/96 + skoXC1 * 1/192)))))))))))),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (560 +
       skoXC1 *
       (-280 +
        skoXC1 *
        (130 +
         skoXC1 *
         (-167/3 + skoXC1 * (25/2 + skoXC1 * (-3/2 + skoXC1 * 1/12))))))) |
     -1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))),
    inference(decision, [], [refute_0_164, refute_0_143, refute_0_11])).

cnf(refute_0_166, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (560 +
       skoXC1 *
       (-280 +
        skoXC1 *
        (130 +
         skoXC1 *
         (-167/3 + skoXC1 * (25/2 + skoXC1 * (-3/2 + skoXC1 * 1/12)))))))),
    inference(resolve, [], [refute_0_94, refute_0_165])).

cnf(refute_0_167, plain,
    (-1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))),
    inference(decision, [],
              [refute_0_166, refute_0_164, refute_0_143, refute_0_11])).

cnf(refute_0_168, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-16800 +
      skoXC1 *
      (-8400 +
       skoXC1 *
       (-4600 +
        skoXC1 *
        (-480 +
         skoXC1 *
         (-450 +
          skoXC1 *
          (610/3 +
           skoXC1 *
           (-190/3 +
            skoXC1 *
            (67/3 +
             skoXC1 * (-31/6 + skoXC1 * (2/3 + skoXC1 * -1/24)))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_88, refute_0_167])).

cnf(refute_0_169, plain,
    (-16800 +
     skoXC1 *
     (-8400 +
      skoXC1 *
      (-4600 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-450 +
         skoXC1 *
         (610/3 +
          skoXC1 *
          (-190/3 +
           skoXC1 *
           (67/3 +
            skoXC1 * (-31/6 + skoXC1 * (2/3 + skoXC1 * -1/24))))))))) <
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) *
     (-10 + skoXC1 * (skoXC1 * -1)) | 0 <= -10 + skoXC1 * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_78, refute_0_168])).

cnf(refute_0_170, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1120 +
       skoXC1 *
       (-560 +
        skoXC1 *
        (260 +
         skoXC1 *
         (-670/3 +
          skoXC1 *
          (187/3 +
           skoXC1 *
           (-67/3 +
            skoXC1 * (31/6 + skoXC1 * (-2/3 + skoXC1 * 1/24))))))))) |
     skoXC1 * skoXC1 <= -10), inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1120 +
       skoXC1 *
       (-560 +
        skoXC1 *
        (260 +
         skoXC1 *
         (-670/3 +
          skoXC1 *
          (187/3 +
           skoXC1 *
           (-67/3 +
            skoXC1 * (31/6 + skoXC1 * (-2/3 + skoXC1 * 1/24))))))))) |
     -10 < skoXC1 * skoXC1),
    inference(decision, [],
              [refute_0_166, refute_0_164, refute_0_143, refute_0_11])).

cnf(refute_0_172, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1120 +
       skoXC1 *
       (-560 +
        skoXC1 *
        (260 +
         skoXC1 *
         (-670/3 +
          skoXC1 *
          (187/3 +
           skoXC1 *
           (-67/3 +
            skoXC1 * (31/6 + skoXC1 * (-2/3 + skoXC1 * 1/24)))))))))),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (221760 +
       skoXC1 *
       (-110880 +
        skoXC1 *
        (52920 +
         skoXC1 *
         (-22764 +
          skoXC1 *
          (5278 +
           skoXC1 *
           (-734 +
            skoXC1 * (383/6 + skoXC1 * (-10/3 + skoXC1 * 1/12))))))))) |
     -665280 <
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))),
    inference(decision, [],
              [refute_0_172, refute_0_166, refute_0_164, refute_0_143,
               refute_0_11])).

cnf(refute_0_174, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (221760 +
       skoXC1 *
       (-110880 +
        skoXC1 *
        (52920 +
         skoXC1 *
         (-22764 +
          skoXC1 *
          (5278 +
           skoXC1 *
           (-734 +
            skoXC1 * (383/6 + skoXC1 * (-10/3 + skoXC1 * 1/12)))))))))),
    inference(resolve, [], [refute_0_77, refute_0_173])).

cnf(refute_0_175, plain,
    (-665280 <
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))),
    inference(decision, [],
              [refute_0_174, refute_0_172, refute_0_166, refute_0_164,
               refute_0_143, refute_0_11])).

cnf(refute_0_176, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-6652800 +
      skoXC1 *
      (-3326400 +
       skoXC1 *
       (-1864800 +
        skoXC1 *
        (-211680 +
         skoXC1 *
         (-189840 +
          skoXC1 *
          (79380 +
           skoXC1 *
           (-26190 +
            skoXC1 *
            (9370 +
             skoXC1 *
             (-6614/3 +
              skoXC1 *
              (958/3 +
               skoXC1 *
               (-29 + skoXC1 * (19/12 + skoXC1 * -1/24)))))))))))) /
     (-10 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_59, refute_0_175])).

cnf(refute_0_177, plain,
    (-6652800 +
     skoXC1 *
     (-3326400 +
      skoXC1 *
      (-1864800 +
       skoXC1 *
       (-211680 +
        skoXC1 *
        (-189840 +
         skoXC1 *
         (79380 +
          skoXC1 *
          (-26190 +
           skoXC1 *
           (9370 +
            skoXC1 *
            (-6614/3 +
             skoXC1 *
             (958/3 +
              skoXC1 * (-29 + skoXC1 * (19/12 + skoXC1 * -1/24))))))))))) <
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) *
     (-10 + skoXC1 * (skoXC1 * -1)) | 0 <= -10 + skoXC1 * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_28, refute_0_176])).

cnf(refute_0_178, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (443520 +
       skoXC1 *
       (-221760 +
        skoXC1 *
        (105840 +
         skoXC1 *
         (-89880 +
          skoXC1 *
          (25340 +
           skoXC1 *
           (-9412 +
            skoXC1 *
            (6611/3 +
             skoXC1 *
             (-958/3 +
              skoXC1 * (29 + skoXC1 * (-19/12 + skoXC1 * 1/24))))))))))) |
     skoXC1 * skoXC1 <= -10), inference(arithmetic, [], [refute_0_177])).

cnf(refute_0_179, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (443520 +
       skoXC1 *
       (-221760 +
        skoXC1 *
        (105840 +
         skoXC1 *
         (-89880 +
          skoXC1 *
          (25340 +
           skoXC1 *
           (-9412 +
            skoXC1 *
            (6611/3 +
             skoXC1 *
             (-958/3 +
              skoXC1 * (29 + skoXC1 * (-19/12 + skoXC1 * 1/24))))))))))) |
     -10 < skoXC1 * skoXC1),
    inference(decision, [],
              [refute_0_174, refute_0_172, refute_0_166, refute_0_164,
               refute_0_143, refute_0_11])).

cnf(refute_0_180, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (443520 +
       skoXC1 *
       (-221760 +
        skoXC1 *
        (105840 +
         skoXC1 *
         (-89880 +
          skoXC1 *
          (25340 +
           skoXC1 *
           (-9412 +
            skoXC1 *
            (6611/3 +
             skoXC1 *
             (-958/3 +
              skoXC1 * (29 + skoXC1 * (-19/12 + skoXC1 * 1/24)))))))))))),
    inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    (-96076792028200960 <
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1))))))))))))))))))))))))))))),
    inference(decision, [],
              [refute_0_180, refute_0_174, refute_0_172, refute_0_166,
               refute_0_164, refute_0_143, refute_0_11])).

cnf(refute_0_182, plain,
    (-96076792028200960 +
     skoXC1 *
     (-96076792028200960 +
      skoXC1 *
      (-512409557461368832/3 +
       skoXC1 *
       (-416332765366059008/3 +
        skoXC1 *
        (-409660765977444352/3 +
         skoXC1 *
         (-278889576678817792/3 +
          skoXC1 *
          (-197441761386168320/3 +
           skoXC1 *
           (-114931368063926272/3 +
            skoXC1 *
            (-21527809123418112 +
             skoXC1 *
             (-10862219843076096 +
              skoXC1 *
              (-5085812237795328 +
               skoXC1 *
               (-2244559021539328 +
                skoXC1 *
                (-899640770625536 +
                 skoXC1 *
                 (-349234651267072 +
                  skoXC1 *
                  (-121711614853120 +
                   skoXC1 *
                   (-41671100727296 +
                    skoXC1 *
                    (-12736077383680 +
                     skoXC1 *
                     (-3846151243776 +
                      skoXC1 *
                      (-1034455569408 +
                       skoXC1 *
                       (-274822047744 +
                        skoXC1 *
                        (-64989789952 +
                         skoXC1 *
                         (-15102540032 +
                          skoXC1 *
                          (-3123732352 +
                           skoXC1 *
                           (-628825600 +
                            skoXC1 *
                            (-112565648 +
                             skoXC1 *
                             (-19329232 +
                              skoXC1 *
                              (-2940920 +
                               skoXC1 *
                               (-420272 +
                                skoXC1 *
                                (-52706 +
                                 skoXC1 *
                                 (-18034/3 +
                                  skoXC1 *
                                  (-1763/3 +
                                   skoXC1 *
                                   (-148/3 +
                                    skoXC1 *
                                    (-10/3 +
                                     skoXC1 *
                                     (-1/6 +
                                      skoXC1 *
                                      -1/
                                      192))))))))))))))))))))))))))))))))) <
     exp(skoXC1) *
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1))))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_27, refute_0_181])).

cnf(refute_0_183, plain,
    (exp(skoXC1) <
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     0 <=
     -96076792028200960 +
     skoXC1 *
     (skoXC1 *
      (-122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (-72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (-26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (-6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (-1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (-155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (-15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (-1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (-71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (-3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (-102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (-2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (-34328 +
                               skoXC1 *
                               (skoXC1 *
                                (-290 +
                                 skoXC1 *
                                 (skoXC1 *
                                  -1)))))))))))))))))))))))))))))),
    inference(split, [], [refute_0_182])).

cnf(refute_0_184, plain,
    (exp(skoXC1) <
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <=
     -96076792028200960), inference(arithmetic, [], [refute_0_183])).

cnf(refute_0_185, plain,
    (exp(skoXC1) <
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1))))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_184, refute_0_181])).

cnf(refute_0_186, plain,
    (1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     (-96076792028200960 +
      skoXC1 *
      (-96076792028200960 +
       skoXC1 *
       (-512409557461368832/3 +
        skoXC1 *
        (-416332765366059008/3 +
         skoXC1 *
         (-409660765977444352/3 +
          skoXC1 *
          (-278889576678817792/3 +
           skoXC1 *
           (-197441761386168320/3 +
            skoXC1 *
            (-114931368063926272/3 +
             skoXC1 *
             (-21527809123418112 +
              skoXC1 *
              (-10862219843076096 +
               skoXC1 *
               (-5085812237795328 +
                skoXC1 *
                (-2244559021539328 +
                 skoXC1 *
                 (-899640770625536 +
                  skoXC1 *
                  (-349234651267072 +
                   skoXC1 *
                   (-121711614853120 +
                    skoXC1 *
                    (-41671100727296 +
                     skoXC1 *
                     (-12736077383680 +
                      skoXC1 *
                      (-3846151243776 +
                       skoXC1 *
                       (-1034455569408 +
                        skoXC1 *
                        (-274822047744 +
                         skoXC1 *
                         (-64989789952 +
                          skoXC1 *
                          (-15102540032 +
                           skoXC1 *
                           (-3123732352 +
                            skoXC1 *
                            (-628825600 +
                             skoXC1 *
                             (-112565648 +
                              skoXC1 *
                              (-19329232 +
                               skoXC1 *
                               (-2940920 +
                                skoXC1 *
                                (-420272 +
                                 skoXC1 *
                                 (-52706 +
                                  skoXC1 *
                                  (-18034/3 +
                                   skoXC1 *
                                   (-1763/3 +
                                    skoXC1 *
                                    (-148/3 +
                                     skoXC1 *
                                     (-10/3 +
                                      skoXC1 *
                                      (-1/6 +
                                       skoXC1 *
                                       -1/
                                       192)))))))))))))))))))))))))))))))))) /
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1))))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_5, refute_0_185])).

cnf(refute_0_187, plain,
    (-96076792028200960 +
     skoXC1 *
     (-96076792028200960 +
      skoXC1 *
      (-512409557461368832/3 +
       skoXC1 *
       (-416332765366059008/3 +
        skoXC1 *
        (-409660765977444352/3 +
         skoXC1 *
         (-278889576678817792/3 +
          skoXC1 *
          (-197441761386168320/3 +
           skoXC1 *
           (-114931368063926272/3 +
            skoXC1 *
            (-21527809123418112 +
             skoXC1 *
             (-10862219843076096 +
              skoXC1 *
              (-5085812237795328 +
               skoXC1 *
               (-2244559021539328 +
                skoXC1 *
                (-899640770625536 +
                 skoXC1 *
                 (-349234651267072 +
                  skoXC1 *
                  (-121711614853120 +
                   skoXC1 *
                   (-41671100727296 +
                    skoXC1 *
                    (-12736077383680 +
                     skoXC1 *
                     (-3846151243776 +
                      skoXC1 *
                      (-1034455569408 +
                       skoXC1 *
                       (-274822047744 +
                        skoXC1 *
                        (-64989789952 +
                         skoXC1 *
                         (-15102540032 +
                          skoXC1 *
                          (-3123732352 +
                           skoXC1 *
                           (-628825600 +
                            skoXC1 *
                            (-112565648 +
                             skoXC1 *
                             (-19329232 +
                              skoXC1 *
                              (-2940920 +
                               skoXC1 *
                               (-420272 +
                                skoXC1 *
                                (-52706 +
                                 skoXC1 *
                                 (-18034/3 +
                                  skoXC1 *
                                  (-1763/3 +
                                   skoXC1 *
                                   (-148/3 +
                                    skoXC1 *
                                    (-10/3 +
                                     skoXC1 *
                                     (-1/6 +
                                      skoXC1 *
                                      -1/
                                      192))))))))))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) *
     (-96076792028200960 +
      skoXC1 *
      (skoXC1 *
       (-122764789783986176 +
        skoXC1 *
        (skoXC1 *
         (-72502393928941568 +
          skoXC1 *
          (skoXC1 *
           (-26226723699818496 +
            skoXC1 *
            (skoXC1 *
             (-6491099226439680 +
              skoXC1 *
              (skoXC1 *
               (-1162709150466048 +
                skoXC1 *
                (skoXC1 *
                 (-155449792856064 +
                  skoXC1 *
                  (skoXC1 *
                   (-15762801917952 +
                    skoXC1 *
                    (skoXC1 *
                     (-1219017590784 +
                      skoXC1 *
                      (skoXC1 *
                       (-71652449280 +
                        skoXC1 *
                        (skoXC1 *
                         (-3160465152 +
                          skoXC1 *
                          (skoXC1 *
                           (-102081408 +
                            skoXC1 *
                            (skoXC1 *
                             (-2315664 +
                              skoXC1 *
                              (skoXC1 *
                               (-34328 +
                                skoXC1 *
                                (skoXC1 *
                                 (-290 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1)))))))))))))))))))))))))))))) |
     0 <=
     -96076792028200960 +
     skoXC1 *
     (skoXC1 *
      (-122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (-72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (-26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (-6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (-1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (-155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (-15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (-1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (-71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (-3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (-102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (-2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (-34328 +
                               skoXC1 *
                               (skoXC1 *
                                (-290 +
                                 skoXC1 *
                                 (skoXC1 *
                                  -1)))))))))))))))))))))))))))))),
    inference(resolve, [], [refute_0_0, refute_0_186])).

cnf(refute_0_188, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (67108864/3 +
       skoXC1 *
       (skoXC1 *
        (-4003199488884736/3 +
         skoXC1 *
         (-2401919800705024/3 +
          skoXC1 *
          (-4181742207234998272/2187 +
           skoXC1 *
           (-11391104652125667328/10935 +
            skoXC1 *
            (-41650288147862388736/32805 +
             skoXC1 *
             (-1669454356240037576704/2657205 +
              skoXC1 *
              (-1377483424929901969408/2657205 +
               skoXC1 *
               (-618975198545902567424/2657205 +
                skoXC1 *
                (-52237351326653056876544/358722675 +
                 skoXC1 *
                 (-2366192940667914616832/39858075 +
                  skoXC1 *
                  (-6420035692991621300224/215233605 +
                   skoXC1 *
                   (-106684870565650038489088/9685512225 +
                    skoXC1 *
                    (-4957665803710978097152/1076168025 +
                     skoXC1 *
                     (-74452315631219316441088/48427561125 +
                      skoXC1 *
                      (-588710249376840752128/1076168025 +
                       skoXC1 *
                       (-7964684484975069097984/48427561125 +
                        skoXC1 *
                        (-6035298279500093696/119574225 +
                         skoXC1 *
                         (-73256403271034159872/5380840125 +
                          skoXC1 *
                          (-52118841844337920/14348907 +
                           skoXC1 *
                           (-2823102979257738368/3228504075 +
                            skoXC1 *
                            (-73013031136146688/358722675 +
                             skoXC1 *
                             (-701410639564214848/16142520375 +
                              skoXC1 *
                              (-3164996091260936/358722675 +
                               skoXC1 *
                               (-26753160550743344/16142520375 +
                                skoXC1 *
                                (-34991172228898/119574225 +
                                 skoXC1 *
                                 (-257923246014214/5380840125 +
                                  skoXC1 *
                                  (-97224380426/13286025 +
                                   skoXC1 *
                                   (-5570866695013/5380840125 +
                                    skoXC1 *
                                    (-48675156143/358722675 +
                                     skoXC1 *
                                     (-106122952961/6457008150 +
                                      skoXC1 *
                                      (-42166162903/22958251200 +
                                       skoXC1 *
                                       (-24425977211/129140163000 +
                                        skoXC1 *
                                        (-102771097/5739562800 +
                                         skoXC1 *
                                         (-401467813/258280326000 +
                                          skoXC1 *
                                          (-2822287/22958251200 +
                                           skoXC1 *
                                           (-3025247/344373768000 +
                                            skoXC1 *
                                            (-154381/275499014400 +
                                             skoXC1 *
                                             (-390169/12397455648000 +
                                              skoXC1 *
                                              (-833/550998028800 +
                                               skoXC1 *
                                               (-301/4958982259200 +
                                                skoXC1 *
                                                (-1/550998028800 +
                                                 skoXC1 *
                                                 -1/
                                                 24794911296000)))))))))))))))))))))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <=
     -96076792028200960), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (skoXC1 *
     (skoXC1 *
      (67108864/3 +
       skoXC1 *
       (skoXC1 *
        (-4003199488884736/3 +
         skoXC1 *
         (-2401919800705024/3 +
          skoXC1 *
          (-4181742207234998272/2187 +
           skoXC1 *
           (-11391104652125667328/10935 +
            skoXC1 *
            (-41650288147862388736/32805 +
             skoXC1 *
             (-1669454356240037576704/2657205 +
              skoXC1 *
              (-1377483424929901969408/2657205 +
               skoXC1 *
               (-618975198545902567424/2657205 +
                skoXC1 *
                (-52237351326653056876544/358722675 +
                 skoXC1 *
                 (-2366192940667914616832/39858075 +
                  skoXC1 *
                  (-6420035692991621300224/215233605 +
                   skoXC1 *
                   (-106684870565650038489088/9685512225 +
                    skoXC1 *
                    (-4957665803710978097152/1076168025 +
                     skoXC1 *
                     (-74452315631219316441088/48427561125 +
                      skoXC1 *
                      (-588710249376840752128/1076168025 +
                       skoXC1 *
                       (-7964684484975069097984/48427561125 +
                        skoXC1 *
                        (-6035298279500093696/119574225 +
                         skoXC1 *
                         (-73256403271034159872/5380840125 +
                          skoXC1 *
                          (-52118841844337920/14348907 +
                           skoXC1 *
                           (-2823102979257738368/3228504075 +
                            skoXC1 *
                            (-73013031136146688/358722675 +
                             skoXC1 *
                             (-701410639564214848/16142520375 +
                              skoXC1 *
                              (-3164996091260936/358722675 +
                               skoXC1 *
                               (-26753160550743344/16142520375 +
                                skoXC1 *
                                (-34991172228898/119574225 +
                                 skoXC1 *
                                 (-257923246014214/5380840125 +
                                  skoXC1 *
                                  (-97224380426/13286025 +
                                   skoXC1 *
                                   (-5570866695013/5380840125 +
                                    skoXC1 *
                                    (-48675156143/358722675 +
                                     skoXC1 *
                                     (-106122952961/6457008150 +
                                      skoXC1 *
                                      (-42166162903/22958251200 +
                                       skoXC1 *
                                       (-24425977211/129140163000 +
                                        skoXC1 *
                                        (-102771097/5739562800 +
                                         skoXC1 *
                                         (-401467813/258280326000 +
                                          skoXC1 *
                                          (-2822287/22958251200 +
                                           skoXC1 *
                                           (-3025247/344373768000 +
                                            skoXC1 *
                                            (-154381/275499014400 +
                                             skoXC1 *
                                             (-390169/12397455648000 +
                                              skoXC1 *
                                              (-833/550998028800 +
                                               skoXC1 *
                                               (-301/4958982259200 +
                                                skoXC1 *
                                                (-1/550998028800 +
                                                 skoXC1 *
                                                 -1/
                                                 24794911296000)))))))))))))))))))))))))))))))))))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <=
     -96076792028200960),
    inference(decision, [],
              [refute_0_180, refute_0_174, refute_0_172, refute_0_166,
               refute_0_164, refute_0_143, refute_0_11])).

cnf(refute_0_190, plain,
    (skoXC1 *
     (skoXC1 *
      (122764789783986176 +
       skoXC1 *
       (skoXC1 *
        (72502393928941568 +
         skoXC1 *
         (skoXC1 *
          (26226723699818496 +
           skoXC1 *
           (skoXC1 *
            (6491099226439680 +
             skoXC1 *
             (skoXC1 *
              (1162709150466048 +
               skoXC1 *
               (skoXC1 *
                (155449792856064 +
                 skoXC1 *
                 (skoXC1 *
                  (15762801917952 +
                   skoXC1 *
                   (skoXC1 *
                    (1219017590784 +
                     skoXC1 *
                     (skoXC1 *
                      (71652449280 +
                       skoXC1 *
                       (skoXC1 *
                        (3160465152 +
                         skoXC1 *
                         (skoXC1 *
                          (102081408 +
                           skoXC1 *
                           (skoXC1 *
                            (2315664 +
                             skoXC1 *
                             (skoXC1 *
                              (34328 +
                               skoXC1 *
                               (skoXC1 *
                                (290 +
                                 skoXC1 *
                                 skoXC1)))))))))))))))))))))))))))) <=
     -96076792028200960),
    inference(resolve, [], [refute_0_189, refute_0_188])).

cnf(refute_0_191, plain, ($false),
    inference(resolve, [], [refute_0_190, refute_0_181])).
% SZS output end CNFRefutation for exp-lower-14-sqrt-weak.tptp
