% SZS output start CNFRefutation for exp-361-4.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf2, axiom,
    (X < 0 | ~ lgen(R, Y, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12)) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf4, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 4 + 20 * X ^ 3 + 180 * X ^ 2 + 840 * X + 1680) /
            (X ^ 4 - 20 * X ^ 3 + 180 * X ^ 2 - 840 * X + 1680)) |
     lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf6, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 6 + 42 * X ^ 5 + 840 * X ^ 4 + 10080 * X ^ 3 +
             75600 * X ^ 2 + 332640 * X + 665280) /
            (X ^ 6 - 42 * X ^ 5 + 840 * X ^ 4 - 10080 * X ^ 3 +
             75600 * X ^ 2 - 332640 * X + 665280)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf5, axiom,
    (X < 0 | 7 < X |
     ~ lgen(R,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240), Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf7, axiom,
    (X < 0 | 19/2 < X |
     ~ lgen(R,
            -(X ^ 7 + 56 * X ^ 6 + 1512 * X ^ 5 + 25200 * X ^ 4 +
             277200 * X ^ 3 + 1995840 * X ^ 2 + 8648640 * X + 17297280) /
            (X ^ 7 - 56 * X ^ 6 + 1512 * X ^ 5 - 25200 * X ^ 4 +
             277200 * X ^ 3 - 1995840 * X ^ 2 + 8648640 * X - 17297280),
            Y) | lgen(R, exp(X), Y))).

fof(a, axiom, (a = 4)).

fof(exp_361, conjecture,
    (! [X] : (0 <= X => (exp(1) * X / a) ^ a <= 10 ^ -3 + exp(X)))).

fof(subgoal_0, plain,
    (! [X] : (0 <= X => (exp(1) * X / a) ^ a <= 10 ^ -3 + exp(X))),
    inference(strip, [], [exp_361])).

fof(negate_0_0, plain,
    (~ ! [X] : (0 <= X => (exp(1) * X / a) ^ a <= 10 ^ -3 + exp(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (10 ^ -3 + exp(X) < (exp(1) * X / a) ^ a & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / a) ^ a & 0 <= skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / a) ^ a),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (a = 4), inference(canonicalize, [], [a])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     28245729/10391023 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     28245729/10391023 <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     49171/18089 <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (~ lgen(0, exp(X_000120), X_000121) | exp(X_000120) <= X_000121),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (7 < X_000120 | X_000120 < 0 |
     ~ lgen(0,
            -(X_000120 ^ 5 + 30 * X_000120 ^ 4 + 420 * X_000120 ^ 3 +
             3360 * X_000120 ^ 2 + 15120 * X_000120 + 30240) /
            (X_000120 ^ 5 - 30 * X_000120 ^ 4 + 420 * X_000120 ^ 3 -
             3360 * X_000120 ^ 2 + 15120 * X_000120 - 30240), X_000121) |
     lgen(0, exp(X_000120), X_000121)),
    inference(subst, [], [exp_upper_bound_cf5])).

cnf(refute_0_4, plain,
    (7 < X_000120 | X_000120 < 0 |
     ~ lgen(0,
            -(X_000120 ^ 5 + 30 * X_000120 ^ 4 + 420 * X_000120 ^ 3 +
             3360 * X_000120 ^ 2 + 15120 * X_000120 + 30240) /
            (X_000120 ^ 5 - 30 * X_000120 ^ 4 + 420 * X_000120 ^ 3 -
             3360 * X_000120 ^ 2 + 15120 * X_000120 - 30240), X_000121) |
     exp(X_000120) <= X_000121),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (7 < X_000120 | X_000120 < 0 |
     X_000121 <
     (-30240 +
      X_000120 *
      (-15120 +
       X_000120 *
       (-3360 + X_000120 * (-420 + X_000120 * (-30 + X_000120 * -1))))) /
     (-30240 +
      X_000120 *
      (15120 +
       X_000120 *
       (-3360 + X_000120 * (420 + X_000120 * (-30 + X_000120))))) |
     exp(X_000120) <= X_000121), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     49171/18089 <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_8, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_9, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     49171/18089 <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_10, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / a) ^ a),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_11, plain, (a = 4),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_12, plain,
    ((exp(1) * skoXC1 / 4) ^ a = (exp(1) * skoXC1 / 4) ^ a),
    introduced(tautology, [refl])).

cnf(refute_0_13, plain,
    ((exp(1) * skoXC1 / 4) ^ a != (exp(1) * skoXC1 / 4) ^ a | a != 4 |
     (exp(1) * skoXC1 / 4) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_14, plain,
    (a != 4 | (exp(1) * skoXC1 / 4) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    inference(resolve, [], [refute_0_12, refute_0_13])).

cnf(refute_0_15, plain,
    ((exp(1) * skoXC1 / 4) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    inference(resolve, [], [refute_0_11, refute_0_14])).

cnf(refute_0_16, plain, (exp(1) * skoXC1 / a = exp(1) * skoXC1 / a),
    introduced(tautology, [refl])).

cnf(refute_0_17, plain,
    (exp(1) * skoXC1 / a != exp(1) * skoXC1 / a | a != 4 |
     exp(1) * skoXC1 / a = exp(1) * skoXC1 / 4),
    introduced(tautology, [equality])).

cnf(refute_0_18, plain,
    (a != 4 | exp(1) * skoXC1 / a = exp(1) * skoXC1 / 4),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain, (exp(1) * skoXC1 / a = exp(1) * skoXC1 / 4),
    inference(resolve, [], [refute_0_11, refute_0_18])).

cnf(refute_0_20, plain,
    ((exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / a) ^ a),
    introduced(tautology, [refl])).

cnf(refute_0_21, plain,
    (exp(1) * skoXC1 / a != exp(1) * skoXC1 / 4 |
     (exp(1) * skoXC1 / a) ^ a != (exp(1) * skoXC1 / a) ^ a |
     (exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ a),
    introduced(tautology, [equality])).

cnf(refute_0_22, plain,
    (exp(1) * skoXC1 / a != exp(1) * skoXC1 / 4 |
     (exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ a),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    ((exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ a),
    inference(resolve, [], [refute_0_19, refute_0_22])).

cnf(refute_0_24, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_25, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_26, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_28, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    ((exp(1) * skoXC1 / 4) ^ a != (exp(1) * skoXC1 / 4) ^ 4 |
     (exp(1) * skoXC1 / a) ^ a != (exp(1) * skoXC1 / 4) ^ a |
     (exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    inference(subst, [], [refute_0_28])).

cnf(refute_0_30, plain,
    ((exp(1) * skoXC1 / 4) ^ a != (exp(1) * skoXC1 / 4) ^ 4 |
     (exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_29])).

cnf(refute_0_31, plain,
    ((exp(1) * skoXC1 / a) ^ a = (exp(1) * skoXC1 / 4) ^ 4),
    inference(resolve, [], [refute_0_15, refute_0_30])).

cnf(refute_0_32, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / 4) ^ 4 |
     (exp(1) * skoXC1 / a) ^ a != (exp(1) * skoXC1 / 4) ^ 4 |
     (exp(1) * skoXC1 / a) ^ a <= 10 ^ -3 + exp(skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / 4) ^ 4 |
     (exp(1) * skoXC1 / a) ^ a <= 10 ^ -3 + exp(skoXC1)),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (10 ^ -3 + exp(skoXC1) < (exp(1) * skoXC1 / 4) ^ 4),
    inference(resolve, [], [refute_0_33, refute_0_10])).

cnf(refute_0_35, plain,
    (exp(skoXC1) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000017, exp(X_000016)) | X_000017 <= exp(X_000016)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3) |
     lgen(0, X_000017, exp(X_000016))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_38, plain,
    (~ lgen(0, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3) | X_000017 <= exp(X_000016)),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (1 +
     X_000016 *
     (1 +
      X_000016 *
      (1/2 +
       X_000016 *
       (1/6 +
        X_000016 *
        (1/24 +
         X_000016 *
         (1/120 +
          X_000016 *
          (121/87480 +
           X_000016 *
           (17/87480 +
            X_000016 *
            (49/2099520 +
             X_000016 *
             (409/170061120 +
              X_000016 *
              (361/1700611200 +
               X_000016 *
               (1/62985600 +
                X_000016 *
                (181/183666009600 +
                 X_000016 *
                 (1/20407334400 +
                  X_000016 *
                  (1/550998028800 +
                   X_000016 * 1/24794911296000)))))))))))))) < X_000017 |
     X_000017 <= exp(X_000016)), inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_40, refute_0_35])).

cnf(refute_0_42, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(split, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_45, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_46, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain,
    (0 <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (0 <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_47, refute_0_35])).

cnf(refute_0_49, plain,
    (1/1000 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (1/1000 < 0 |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))),
    inference(split, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_53, plain,
    (0 < exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(split, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (0 < 0 | skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) < 0 |
     0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(split, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) |
     0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (exp(1) < 0 / (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))) |
     0 <= skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(split, [], [refute_0_53])).

cnf(refute_0_57, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))) <= 0),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(resolve, [], [refute_0_57, refute_0_55])).

cnf(refute_0_59, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_60, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) <= 0),
    inference(decision, [], [refute_0_11, refute_0_59])).

cnf(refute_0_61, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256)))))) <= 0),
    inference(resolve, [], [refute_0_60, refute_0_58])).

cnf(refute_0_62, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))),
    inference(resolve, [], [refute_0_61, refute_0_51])).

cnf(refute_0_63, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1)), inference(resolve, [], [refute_0_43, refute_0_62])).

cnf(refute_0_64, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_65, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_64, refute_0_63])).

cnf(refute_0_66, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(resolve, [], [refute_0_9, refute_0_66])).

cnf(refute_0_68, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(resolve, [], [refute_0_68, refute_0_62])).

cnf(refute_0_70, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(split, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     49171/18089 <=
     1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_72, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(split, [], [refute_0_49])).

cnf(refute_0_73, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1)), inference(resolve, [], [refute_0_72, refute_0_62])).

cnf(refute_0_74, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_75, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_74, refute_0_73])).

cnf(refute_0_76, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(resolve, [], [refute_0_71, refute_0_76])).

cnf(refute_0_78, plain,
    (1/1000 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (1/1000 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(resolve, [], [refute_0_78, refute_0_62])).

cnf(refute_0_80, plain,
    (1/1000 < 0 |
     exp(1) *
     (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <
     0 |
     0 <
     exp(1) *
     (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))),
    inference(split, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))))) |
     0 <
     exp(1) *
     (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_83, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))))) <= 0),
    inference(split, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (0 < exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))),
    inference(split, [], [refute_0_62])).

cnf(refute_0_85, plain,
    (0 < 0 | skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))) < 0 |
     0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))),
    inference(split, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/256))) |
     0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))),
    inference(resolve, [], [refute_0_60, refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))) <= 0),
    inference(decision, [], [refute_0_11, refute_0_59, refute_0_87])).

cnf(refute_0_89, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -49171/4630784))))) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_83])).

cnf(refute_0_90, plain,
    (0 <
     exp(1) *
     (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))),
    inference(resolve, [], [refute_0_89, refute_0_81])).

cnf(refute_0_91, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1)), inference(resolve, [], [refute_0_70, refute_0_90])).

cnf(refute_0_92, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_8, refute_0_91])).

cnf(refute_0_93, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     49171/18089), inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(resolve, [], [refute_0_7, refute_0_93])).

cnf(refute_0_95, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(resolve, [], [refute_0_95, refute_0_90])).

cnf(refute_0_97, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(split, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     49171/18089 <=
     1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_99, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_100, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(split, [], [refute_0_79])).

cnf(refute_0_101, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1)), inference(resolve, [], [refute_0_100, refute_0_90])).

cnf(refute_0_102, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_99, refute_0_101])).

cnf(refute_0_103, plain,
    (1/1000 /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     49171/18089), inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(resolve, [], [refute_0_98, refute_0_103])).

cnf(refute_0_105, plain,
    (1/1000 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (1/1000 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(resolve, [], [refute_0_105, refute_0_90])).

cnf(refute_0_107, plain,
    (1/1000 < 0 |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <
     0 |
     0 <
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))),
    inference(split, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (0 <
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) |
     0 <
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (0 <
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_110, plain,
    (0 < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) <
     0 |
     0 <
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), inference(split, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (0 <
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) |
     0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (exp(1) <
     0 /
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) |
     0 <=
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), inference(split, [], [refute_0_109])).

cnf(refute_0_113, plain,
    (exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))) <=
     0), inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (0 <
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), inference(resolve, [], [refute_0_113, refute_0_111])).

cnf(refute_0_115, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776))) <=
     0), inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_116, plain,
    (exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -2417787241/83766251776)))) <=
     0), inference(resolve, [], [refute_0_115, refute_0_114])).

cnf(refute_0_117, plain,
    (0 <
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))),
    inference(resolve, [], [refute_0_116, refute_0_108])).

cnf(refute_0_118, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1)), inference(resolve, [], [refute_0_97, refute_0_117])).

cnf(refute_0_119, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_6, refute_0_118])).

cnf(refute_0_120, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     49171/18089), inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(resolve, [], [refute_0_1, refute_0_120])).

cnf(refute_0_122, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(resolve, [], [refute_0_122, refute_0_117])).

cnf(refute_0_124, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(split, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))),
    inference(decision, [], [refute_0_11, refute_0_59, refute_0_87])).

cnf(refute_0_126, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1)), inference(resolve, [], [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (~ lgen(0, exp(X_000127), X_000128) | exp(X_000127) <= X_000128),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_128, plain,
    (19/2 < X_000127 | X_000127 < 0 |
     ~ lgen(0,
            -(X_000127 ^ 7 + 56 * X_000127 ^ 6 + 1512 * X_000127 ^ 5 +
             25200 * X_000127 ^ 4 + 277200 * X_000127 ^ 3 +
             1995840 * X_000127 ^ 2 + 8648640 * X_000127 + 17297280) /
            (X_000127 ^ 7 - 56 * X_000127 ^ 6 + 1512 * X_000127 ^ 5 -
             25200 * X_000127 ^ 4 + 277200 * X_000127 ^ 3 -
             1995840 * X_000127 ^ 2 + 8648640 * X_000127 - 17297280),
            X_000128) | lgen(0, exp(X_000127), X_000128)),
    inference(subst, [], [exp_upper_bound_cf7])).

cnf(refute_0_129, plain,
    (19/2 < X_000127 | X_000127 < 0 |
     ~ lgen(0,
            -(X_000127 ^ 7 + 56 * X_000127 ^ 6 + 1512 * X_000127 ^ 5 +
             25200 * X_000127 ^ 4 + 277200 * X_000127 ^ 3 +
             1995840 * X_000127 ^ 2 + 8648640 * X_000127 + 17297280) /
            (X_000127 ^ 7 - 56 * X_000127 ^ 6 + 1512 * X_000127 ^ 5 -
             25200 * X_000127 ^ 4 + 277200 * X_000127 ^ 3 -
             1995840 * X_000127 ^ 2 + 8648640 * X_000127 - 17297280),
            X_000128) | exp(X_000127) <= X_000128),
    inference(resolve, [], [refute_0_128, refute_0_127])).

cnf(refute_0_130, plain,
    (19/2 < X_000127 | X_000127 < 0 |
     X_000128 <
     (-17297280 +
      X_000127 *
      (-8648640 +
       X_000127 *
       (-1995840 +
        X_000127 *
        (-277200 +
         X_000127 *
         (-25200 +
          X_000127 * (-1512 + X_000127 * (-56 + X_000127 * -1))))))) /
     (-17297280 +
      X_000127 *
      (8648640 +
       X_000127 *
       (-1995840 +
        X_000127 *
        (277200 +
         X_000127 *
         (-25200 + X_000127 * (1512 + X_000127 * (-56 + X_000127))))))) |
     exp(X_000127) <= X_000128),
    inference(arithmetic, [], [refute_0_129])).

cnf(refute_0_131, plain,
    (1 < 0 | 19/2 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     exp(1) <=
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(subst, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (1 < 0 | 19/2 < 1 |
     (1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1)))))))),
    inference(resolve, [], [refute_0_131, refute_0_126])).

cnf(refute_0_133, plain,
    ((1001/1000 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     28245729/10391023), inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (1001/1000 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <
     28245729/10391023 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_133])).

cnf(refute_0_135, plain,
    (1001/1000 <
     skoXC1 *
     (-1 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (-1/6 +
        skoXC1 *
        (8105860118958971186273/47234921988760301020416 +
         skoXC1 *
         (-1/120 +
          skoXC1 *
          (-121/87480 +
           skoXC1 *
           (-17/87480 +
            skoXC1 *
            (-49/2099520 +
             skoXC1 *
             (-409/170061120 +
              skoXC1 *
              (-361/1700611200 +
               skoXC1 *
               (-1/62985600 +
                skoXC1 *
                (-181/183666009600 +
                 skoXC1 *
                 (-1/20407334400 +
                  skoXC1 *
                  (-1/550998028800 +
                   skoXC1 * -1/24794911296000)))))))))))))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(arithmetic, [], [refute_0_134])).

cnf(refute_0_136, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     28245729/10391023 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     28245729/10391023 <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_137, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) |
     49171/18089 <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_138, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776)))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_139, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) |
     49171/18089 <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_140, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_141, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) |
     49171/18089 <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_142, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_143, plain,
    (~ lgen(0, X_000090, exp(X_000089)) | X_000090 <= exp(X_000089)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_144, plain,
    (X_000089 < 0 |
     ~ lgen(0, X_000090,
            (X_000089 ^ 6 + 42 * X_000089 ^ 5 + 840 * X_000089 ^ 4 +
             10080 * X_000089 ^ 3 + 75600 * X_000089 ^ 2 +
             332640 * X_000089 + 665280) /
            (X_000089 ^ 6 - 42 * X_000089 ^ 5 + 840 * X_000089 ^ 4 -
             10080 * X_000089 ^ 3 + 75600 * X_000089 ^ 2 -
             332640 * X_000089 + 665280)) |
     lgen(0, X_000090, exp(X_000089))),
    inference(subst, [], [exp_lower_bound_cf6])).

cnf(refute_0_145, plain,
    (X_000089 < 0 |
     ~ lgen(0, X_000090,
            (X_000089 ^ 6 + 42 * X_000089 ^ 5 + 840 * X_000089 ^ 4 +
             10080 * X_000089 ^ 3 + 75600 * X_000089 ^ 2 +
             332640 * X_000089 + 665280) /
            (X_000089 ^ 6 - 42 * X_000089 ^ 5 + 840 * X_000089 ^ 4 -
             10080 * X_000089 ^ 3 + 75600 * X_000089 ^ 2 -
             332640 * X_000089 + 665280)) | X_000090 <= exp(X_000089)),
    inference(resolve, [], [refute_0_144, refute_0_143])).

cnf(refute_0_146, plain,
    ((665280 +
      X_000089 *
      (332640 +
       X_000089 *
       (75600 +
        X_000089 *
        (10080 + X_000089 * (840 + X_000089 * (42 + X_000089)))))) /
     (665280 +
      X_000089 *
      (-332640 +
       X_000089 *
       (75600 +
        X_000089 *
        (-10080 + X_000089 * (840 + X_000089 * (-42 + X_000089)))))) <
     X_000090 | X_000089 < 0 | X_000090 <= exp(X_000089)),
    inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_146])).

cnf(refute_0_148, plain,
    (skoXC1 < 0 |
     (665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_147, refute_0_35])).

cnf(refute_0_149, plain,
    ((665280 +
      skoXC1 *
      (332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1)))))) /
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_59, refute_0_148])).

cnf(refute_0_150, plain,
    (665280 +
     skoXC1 *
     (332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (10080 + skoXC1 * (840 + skoXC1 * (42 + skoXC1))))) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (665280 +
      skoXC1 *
      (-332640 +
       skoXC1 *
       (75600 +
        skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))) |
     665280 +
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_142, refute_0_149])).

cnf(refute_0_151, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) |
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1))))) <=
     -665280), inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (-665280 <
     skoXC1 *
     (-332640 +
      skoXC1 *
      (75600 +
       skoXC1 * (-10080 + skoXC1 * (840 + skoXC1 * (-42 + skoXC1)))))),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_153, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))))),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) <=
     0), inference(split, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) <
     0 |
     16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))),
    inference(split, [], [refute_0_153])).

cnf(refute_0_156, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) |
     16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (-21021/25 + skoXC1 * (-20979/500 + skoXC1 * -1001/1000)))))),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    (skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (-21021/25 + skoXC1 * (-20979/500 + skoXC1 * -1001/1000))))) <=
     16648632/25),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_158, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))),
    inference(resolve, [], [refute_0_157, refute_0_156])).

cnf(refute_0_159, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_160, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-10395/4 +
          skoXC1 *
          (10395/8 +
           skoXC1 *
           (-4725/16 +
            skoXC1 *
            (315/8 +
             skoXC1 *
             (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(split, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(split, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (10395/4 +
         skoXC1 *
         (-10395/8 +
          skoXC1 *
          (4725/16 +
           skoXC1 *
           (-315/8 +
            skoXC1 *
            (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-10395/4 +
          skoXC1 *
          (10395/8 +
           skoXC1 *
           (-4725/16 +
            skoXC1 *
            (315/8 +
             skoXC1 *
             (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(split, [], [refute_0_160])).

cnf(refute_0_164, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (10395/4 +
         skoXC1 *
         (-10395/8 +
          skoXC1 *
          (4725/16 +
           skoXC1 *
           (-315/8 +
            skoXC1 *
            (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))) <= 0),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(resolve, [], [refute_0_164, refute_0_162])).

cnf(refute_0_166, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-10395/4 +
         skoXC1 *
         (10395/8 +
          skoXC1 *
          (-4725/16 +
           skoXC1 *
           (315/8 +
            skoXC1 *
            (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256))))))))) <= 0),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_167, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-10395/4 +
            skoXC1 *
            (10395/8 +
             skoXC1 *
             (-4725/16 +
              skoXC1 *
              (315/8 +
               skoXC1 *
               (-105/32 + skoXC1 * (21/128 + skoXC1 * -1/256)))))))))))) <=
     0), inference(resolve, [], [refute_0_166, refute_0_165])).

cnf(refute_0_168, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))),
    inference(resolve, [], [refute_0_167, refute_0_158])).

cnf(refute_0_169, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) <
     exp(1)), inference(resolve, [], [refute_0_154, refute_0_168])).

cnf(refute_0_170, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_171, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_170, refute_0_169])).

cnf(refute_0_172, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_171])).

cnf(refute_0_173, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (10395/4 +
             skoXC1 *
             (-10395/8 +
              skoXC1 *
              (4725/16 +
               skoXC1 *
               (-315/8 +
                skoXC1 *
                (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256))))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) <=
     0), inference(resolve, [], [refute_0_141, refute_0_172])).

cnf(refute_0_174, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (10395/4 +
            skoXC1 *
            (-10395/8 +
             skoXC1 *
             (4725/16 +
              skoXC1 *
              (-315/8 +
               skoXC1 *
               (105/32 + skoXC1 * (-21/128 + skoXC1 * 1/256)))))))))))) <=
     0), inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))))),
    inference(resolve, [], [refute_0_174, refute_0_168])).

cnf(refute_0_176, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) <= 0),
    inference(split, [], [refute_0_175])).

cnf(refute_0_177, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) < 0 |
     16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))),
    inference(split, [], [refute_0_175])).

cnf(refute_0_178, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) |
     16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (-21021/25 + skoXC1 * (-20979/500 + skoXC1 * -1001/1000)))))),
    inference(arithmetic, [], [refute_0_177])).

cnf(refute_0_179, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))),
    inference(resolve, [], [refute_0_157, refute_0_178])).

cnf(refute_0_180, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_181, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-511132545/72356 +
         skoXC1 *
         (511132545/144712 +
          skoXC1 *
          (-232332975/289424 +
           skoXC1 *
           (15488865/144712 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 * (1032591/2315392 + skoXC1 * -49171/4630784))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) <= 0),
    inference(split, [], [refute_0_180])).

cnf(refute_0_182, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) |
     49171/18089 <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_183, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_184, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) |
     49171/18089 <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_185, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776)))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_186, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) |
     49171/18089 <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_187, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_188, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) |
     49171/18089 <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 *
             (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_189, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_190, plain,
    (~ lgen(0, X_000069, exp(X_000068)) | X_000069 <= exp(X_000068)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_191, plain,
    (X_000068 < 0 |
     ~ lgen(0, X_000069,
            (X_000068 ^ 4 + 20 * X_000068 ^ 3 + 180 * X_000068 ^ 2 +
             840 * X_000068 + 1680) /
            (X_000068 ^ 4 - 20 * X_000068 ^ 3 + 180 * X_000068 ^ 2 -
             840 * X_000068 + 1680)) | lgen(0, X_000069, exp(X_000068))),
    inference(subst, [], [exp_lower_bound_cf4])).

cnf(refute_0_192, plain,
    (X_000068 < 0 |
     ~ lgen(0, X_000069,
            (X_000068 ^ 4 + 20 * X_000068 ^ 3 + 180 * X_000068 ^ 2 +
             840 * X_000068 + 1680) /
            (X_000068 ^ 4 - 20 * X_000068 ^ 3 + 180 * X_000068 ^ 2 -
             840 * X_000068 + 1680)) | X_000069 <= exp(X_000068)),
    inference(resolve, [], [refute_0_191, refute_0_190])).

cnf(refute_0_193, plain,
    ((1680 +
      X_000068 * (840 + X_000068 * (180 + X_000068 * (20 + X_000068)))) /
     (1680 +
      X_000068 * (-840 + X_000068 * (180 + X_000068 * (-20 + X_000068)))) <
     X_000069 | X_000068 < 0 | X_000069 <= exp(X_000068)),
    inference(arithmetic, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_194, refute_0_35])).

cnf(refute_0_196, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_59, refute_0_195])).

cnf(refute_0_197, plain,
    (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0), inference(resolve, [], [refute_0_189, refute_0_196])).

cnf(refute_0_198, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (-1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_200, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))))),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 *
             (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) <= 0),
    inference(split, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) <
     0 |
     42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))),
    inference(split, [], [refute_0_200])).

cnf(refute_0_203, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) |
     42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 * (-9009/50 + skoXC1 * (-999/50 + skoXC1 * -1001/1000)))),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (skoXC1 *
     (-20979/25 +
      skoXC1 * (-9009/50 + skoXC1 * (-999/50 + skoXC1 * -1001/1000))) <=
     42042/25),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_205, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))),
    inference(resolve, [], [refute_0_204, refute_0_203])).

cnf(refute_0_206, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_207, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-105/16 +
          skoXC1 *
          (105/32 +
           skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(split, [], [refute_0_206])).

cnf(refute_0_208, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(split, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (105/16 +
         skoXC1 *
         (-105/32 +
          skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-105/16 +
          skoXC1 *
          (105/32 +
           skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(split, [], [refute_0_207])).

cnf(refute_0_211, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (105/16 +
         skoXC1 *
         (-105/32 +
          skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))) <= 0),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(resolve, [], [refute_0_211, refute_0_209])).

cnf(refute_0_213, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-105/16 +
         skoXC1 *
         (105/32 +
          skoXC1 * (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256))))))) <= 0),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_214, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (-105/16 +
            skoXC1 *
            (105/32 +
             skoXC1 *
             (-45/64 + skoXC1 * (5/64 + skoXC1 * -1/256)))))))))) <= 0),
    inference(resolve, [], [refute_0_213, refute_0_212])).

cnf(refute_0_215, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 * (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))),
    inference(resolve, [], [refute_0_214, refute_0_205])).

cnf(refute_0_216, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) <
     exp(1)), inference(resolve, [], [refute_0_201, refute_0_215])).

cnf(refute_0_217, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_218, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_217, refute_0_216])).

cnf(refute_0_219, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (105/16 +
             skoXC1 *
             (-105/32 +
              skoXC1 *
              (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256))))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 *
             (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) <= 0),
    inference(resolve, [], [refute_0_188, refute_0_219])).

cnf(refute_0_221, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (105/16 +
            skoXC1 *
            (-105/32 +
             skoXC1 *
             (45/64 + skoXC1 * (-5/64 + skoXC1 * 1/256)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784))))))))))),
    inference(resolve, [], [refute_0_221, refute_0_215])).

cnf(refute_0_223, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) <=
     0), inference(split, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) <
     0 |
     42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784)))))))))),
    inference(split, [], [refute_0_222])).

cnf(refute_0_225, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) |
     42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 * (-9009/50 + skoXC1 * (-999/50 + skoXC1 * -1001/1000)))),
    inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784)))))))))),
    inference(resolve, [], [refute_0_204, refute_0_225])).

cnf(refute_0_227, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_228, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-5162955/289424 +
         skoXC1 *
         (5162955/578848 +
          skoXC1 *
          (-2212695/1157696 +
           skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) <=
     0), inference(split, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     49171/18089 <=
     (1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_230, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_231, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     49171/18089 <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_232, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_233, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     49171/18089 <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_234, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_235, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     49171/18089 <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_236, plain,
    (~ lgen(0, X_000048, exp(X_000047)) | X_000048 <= exp(X_000047)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_237, plain,
    (X_000047 < -1 | ~ lgen(0, X_000048, 1 + X_000047) |
     lgen(0, X_000048, exp(X_000047))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_238, plain,
    (X_000047 < -1 | ~ lgen(0, X_000048, 1 + X_000047) |
     X_000048 <= exp(X_000047)),
    inference(resolve, [], [refute_0_237, refute_0_236])).

cnf(refute_0_239, plain,
    (1 + X_000047 < X_000048 | X_000047 < -1 | X_000048 <= exp(X_000047)),
    inference(arithmetic, [], [refute_0_238])).

cnf(refute_0_240, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_239])).

cnf(refute_0_241, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_240, refute_0_35])).

cnf(refute_0_242, plain,
    (skoXC1 < -1 |
     1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain, (-1 <= skoXC1),
    inference(decision, [], [refute_0_11, refute_0_59, refute_0_87])).

cnf(refute_0_244, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_243, refute_0_242])).

cnf(refute_0_245, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(split, [], [refute_0_244])).

cnf(refute_0_246, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     exp(1)), inference(resolve, [], [refute_0_245, refute_0_62])).

cnf(refute_0_247, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_248, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_247, refute_0_246])).

cnf(refute_0_249, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_248])).

cnf(refute_0_250, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(resolve, [], [refute_0_235, refute_0_249])).

cnf(refute_0_251, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))) <= 0),
    inference(arithmetic, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))))),
    inference(resolve, [], [refute_0_251, refute_0_62])).

cnf(refute_0_253, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(split, [], [refute_0_252])).

cnf(refute_0_254, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     exp(1)), inference(resolve, [], [refute_0_253, refute_0_90])).

cnf(refute_0_255, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_234, refute_0_254])).

cnf(refute_0_256, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) <
     49171/18089), inference(arithmetic, [], [refute_0_255])).

cnf(refute_0_257, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784)))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(resolve, [], [refute_0_233, refute_0_256])).

cnf(refute_0_258, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 49171/4630784))))) <= 0),
    inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(resolve, [], [refute_0_258, refute_0_90])).

cnf(refute_0_260, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(split, [], [refute_0_259])).

cnf(refute_0_261, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1)), inference(resolve, [], [refute_0_260, refute_0_117])).

cnf(refute_0_262, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_232, refute_0_261])).

cnf(refute_0_263, plain,
    ((1001/1000 + skoXC1) /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     49171/18089), inference(arithmetic, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(resolve, [], [refute_0_231, refute_0_263])).

cnf(refute_0_265, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(arithmetic, [], [refute_0_264])).

cnf(refute_0_266, plain,
    (1001/1000 + skoXC1 <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(resolve, [], [refute_0_265, refute_0_117])).

cnf(refute_0_267, plain,
    ((1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(split, [], [refute_0_266])).

cnf(refute_0_268, plain,
    ((1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1)), inference(resolve, [], [refute_0_267, refute_0_125])).

cnf(refute_0_269, plain,
    (1 < 0 | 7 < 1 |
     (1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_230, refute_0_268])).

cnf(refute_0_270, plain,
    ((1001/1000 + skoXC1) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     49171/18089), inference(arithmetic, [], [refute_0_269])).

cnf(refute_0_271, plain,
    (1001/1000 + skoXC1 <
     49171/18089 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(resolve, [], [refute_0_229, refute_0_270])).

cnf(refute_0_272, plain,
    (1001/1000 <
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(arithmetic, [], [refute_0_271])).

cnf(refute_0_273, plain,
    (1/1000 <
     49171/18089 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     49171/18089 <=
     1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_274, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_275, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     49171/18089 <=
     1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_276, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_277, plain,
    (1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(split, [], [refute_0_106])).

cnf(refute_0_278, plain,
    (1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     exp(1)), inference(resolve, [], [refute_0_277, refute_0_117])).

cnf(refute_0_279, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_276, refute_0_278])).

cnf(refute_0_280, plain,
    (1/1000 /
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) <
     49171/18089), inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (1/1000 <
     49171/18089 *
     (exp(1) *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(resolve, [], [refute_0_275, refute_0_280])).

cnf(refute_0_282, plain,
    (1/1000 <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     exp(1) *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 2417787241/83766251776)))) <=
     0), inference(arithmetic, [], [refute_0_281])).

cnf(refute_0_283, plain,
    (1/1000 <
     exp(1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))))),
    inference(resolve, [], [refute_0_282, refute_0_117])).

cnf(refute_0_284, plain,
    (1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(split, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     exp(1)), inference(resolve, [], [refute_0_284, refute_0_125])).

cnf(refute_0_286, plain,
    (1 < 0 | 7 < 1 |
     1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_274, refute_0_285])).

cnf(refute_0_287, plain,
    (1/1000 /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) <
     49171/18089), inference(arithmetic, [], [refute_0_286])).

cnf(refute_0_288, plain,
    (1/1000 <
     49171/18089 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064)))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(resolve, [], [refute_0_273, refute_0_287])).

cnf(refute_0_289, plain,
    (1/1000 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(arithmetic, [], [refute_0_288])).

cnf(refute_0_290, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) |
     1/1000 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696)))),
    inference(decision, [], [refute_0_11, refute_0_59, refute_0_87])).

cnf(refute_0_291, plain,
    (1/1000 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696)))),
    inference(resolve, [], [refute_0_289, refute_0_290])).

cnf(refute_0_292, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) |
     1001/1000 <
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696)))),
    inference(decision, [],
              [refute_0_87, refute_0_59, refute_0_11, refute_0_291])).

cnf(refute_0_293, plain,
    (1001/1000 <
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 * (skoXC1 * 5845695142742392081/27409316158594621696)))),
    inference(resolve, [], [refute_0_272, refute_0_292])).

cnf(refute_0_294, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) |
     49171/18089 <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_295, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_296, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) |
     49171/18089 <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_297, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776)))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_298, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) |
     49171/18089 <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_299, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_300, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) |
     49171/18089 <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_301, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (12 + skoXC1 * (-6 + skoXC1)) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) |
     12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_302, plain,
    (~ lgen(0, X_000055, exp(X_000054)) | X_000055 <= exp(X_000054)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_303, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            (X_000054 ^ 2 + 6 * X_000054 + 12) /
            (X_000054 ^ 2 - 6 * X_000054 + 12)) |
     lgen(0, X_000055, exp(X_000054))),
    inference(subst, [], [exp_lower_bound_cf2])).

cnf(refute_0_304, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            (X_000054 ^ 2 + 6 * X_000054 + 12) /
            (X_000054 ^ 2 - 6 * X_000054 + 12)) |
     X_000055 <= exp(X_000054)),
    inference(resolve, [], [refute_0_303, refute_0_302])).

cnf(refute_0_305, plain,
    ((12 + X_000054 * (6 + X_000054)) / (12 + X_000054 * (-6 + X_000054)) <
     X_000055 | X_000054 < 0 | X_000055 <= exp(X_000054)),
    inference(arithmetic, [], [refute_0_304])).

cnf(refute_0_306, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) |
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256))))))) <=
     exp(skoXC1)), inference(subst, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (skoXC1 < 0 |
     (12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_306, refute_0_35])).

cnf(refute_0_308, plain,
    ((12 + skoXC1 * (6 + skoXC1)) / (12 + skoXC1 * (-6 + skoXC1)) <
     -1/1000 +
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))),
    inference(resolve, [], [refute_0_59, refute_0_307])).

cnf(refute_0_309, plain,
    (12 + skoXC1 * (6 + skoXC1) <
     (-1/1000 +
      exp(1) *
      (exp(1) *
       (exp(1) *
        (exp(1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/256)))))))) *
     (12 + skoXC1 * (-6 + skoXC1)) | 12 + skoXC1 * (-6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_301, refute_0_308])).

cnf(refute_0_310, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) |
     skoXC1 * (-6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_309])).

cnf(refute_0_311, plain, (-12 < skoXC1 * (-6 + skoXC1)),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_312, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))))),
    inference(resolve, [], [refute_0_310, refute_0_311])).

cnf(refute_0_313, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) <=
     0), inference(split, [], [refute_0_312])).

cnf(refute_0_314, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) < 0 |
     3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))),
    inference(split, [], [refute_0_312])).

cnf(refute_0_315, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) |
     3003/250 < skoXC1 * (-2997/500 + skoXC1 * -1001/1000)),
    inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (skoXC1 * (-2997/500 + skoXC1 * -1001/1000) <= 3003/250),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_317, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))),
    inference(resolve, [], [refute_0_316, refute_0_315])).

cnf(refute_0_318, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_319, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(split, [], [refute_0_318])).

cnf(refute_0_320, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) <
     0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(split, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(arithmetic, [], [refute_0_320])).

cnf(refute_0_322, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(split, [], [refute_0_319])).

cnf(refute_0_323, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))) <=
     0), inference(arithmetic, [], [refute_0_322])).

cnf(refute_0_324, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(resolve, [], [refute_0_323, refute_0_321])).

cnf(refute_0_325, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256))))) <= 0),
    inference(decision, [], [refute_0_87, refute_0_59, refute_0_11])).

cnf(refute_0_326, plain,
    (exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (-3/64 + skoXC1 * (3/128 + skoXC1 * -1/256)))))))) <=
     0), inference(resolve, [], [refute_0_325, refute_0_324])).

cnf(refute_0_327, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))),
    inference(resolve, [], [refute_0_326, refute_0_317])).

cnf(refute_0_328, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) <
     exp(1)), inference(resolve, [], [refute_0_313, refute_0_327])).

cnf(refute_0_329, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_330, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_329, refute_0_328])).

cnf(refute_0_331, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_330])).

cnf(refute_0_332, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (exp(1) *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256))))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) <=
     0), inference(resolve, [], [refute_0_300, refute_0_331])).

cnf(refute_0_333, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) |
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (3/64 + skoXC1 * (-3/128 + skoXC1 * 1/256)))))))) <=
     0), inference(arithmetic, [], [refute_0_332])).

cnf(refute_0_334, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))))),
    inference(resolve, [], [refute_0_333, refute_0_327])).

cnf(refute_0_335, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) <
     exp(1) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) <= 0),
    inference(split, [], [refute_0_334])).

cnf(refute_0_336, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) < 0 |
     3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) < 0 |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))),
    inference(split, [], [refute_0_334])).

cnf(refute_0_337, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) |
     3003/250 < skoXC1 * (-2997/500 + skoXC1 * -1001/1000)),
    inference(arithmetic, [], [refute_0_336])).

cnf(refute_0_338, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) |
     0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))),
    inference(resolve, [], [refute_0_316, refute_0_337])).

cnf(refute_0_339, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_340, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-147513/1157696 +
         skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) <= 0),
    inference(split, [], [refute_0_339])).

cnf(refute_0_341, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-147513/1157696 +
         skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))) <= 0),
    inference(decision, [],
              [refute_0_291, refute_0_11, refute_0_59, refute_0_87,
               refute_0_293])).

cnf(refute_0_342, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-147513/1157696 +
           skoXC1 * (147513/2315392 + skoXC1 * -49171/4630784))))))) <= 0),
    inference(resolve, [], [refute_0_341, refute_0_340])).

cnf(refute_0_343, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))),
    inference(resolve, [], [refute_0_342, refute_0_338])).

cnf(refute_0_344, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) <
     exp(1)), inference(resolve, [], [refute_0_335, refute_0_343])).

cnf(refute_0_345, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_299, refute_0_344])).

cnf(refute_0_346, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_345])).

cnf(refute_0_347, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (147513/1157696 +
            skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784)))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) <= 0),
    inference(resolve, [], [refute_0_298, refute_0_346])).

cnf(refute_0_348, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (147513/1157696 +
           skoXC1 * (-147513/2315392 + skoXC1 * 49171/4630784))))))) <= 0),
    inference(arithmetic, [], [refute_0_347])).

cnf(refute_0_349, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776)))))))),
    inference(resolve, [], [refute_0_348, refute_0_343])).

cnf(refute_0_350, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) < exp(1) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) <= 0),
    inference(split, [], [refute_0_349])).

cnf(refute_0_351, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) < 0 |
     3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) < 0 |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776))))))),
    inference(split, [], [refute_0_349])).

cnf(refute_0_352, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) |
     3003/250 < skoXC1 * (-2997/500 + skoXC1 * -1001/1000)),
    inference(arithmetic, [], [refute_0_351])).

cnf(refute_0_353, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776))))))),
    inference(resolve, [], [refute_0_316, refute_0_352])).

cnf(refute_0_354, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_355, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) <
     0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    inference(split, [], [refute_0_354])).

cnf(refute_0_356, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (7253361723/20941562944 +
         skoXC1 *
         (-7253361723/41883125888 + skoXC1 * 2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    inference(arithmetic, [], [refute_0_355])).

cnf(refute_0_357, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    inference(split, [], [refute_0_354])).

cnf(refute_0_358, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (7253361723/20941562944 +
         skoXC1 *
         (-7253361723/41883125888 + skoXC1 * 2417787241/83766251776))))) <=
     0), inference(arithmetic, [], [refute_0_357])).

cnf(refute_0_359, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    inference(resolve, [], [refute_0_358, refute_0_356])).

cnf(refute_0_360, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-7253361723/20941562944 +
         skoXC1 *
         (7253361723/41883125888 + skoXC1 * -2417787241/83766251776))))) <=
     0),
    inference(decision, [],
              [refute_0_293, refute_0_87, refute_0_59, refute_0_11,
               refute_0_291])).

cnf(refute_0_361, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-7253361723/20941562944 +
          skoXC1 *
          (7253361723/41883125888 +
           skoXC1 * -2417787241/83766251776)))))) <= 0),
    inference(resolve, [], [refute_0_360, refute_0_359])).

cnf(refute_0_362, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776))))))),
    inference(resolve, [], [refute_0_361, refute_0_353])).

cnf(refute_0_363, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) < exp(1)),
    inference(resolve, [], [refute_0_350, refute_0_362])).

cnf(refute_0_364, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_297, refute_0_363])).

cnf(refute_0_365, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) < 49171/18089),
    inference(arithmetic, [], [refute_0_364])).

cnf(refute_0_366, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (7253361723/20941562944 +
           skoXC1 *
           (-7253361723/41883125888 +
            skoXC1 * 2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) <= 0),
    inference(resolve, [], [refute_0_296, refute_0_365])).

cnf(refute_0_367, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (7253361723/20941562944 +
          skoXC1 *
          (-7253361723/41883125888 +
           skoXC1 * 2417787241/83766251776)))))) <= 0),
    inference(arithmetic, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064))))))),
    inference(resolve, [], [refute_0_367, refute_0_362])).

cnf(refute_0_369, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) < exp(1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064))))) <= 0),
    inference(split, [], [refute_0_368])).

cnf(refute_0_370, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064)))))),
    inference(decision, [],
              [refute_0_291, refute_0_11, refute_0_59, refute_0_87,
               refute_0_293])).

cnf(refute_0_371, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) < exp(1)),
    inference(resolve, [], [refute_0_369, refute_0_370])).

cnf(refute_0_372, plain,
    (1 < 0 | 7 < 1 |
     (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_295, refute_0_371])).

cnf(refute_0_373, plain,
    ((3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000)) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) < 49171/18089),
    inference(arithmetic, [], [refute_0_372])).

cnf(refute_0_374, plain,
    (3003/250 + skoXC1 * (2997/500 + skoXC1 * 1001/1000) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (356655049281633/378811932094016 +
          skoXC1 *
          (-356655049281633/757623864188032 +
           skoXC1 * 118885016427211/1515247728376064)))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064))))) <= 0),
    inference(resolve, [], [refute_0_294, refute_0_373])).

cnf(refute_0_375, plain,
    (3003/250 <
     skoXC1 *
     (-2997/500 +
      skoXC1 *
      (-1001/1000 +
       skoXC1 *
       (skoXC1 *
        (17537085428227176243/6852329039648655424 +
         skoXC1 *
         (-17537085428227176243/13704658079297310848 +
          skoXC1 * 5845695142742392081/27409316158594621696))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064))))) <= 0),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (356655049281633/378811932094016 +
         skoXC1 *
         (-356655049281633/757623864188032 +
          skoXC1 * 118885016427211/1515247728376064))))) |
     3003/250 <
     skoXC1 *
     (-2997/500 +
      skoXC1 *
      (-1001/1000 +
       skoXC1 *
       (skoXC1 *
        (17537085428227176243/6852329039648655424 +
         skoXC1 *
         (-17537085428227176243/13704658079297310848 +
          skoXC1 * 5845695142742392081/27409316158594621696)))))),
    inference(decision, [],
              [refute_0_291, refute_0_11, refute_0_59, refute_0_87,
               refute_0_293])).

cnf(refute_0_377, plain,
    (3003/250 <
     skoXC1 *
     (-2997/500 +
      skoXC1 *
      (-1001/1000 +
       skoXC1 *
       (skoXC1 *
        (17537085428227176243/6852329039648655424 +
         skoXC1 *
         (-17537085428227176243/13704658079297310848 +
          skoXC1 * 5845695142742392081/27409316158594621696)))))),
    inference(resolve, [], [refute_0_375, refute_0_376])).

cnf(refute_0_378, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-5162955/289424 +
         skoXC1 *
         (5162955/578848 +
          skoXC1 *
          (-2212695/1157696 +
           skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))) <= 0),
    inference(decision, [],
              [refute_0_293, refute_0_87, refute_0_59, refute_0_11,
               refute_0_291, refute_0_377])).

cnf(refute_0_379, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-5162955/289424 +
           skoXC1 *
           (5162955/578848 +
            skoXC1 *
            (-2212695/1157696 +
             skoXC1 * (245855/1157696 + skoXC1 * -49171/4630784))))))))) <=
     0), inference(resolve, [], [refute_0_378, refute_0_228])).

cnf(refute_0_380, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784)))))))))),
    inference(resolve, [], [refute_0_379, refute_0_226])).

cnf(refute_0_381, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) <
     exp(1)), inference(resolve, [], [refute_0_223, refute_0_380])).

cnf(refute_0_382, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_187, refute_0_381])).

cnf(refute_0_383, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_382])).

cnf(refute_0_384, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (5162955/289424 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (2212695/1157696 +
              skoXC1 *
              (-245855/1157696 + skoXC1 * 49171/4630784)))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) <=
     0), inference(resolve, [], [refute_0_186, refute_0_383])).

cnf(refute_0_385, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (5162955/289424 +
           skoXC1 *
           (-5162955/578848 +
            skoXC1 *
            (2212695/1157696 +
             skoXC1 * (-245855/1157696 + skoXC1 * 49171/4630784))))))))) <=
     0), inference(arithmetic, [], [refute_0_384])).

cnf(refute_0_386, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776)))))))))),
    inference(resolve, [], [refute_0_385, refute_0_380])).

cnf(refute_0_387, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) < exp(1) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) <= 0),
    inference(split, [], [refute_0_386])).

cnf(refute_0_388, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) < 0 |
     42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) < 0 |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776))))))))),
    inference(split, [], [refute_0_386])).

cnf(refute_0_389, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) |
     42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 * (-9009/50 + skoXC1 * (-999/50 + skoXC1 * -1001/1000)))),
    inference(arithmetic, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776))))))))),
    inference(resolve, [], [refute_0_204, refute_0_389])).

cnf(refute_0_391, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_392, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    inference(split, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (253867660305/5235390736 +
         skoXC1 *
         (-253867660305/10470781472 +
          skoXC1 *
          (108800425845/20941562944 +
           skoXC1 *
           (-12088936205/20941562944 +
            skoXC1 * 2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    inference(arithmetic, [], [refute_0_392])).

cnf(refute_0_394, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    inference(split, [], [refute_0_391])).

cnf(refute_0_395, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (253867660305/5235390736 +
         skoXC1 *
         (-253867660305/10470781472 +
          skoXC1 *
          (108800425845/20941562944 +
           skoXC1 *
           (-12088936205/20941562944 +
            skoXC1 * 2417787241/83766251776))))))) <= 0),
    inference(arithmetic, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    inference(resolve, [], [refute_0_395, refute_0_393])).

cnf(refute_0_397, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-253867660305/5235390736 +
         skoXC1 *
         (253867660305/10470781472 +
          skoXC1 *
          (-108800425845/20941562944 +
           skoXC1 *
           (12088936205/20941562944 +
            skoXC1 * -2417787241/83766251776))))))) <= 0),
    inference(decision, [],
              [refute_0_377, refute_0_291, refute_0_11, refute_0_59,
               refute_0_87, refute_0_293])).

cnf(refute_0_398, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-253867660305/5235390736 +
          skoXC1 *
          (253867660305/10470781472 +
           skoXC1 *
           (-108800425845/20941562944 +
            skoXC1 *
            (12088936205/20941562944 +
             skoXC1 * -2417787241/83766251776)))))))) <= 0),
    inference(resolve, [], [refute_0_397, refute_0_396])).

cnf(refute_0_399, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776))))))))),
    inference(resolve, [], [refute_0_398, refute_0_390])).

cnf(refute_0_400, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) < exp(1)),
    inference(resolve, [], [refute_0_387, refute_0_399])).

cnf(refute_0_401, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_185, refute_0_400])).

cnf(refute_0_402, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) < 49171/18089),
    inference(arithmetic, [], [refute_0_401])).

cnf(refute_0_403, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (253867660305/5235390736 +
           skoXC1 *
           (-253867660305/10470781472 +
            skoXC1 *
            (108800425845/20941562944 +
             skoXC1 *
             (-12088936205/20941562944 +
              skoXC1 * 2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) <= 0),
    inference(resolve, [], [refute_0_184, refute_0_402])).

cnf(refute_0_404, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (253867660305/5235390736 +
          skoXC1 *
          (-253867660305/10470781472 +
           skoXC1 *
           (108800425845/20941562944 +
            skoXC1 *
            (-12088936205/20941562944 +
             skoXC1 * 2417787241/83766251776)))))))) <= 0),
    inference(arithmetic, [], [refute_0_403])).

cnf(refute_0_405, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064))))))))),
    inference(resolve, [], [refute_0_404, refute_0_399])).

cnf(refute_0_406, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) < exp(1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064))))))) <= 0),
    inference(split, [], [refute_0_405])).

cnf(refute_0_407, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064)))))))),
    inference(decision, [],
              [refute_0_293, refute_0_87, refute_0_59, refute_0_11,
               refute_0_291, refute_0_377])).

cnf(refute_0_408, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) < exp(1)),
    inference(resolve, [], [refute_0_406, refute_0_407])).

cnf(refute_0_409, plain,
    (1 < 0 | 7 < 1 |
     (42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_183, refute_0_408])).

cnf(refute_0_410, plain,
    ((42042/25 +
      skoXC1 *
      (20979/25 +
       skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000)))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_409])).

cnf(refute_0_411, plain,
    (42042/25 +
     skoXC1 *
     (20979/25 +
      skoXC1 * (9009/50 + skoXC1 * (999/50 + skoXC1 * 1001/1000))) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (12482926724857155/94702983023504 +
          skoXC1 *
          (-12482926724857155/189405966047008 +
           skoXC1 *
           (5349825739224495/378811932094016 +
            skoXC1 *
            (-594425082136055/378811932094016 +
             skoXC1 * 118885016427211/1515247728376064)))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064))))))) <= 0),
    inference(resolve, [], [refute_0_182, refute_0_410])).

cnf(refute_0_412, plain,
    (42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 *
      (-9009/50 +
       skoXC1 *
       (-999/50 +
        skoXC1 *
        (76510399330722386560643/214135282489020482000 +
         skoXC1 *
         (-613797989987951168505/3426164519824327712 +
          skoXC1 *
          (263056281423407643645/6852329039648655424 +
           skoXC1 *
           (-29228475713711960405/6852329039648655424 +
            skoXC1 * 5845695142742392081/27409316158594621696))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064))))))) <= 0),
    inference(arithmetic, [], [refute_0_411])).

cnf(refute_0_413, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (12482926724857155/94702983023504 +
         skoXC1 *
         (-12482926724857155/189405966047008 +
          skoXC1 *
          (5349825739224495/378811932094016 +
           skoXC1 *
           (-594425082136055/378811932094016 +
            skoXC1 * 118885016427211/1515247728376064))))))) |
     42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 *
      (-9009/50 +
       skoXC1 *
       (-999/50 +
        skoXC1 *
        (76510399330722386560643/214135282489020482000 +
         skoXC1 *
         (-613797989987951168505/3426164519824327712 +
          skoXC1 *
          (263056281423407643645/6852329039648655424 +
           skoXC1 *
           (-29228475713711960405/6852329039648655424 +
            skoXC1 * 5845695142742392081/27409316158594621696)))))))),
    inference(decision, [],
              [refute_0_293, refute_0_87, refute_0_59, refute_0_11,
               refute_0_291, refute_0_377])).

cnf(refute_0_414, plain,
    (42042/25 <
     skoXC1 *
     (-20979/25 +
      skoXC1 *
      (-9009/50 +
       skoXC1 *
       (-999/50 +
        skoXC1 *
        (76510399330722386560643/214135282489020482000 +
         skoXC1 *
         (-613797989987951168505/3426164519824327712 +
          skoXC1 *
          (263056281423407643645/6852329039648655424 +
           skoXC1 *
           (-29228475713711960405/6852329039648655424 +
            skoXC1 * 5845695142742392081/27409316158594621696)))))))),
    inference(resolve, [], [refute_0_412, refute_0_413])).

cnf(refute_0_415, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-511132545/72356 +
         skoXC1 *
         (511132545/144712 +
          skoXC1 *
          (-232332975/289424 +
           skoXC1 *
           (15488865/144712 +
            skoXC1 *
            (-5162955/578848 +
             skoXC1 *
             (1032591/2315392 + skoXC1 * -49171/4630784))))))))) <= 0),
    inference(decision, [],
              [refute_0_414, refute_0_293, refute_0_87, refute_0_59,
               refute_0_11, refute_0_291, refute_0_377])).

cnf(refute_0_416, plain,
    (exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-511132545/72356 +
           skoXC1 *
           (511132545/144712 +
            skoXC1 *
            (-232332975/289424 +
             skoXC1 *
             (15488865/144712 +
              skoXC1 *
              (-5162955/578848 +
               skoXC1 *
               (1032591/2315392 + skoXC1 * -49171/4630784))))))))))) <= 0),
    inference(resolve, [], [refute_0_415, refute_0_181])).

cnf(refute_0_417, plain,
    (0 <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))),
    inference(resolve, [], [refute_0_416, refute_0_179])).

cnf(refute_0_418, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) <
     exp(1)), inference(resolve, [], [refute_0_176, refute_0_417])).

cnf(refute_0_419, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_140, refute_0_418])).

cnf(refute_0_420, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_419])).

cnf(refute_0_421, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (exp(1) *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (511132545/72356 +
            skoXC1 *
            (-511132545/144712 +
             skoXC1 *
             (232332975/289424 +
              skoXC1 *
              (-15488865/144712 +
               skoXC1 *
               (5162955/578848 +
                skoXC1 *
                (-1032591/2315392 + skoXC1 * 49171/4630784)))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) <= 0),
    inference(resolve, [], [refute_0_139, refute_0_420])).

cnf(refute_0_422, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) |
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (511132545/72356 +
           skoXC1 *
           (-511132545/144712 +
            skoXC1 *
            (232332975/289424 +
             skoXC1 *
             (-15488865/144712 +
              skoXC1 *
              (5162955/578848 +
               skoXC1 *
               (-1032591/2315392 + skoXC1 * 49171/4630784))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_421])).

cnf(refute_0_423, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776)))))))))))),
    inference(resolve, [], [refute_0_422, refute_0_417])).

cnf(refute_0_424, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) < exp(1) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) <= 0),
    inference(split, [], [refute_0_423])).

cnf(refute_0_425, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) < 0 |
     16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) < 0 |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776))))))))))),
    inference(split, [], [refute_0_423])).

cnf(refute_0_426, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) |
     16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (-21021/25 + skoXC1 * (-20979/500 + skoXC1 * -1001/1000)))))),
    inference(arithmetic, [], [refute_0_425])).

cnf(refute_0_427, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) |
     0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776))))))))))),
    inference(resolve, [], [refute_0_157, refute_0_426])).

cnf(refute_0_428, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_429, plain,
    (0 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) < 0 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    inference(split, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (25132898370195/1308847684 +
         skoXC1 *
         (-25132898370195/2617695368 +
          skoXC1 *
          (11424044713725/5235390736 +
           skoXC1 *
           (-761602980915/2617695368 +
            skoXC1 *
            (253867660305/10470781472 +
             skoXC1 *
             (-50773532061/41883125888 +
              skoXC1 * 2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_429])).

cnf(refute_0_431, plain,
    (exp(1) <
     0 /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    inference(split, [], [refute_0_428])).

cnf(refute_0_432, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (25132898370195/1308847684 +
         skoXC1 *
         (-25132898370195/2617695368 +
          skoXC1 *
          (11424044713725/5235390736 +
           skoXC1 *
           (-761602980915/2617695368 +
            skoXC1 *
            (253867660305/10470781472 +
             skoXC1 *
             (-50773532061/41883125888 +
              skoXC1 * 2417787241/83766251776))))))))) <= 0),
    inference(arithmetic, [], [refute_0_431])).

cnf(refute_0_433, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    inference(resolve, [], [refute_0_432, refute_0_430])).

cnf(refute_0_434, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-25132898370195/1308847684 +
         skoXC1 *
         (25132898370195/2617695368 +
          skoXC1 *
          (-11424044713725/5235390736 +
           skoXC1 *
           (761602980915/2617695368 +
            skoXC1 *
            (-253867660305/10470781472 +
             skoXC1 *
             (50773532061/41883125888 +
              skoXC1 * -2417787241/83766251776))))))))) <= 0),
    inference(decision, [],
              [refute_0_377, refute_0_291, refute_0_11, refute_0_59,
               refute_0_87, refute_0_293, refute_0_414])).

cnf(refute_0_435, plain,
    (exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-25132898370195/1308847684 +
          skoXC1 *
          (25132898370195/2617695368 +
           skoXC1 *
           (-11424044713725/5235390736 +
            skoXC1 *
            (761602980915/2617695368 +
             skoXC1 *
             (-253867660305/10470781472 +
              skoXC1 *
              (50773532061/41883125888 +
               skoXC1 * -2417787241/83766251776)))))))))) <= 0),
    inference(resolve, [], [refute_0_434, refute_0_433])).

cnf(refute_0_436, plain,
    (0 <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776))))))))))),
    inference(resolve, [], [refute_0_435, refute_0_427])).

cnf(refute_0_437, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) < exp(1)),
    inference(resolve, [], [refute_0_424, refute_0_436])).

cnf(refute_0_438, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_138, refute_0_437])).

cnf(refute_0_439, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) < 49171/18089),
    inference(arithmetic, [], [refute_0_438])).

cnf(refute_0_440, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (exp(1) *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (25132898370195/1308847684 +
           skoXC1 *
           (-25132898370195/2617695368 +
            skoXC1 *
            (11424044713725/5235390736 +
             skoXC1 *
             (-761602980915/2617695368 +
              skoXC1 *
              (253867660305/10470781472 +
               skoXC1 *
               (-50773532061/41883125888 +
                skoXC1 * 2417787241/83766251776))))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) <= 0),
    inference(resolve, [], [refute_0_137, refute_0_439])).

cnf(refute_0_441, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (25132898370195/1308847684 +
          skoXC1 *
          (-25132898370195/2617695368 +
           skoXC1 *
           (11424044713725/5235390736 +
            skoXC1 *
            (-761602980915/2617695368 +
             skoXC1 *
             (253867660305/10470781472 +
              skoXC1 *
              (-50773532061/41883125888 +
               skoXC1 * 2417787241/83766251776)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_440])).

cnf(refute_0_442, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064))))))))))),
    inference(resolve, [], [refute_0_441, refute_0_436])).

cnf(refute_0_443, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     exp(1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(split, [], [refute_0_442])).

cnf(refute_0_444, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064)))))))))),
    inference(decision, [],
              [refute_0_414, refute_0_293, refute_0_87, refute_0_59,
               refute_0_11, refute_0_291, refute_0_377])).

cnf(refute_0_445, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     exp(1)), inference(resolve, [], [refute_0_443, refute_0_444])).

cnf(refute_0_446, plain,
    (1 < 0 | 19/2 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     exp(1) <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064))))))))))),
    inference(subst, [], [refute_0_130])).

cnf(refute_0_447, plain,
    (1 < 0 | 19/2 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1)))))))),
    inference(resolve, [], [refute_0_446, refute_0_445])).

cnf(refute_0_448, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     28245729/10391023), inference(arithmetic, [], [refute_0_447])).

cnf(refute_0_449, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     28245729/10391023 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(resolve, [], [refute_0_136, refute_0_448])).

cnf(refute_0_450, plain,
    (16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (867487193445889411924430517/6150380467286497528700 +
         skoXC1 *
         (-4365873973426477021165405071/61503804672864975287000 +
          skoXC1 *
          (1983192049742141720925072301/123007609345729950574000 +
           skoXC1 *
           (-1057768096191518291522985/492030437382919802296 +
            skoXC1 *
            (352589365397172763840995/1968121749531679209184 +
             skoXC1 *
             (-70517873079434552768199/7872486998126716836736 +
              skoXC1 *
              3357993956163550131819/15744973996253433673472))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(arithmetic, [], [refute_0_449])).

cnf(refute_0_451, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     49171/18089 <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_452, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064))))))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_453, plain,
    (1 < 0 | 7 < 1 |
     (16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1)))))),
    inference(resolve, [], [refute_0_452, refute_0_445])).

cnf(refute_0_454, plain,
    ((16648632/25 +
      skoXC1 *
      (8307684/25 +
       skoXC1 *
       (378378/5 +
        skoXC1 *
        (251748/25 +
         skoXC1 *
         (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000)))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) <
     49171/18089), inference(arithmetic, [], [refute_0_453])).

cnf(refute_0_455, plain,
    (16648632/25 +
     skoXC1 *
     (8307684/25 +
      skoXC1 *
      (378378/5 +
       skoXC1 *
       (251748/25 +
        skoXC1 *
        (21021/25 + skoXC1 * (20979/500 + skoXC1 * 1001/1000))))) <
     49171/18089 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (1235809745760858345/23675745755876 +
          skoXC1 *
          (-1235809745760858345/47351491511752 +
           skoXC1 *
           (561731702618571975/94702983023504 +
            skoXC1 *
            (-37448780174571465/47351491511752 +
             skoXC1 *
             (12482926724857155/189405966047008 +
              skoXC1 *
              (-2496585344971431/757623864188032 +
               skoXC1 * 118885016427211/1515247728376064)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(resolve, [], [refute_0_451, refute_0_454])).

cnf(refute_0_456, plain,
    (16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (1510147349673775742945631/10706764124451024100 +
         skoXC1 *
         (-7600242470192232870941253/107067641244510241000 +
          skoXC1 *
          (3452399344264453813338143/214135282489020482000 +
           skoXC1 *
           (-1841393969963853505515/856541129956081928 +
            skoXC1 *
            (613797989987951168505/3426164519824327712 +
             skoXC1 *
             (-122759597997590233701/13704658079297310848 +
              skoXC1 * 5845695142742392081/27409316158594621696))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) <= 0),
    inference(arithmetic, [], [refute_0_455])).

cnf(refute_0_457, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) |
     16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (1510147349673775742945631/10706764124451024100 +
         skoXC1 *
         (-7600242470192232870941253/107067641244510241000 +
          skoXC1 *
          (3452399344264453813338143/214135282489020482000 +
           skoXC1 *
           (-1841393969963853505515/856541129956081928 +
            skoXC1 *
            (613797989987951168505/3426164519824327712 +
             skoXC1 *
             (-122759597997590233701/13704658079297310848 +
              skoXC1 * 5845695142742392081/27409316158594621696)))))))))),
    inference(decision, [],
              [refute_0_414, refute_0_293, refute_0_87, refute_0_59,
               refute_0_11, refute_0_291, refute_0_377])).

cnf(refute_0_458, plain,
    (16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (1510147349673775742945631/10706764124451024100 +
         skoXC1 *
         (-7600242470192232870941253/107067641244510241000 +
          skoXC1 *
          (3452399344264453813338143/214135282489020482000 +
           skoXC1 *
           (-1841393969963853505515/856541129956081928 +
            skoXC1 *
            (613797989987951168505/3426164519824327712 +
             skoXC1 *
             (-122759597997590233701/13704658079297310848 +
              skoXC1 * 5845695142742392081/27409316158594621696)))))))))),
    inference(resolve, [], [refute_0_456, refute_0_457])).

cnf(refute_0_459, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (1235809745760858345/23675745755876 +
         skoXC1 *
         (-1235809745760858345/47351491511752 +
          skoXC1 *
          (561731702618571975/94702983023504 +
           skoXC1 *
           (-37448780174571465/47351491511752 +
            skoXC1 *
            (12482926724857155/189405966047008 +
             skoXC1 *
             (-2496585344971431/757623864188032 +
              skoXC1 * 118885016427211/1515247728376064))))))))) |
     16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (867487193445889411924430517/6150380467286497528700 +
         skoXC1 *
         (-4365873973426477021165405071/61503804672864975287000 +
          skoXC1 *
          (1983192049742141720925072301/123007609345729950574000 +
           skoXC1 *
           (-1057768096191518291522985/492030437382919802296 +
            skoXC1 *
            (352589365397172763840995/1968121749531679209184 +
             skoXC1 *
             (-70517873079434552768199/7872486998126716836736 +
              skoXC1 *
              3357993956163550131819/15744973996253433673472)))))))))),
    inference(decision, [],
              [refute_0_458, refute_0_414, refute_0_293, refute_0_87,
               refute_0_59, refute_0_11, refute_0_291, refute_0_377])).

cnf(refute_0_460, plain,
    (16648632/25 <
     skoXC1 *
     (-8307684/25 +
      skoXC1 *
      (-378378/5 +
       skoXC1 *
       (-251748/25 +
        skoXC1 *
        (867487193445889411924430517/6150380467286497528700 +
         skoXC1 *
         (-4365873973426477021165405071/61503804672864975287000 +
          skoXC1 *
          (1983192049742141720925072301/123007609345729950574000 +
           skoXC1 *
           (-1057768096191518291522985/492030437382919802296 +
            skoXC1 *
            (352589365397172763840995/1968121749531679209184 +
             skoXC1 *
             (-70517873079434552768199/7872486998126716836736 +
              skoXC1 *
              3357993956163550131819/15744973996253433673472)))))))))),
    inference(resolve, [], [refute_0_450, refute_0_459])).

cnf(refute_0_461, plain,
    (skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0 |
     skoXC1 *
     (-1 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (-1/6 +
        skoXC1 *
        (8105860118958971186273/47234921988760301020416 +
         skoXC1 *
         (-1/120 +
          skoXC1 *
          (-121/87480 +
           skoXC1 *
           (-17/87480 +
            skoXC1 *
            (-49/2099520 +
             skoXC1 *
             (-409/170061120 +
              skoXC1 *
              (-361/1700611200 +
               skoXC1 *
               (-1/62985600 +
                skoXC1 *
                (-181/183666009600 +
                 skoXC1 *
                 (-1/20407334400 +
                  skoXC1 *
                  (-1/550998028800 +
                   skoXC1 * -1/24794911296000)))))))))))))) <= 1001/1000),
    inference(decision, [],
              [refute_0_460, refute_0_458, refute_0_414, refute_0_293,
               refute_0_87, refute_0_59, refute_0_11, refute_0_291,
               refute_0_377])).

cnf(refute_0_462, plain,
    (skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * 118885016427211/1515247728376064))) <=
     0), inference(resolve, [], [refute_0_461, refute_0_135])).

cnf(refute_0_463, plain, ($false),
    inference(resolve, [], [refute_0_462, refute_0_125])).
% SZS output end CNFRefutation for exp-361-4.tptp
