% SZS output start CNFRefutation for exp-hyperbolic-2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_lower_bound_cf4, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 4 + 20 * X ^ 3 + 180 * X ^ 2 + 840 * X + 1680) /
            (X ^ 4 - 20 * X ^ 3 + 180 * X ^ 2 - 840 * X + 1680)) |
     lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf4, axiom,
    (0 <= X |
     ~ lgen(R,
            (X ^ 4 + 20 * X ^ 3 + 180 * X ^ 2 + 840 * X + 1680) /
            (X ^ 4 - 20 * X ^ 3 + 180 * X ^ 2 - 840 * X + 1680), Y) |
     lgen(R, exp(X), Y))).

fof(exp_hyp, conjecture,
    (! [X] :
       (0 < X =>
        X <
        (3 + X ^ 2 / 10) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 10)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 < X =>
        X <
        (3 + X ^ 2 / 10) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 10))),
    inference(strip, [], [exp_hyp])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 < X =>
          X <
          (3 + X ^ 2 / 10) * ((exp(X) - exp(-X)) / 2) /
          (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 10))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X &
        (3 + X ^ 2 / 10) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 10) <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 < skoXC1 &
     (3 + skoXC1 ^ 2 / 10) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 10) <= skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    ((3 + skoXC1 ^ 2 / 10) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 10) <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <
     1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) |
     (1 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) |
     2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_1, plain,
    (X_000017 < exp(X_000016) | ~ lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain,
    (~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3) |
     lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_3, plain,
    (X_000017 < exp(X_000016) |
     ~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000017 < exp(X_000016) |
     1 +
     X_000016 *
     (1 +
      X_000016 *
      (1/2 +
       X_000016 *
       (1/6 +
        X_000016 *
        (1/24 +
         X_000016 *
         (1/120 +
          X_000016 *
          (121/87480 +
           X_000016 *
           (17/87480 +
            X_000016 *
            (49/2099520 +
             X_000016 *
             (409/170061120 +
              X_000016 *
              (361/1700611200 +
               X_000016 *
               (1/62985600 +
                X_000016 *
                (181/183666009600 +
                 X_000016 *
                 (1/20407334400 +
                  X_000016 *
                  (1/550998028800 +
                   X_000016 * 1/24794911296000)))))))))))))) <= X_000017),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <
     exp(skoXC1) |
     1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20)) <
     (-2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (66 +
        skoXC1 *
        (-18 + skoXC1 * (-1/2 + skoXC1 * (1/2 + skoXC1 * -1/20)))))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) |
     -2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (66 +
       skoXC1 *
       (-18 + skoXC1 * (-1/2 + skoXC1 * (1/2 + skoXC1 * -1/20))))) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) |
     1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_7, plain,
    ((1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) |
     0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_8, plain,
    ((3 + skoXC1 ^ 2 / 10) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 10) <= skoXC1),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    ((exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/20)) +
      exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/20))) /
     (2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) <= skoXC1),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (skoXC1 <
     (exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/20)) +
      exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/20))) /
     (2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) |
     exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/20)) +
     exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/20)) <=
     skoXC1 *
     (2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) |
     2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
     exp(skoXC1 * -1) * 1/2 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_11, plain,
    (exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/20)) +
     exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/20)) <=
     skoXC1 *
     (2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) |
     2 + skoXC1 * (skoXC1 * 1/10) + exp(skoXC1) * 1/2 +
     exp(skoXC1 * -1) * 1/2 <= 0),
    inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (exp(skoXC1 * -1) * (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20)) |
     exp(skoXC1 * -1) <= -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_14, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_15, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -1/5) |
     -4 + skoXC1 * (skoXC1 * -1/5) <= exp(skoXC1)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_18, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_19, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 |
     -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (exp(skoXC1) < -4 + skoXC1 * (skoXC1 * -1/5) |
     -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -1/5) |
     -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_16, refute_0_21])).

cnf(refute_0_23, plain,
    (4 < skoXC1 * (skoXC1 * -1/5) |
     -4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_25, plain, (skoXC1 * (skoXC1 * -1/5) <= 4),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (-4 + skoXC1 * (skoXC1 * -1/5) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_25, refute_0_23])).

cnf(refute_0_27, plain,
    (exp(skoXC1 * -1) * (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_12, refute_0_26])).

cnf(refute_0_28, plain,
    (0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <= exp(skoXC1 * -1)),
    inference(split, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/20) <= -3/2 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <= exp(skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain, (-3/2 < skoXC1 * (1/2 + skoXC1 * 1/20)),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_31, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <= exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain,
    (exp(X_000096) < X_000097 | ~ lgen(1, exp(X_000096), X_000097)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_33, plain,
    (~ lgen(1,
            (X_000096 ^ 4 + 20 * X_000096 ^ 3 + 180 * X_000096 ^ 2 +
             840 * X_000096 + 1680) /
            (X_000096 ^ 4 - 20 * X_000096 ^ 3 + 180 * X_000096 ^ 2 -
             840 * X_000096 + 1680), X_000097) | 0 <= X_000096 |
     lgen(1, exp(X_000096), X_000097)),
    inference(subst, [], [exp_upper_bound_cf4])).

cnf(refute_0_34, plain,
    (exp(X_000096) < X_000097 |
     ~ lgen(1,
            (X_000096 ^ 4 + 20 * X_000096 ^ 3 + 180 * X_000096 ^ 2 +
             840 * X_000096 + 1680) /
            (X_000096 ^ 4 - 20 * X_000096 ^ 3 + 180 * X_000096 ^ 2 -
             840 * X_000096 + 1680), X_000097) | 0 <= X_000096),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (exp(X_000096) < X_000097 |
     X_000097 <=
     (1680 +
      X_000096 * (840 + X_000096 * (180 + X_000096 * (20 + X_000096)))) /
     (1680 +
      X_000096 * (-840 + X_000096 * (180 + X_000096 * (-20 + X_000096)))) |
     0 <= X_000096), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (exp(skoXC1 * -1) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) | 0 <= skoXC1 * -1 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (1680 +
      skoXC1 * -1 *
      (840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (20 + skoXC1 * -1)))) /
     (1680 +
      skoXC1 * -1 *
      (-840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (-20 + skoXC1 * -1))))),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (0 <= skoXC1 * -1 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (1680 +
      skoXC1 * -1 *
      (840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (20 + skoXC1 * -1)))) /
     (1680 +
      skoXC1 * -1 *
      (-840 + skoXC1 * -1 * (180 + skoXC1 * -1 * (-20 + skoXC1 * -1))))),
    inference(resolve, [], [refute_0_31, refute_0_36])).

cnf(refute_0_38, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))))),
    inference(resolve, [], [refute_0_38, refute_0_24])).

cnf(refute_0_40, plain,
    (0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_39, refute_0_7])).

cnf(refute_0_41, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/20) <= -3/2 |
     (-2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (66 +
        skoXC1 *
        (-18 + skoXC1 * (-1/2 + skoXC1 * (1/2 + skoXC1 * -1/20)))))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    ((-2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (66 +
        skoXC1 *
        (-18 + skoXC1 * (-1/2 + skoXC1 * (1/2 + skoXC1 * -1/20)))))) /
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_41, refute_0_30])).

cnf(refute_0_43, plain,
    (-2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (66 +
       skoXC1 *
       (-18 + skoXC1 * (-1/2 + skoXC1 * (1/2 + skoXC1 * -1/20))))) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) |
     1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= 0),
    inference(resolve, [], [refute_0_42, refute_0_6])).

cnf(refute_0_44, plain,
    (exp(skoXC1) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) |
     skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10)))))))) <
     (5806080 +
      skoXC1 *
      (6773760 +
       skoXC1 *
       (3718656 +
        skoXC1 *
        (1257984 +
         skoXC1 *
         (256608 +
          skoXC1 *
          (21744 +
           skoXC1 *
           (-15264/5 +
            skoXC1 *
            (-5004/5 +
             skoXC1 *
             (48/5 +
              skoXC1 *
              (246/5 +
               skoXC1 *
               (107/10 + skoXC1 * (11/10 + skoXC1 * 1/20)))))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     5806080 +
     skoXC1 *
     (6773760 +
      skoXC1 *
      (3718656 +
       skoXC1 *
       (1257984 +
        skoXC1 *
        (256608 +
         skoXC1 *
         (21744 +
          skoXC1 *
          (-15264/5 +
           skoXC1 *
           (-5004/5 +
            skoXC1 *
            (48/5 +
             skoXC1 *
             (246/5 +
              skoXC1 *
              (107/10 + skoXC1 * (11/10 + skoXC1 * 1/20))))))))))) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_46, plain,
    ((3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10)))))))) |
     3456 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 +
          skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_47, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20)) <
     (-3456 + skoXC1 * (-1152 + skoXC1 * -576/5)) /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) |
     -3456 + skoXC1 * (-1152 + skoXC1 * -576/5) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) |
     2304 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_48, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) |
     0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_49, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_50, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_51, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (0 < skoXC1 * -1 |
     exp(skoXC1 * -1) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     2304 /
     (2304 +
      skoXC1 * -1 *
      (-2304 +
       skoXC1 * -1 *
       (1152 +
        skoXC1 * -1 *
        (-384 +
         skoXC1 * -1 * (84 + skoXC1 * -1 * (-12 + skoXC1 * -1))))))),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (0 < skoXC1 * -1 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     2304 /
     (2304 +
      skoXC1 * -1 *
      (-2304 +
       skoXC1 * -1 *
       (1152 +
        skoXC1 * -1 *
        (-384 +
         skoXC1 * -1 * (84 + skoXC1 * -1 * (-12 + skoXC1 * -1))))))),
    inference(resolve, [], [refute_0_31, refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 < 0 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_57, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))))),
    inference(resolve, [], [refute_0_56, refute_0_55])).

cnf(refute_0_58, plain,
    (0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     2304 /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_57, refute_0_48])).

cnf(refute_0_59, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/20) <= -3/2 |
     (-3456 + skoXC1 * (-1152 + skoXC1 * -576/5)) /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    ((-3456 + skoXC1 * (-1152 + skoXC1 * -576/5)) /
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 +
        skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_59, refute_0_30])).

cnf(refute_0_61, plain,
    (-3456 + skoXC1 * (-1152 + skoXC1 * -576/5) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (2304 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))) |
     2304 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_60, refute_0_47])).

cnf(refute_0_62, plain,
    (exp(skoXC1) *
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10)))))))) |
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (-2304 <
     skoXC1 *
     (2304 +
      skoXC1 *
      (1152 + skoXC1 * (384 + skoXC1 * (84 + skoXC1 * (12 + skoXC1)))))),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_64, plain,
    (exp(skoXC1) *
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain,
    (3456 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 +
          skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))) <= 0 |
     exp(skoXC1) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(split, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 +
          skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))) <=
     -3456 |
     exp(skoXC1) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (-3456 <
     skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 + skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_68, plain,
    (exp(skoXC1) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (X_000090 < exp(X_000089) | ~ lgen(1, X_000090, exp(X_000089))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_70, plain,
    (X_000089 < 0 |
     ~ lgen(1, X_000090,
            (X_000089 ^ 4 + 20 * X_000089 ^ 3 + 180 * X_000089 ^ 2 +
             840 * X_000089 + 1680) /
            (X_000089 ^ 4 - 20 * X_000089 ^ 3 + 180 * X_000089 ^ 2 -
             840 * X_000089 + 1680)) | lgen(1, X_000090, exp(X_000089))),
    inference(subst, [], [exp_lower_bound_cf4])).

cnf(refute_0_71, plain,
    (X_000089 < 0 | X_000090 < exp(X_000089) |
     ~ lgen(1, X_000090,
            (X_000089 ^ 4 + 20 * X_000089 ^ 3 + 180 * X_000089 ^ 2 +
             840 * X_000089 + 1680) /
            (X_000089 ^ 4 - 20 * X_000089 ^ 3 + 180 * X_000089 ^ 2 -
             840 * X_000089 + 1680))),
    inference(resolve, [], [refute_0_70, refute_0_69])).

cnf(refute_0_72, plain,
    (X_000089 < 0 | X_000090 < exp(X_000089) |
     (1680 +
      X_000089 * (840 + X_000089 * (180 + X_000089 * (20 + X_000089)))) /
     (1680 +
      X_000089 *
      (-840 + X_000089 * (180 + X_000089 * (-20 + X_000089)))) <=
     X_000090), inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 < 0 |
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <
     exp(skoXC1) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(resolve, [], [refute_0_68, refute_0_73])).

cnf(refute_0_75, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) /
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))))),
    inference(resolve, [], [refute_0_56, refute_0_74])).

cnf(refute_0_76, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (3456 +
      skoXC1 *
      (2304 +
       skoXC1 *
       (3456/5 +
        skoXC1 *
        (576/5 +
         skoXC1 *
         (-42/5 +
          skoXC1 *
          (-24/5 +
           skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20)))))))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10)))))))) |
     3456 +
     skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 +
          skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))) <= 0),
    inference(resolve, [], [refute_0_75, refute_0_46])).

cnf(refute_0_77, plain,
    (skoXC1 *
     (2304 +
      skoXC1 *
      (3456/5 +
       skoXC1 *
       (576/5 +
        skoXC1 *
        (-42/5 +
         skoXC1 *
         (-24/5 +
          skoXC1 * (-3/10 + skoXC1 * (1/10 + skoXC1 * 1/20))))))) <=
     -3456 |
     (5806080 +
      skoXC1 *
      (6773760 +
       skoXC1 *
       (3718656 +
        skoXC1 *
        (1257984 +
         skoXC1 *
         (256608 +
          skoXC1 *
          (21744 +
           skoXC1 *
           (-15264/5 +
            skoXC1 *
            (-5004/5 +
             skoXC1 *
             (48/5 +
              skoXC1 *
              (246/5 +
               skoXC1 *
               (107/10 + skoXC1 * (11/10 + skoXC1 * 1/20)))))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    ((5806080 +
      skoXC1 *
      (6773760 +
       skoXC1 *
       (3718656 +
        skoXC1 *
        (1257984 +
         skoXC1 *
         (256608 +
          skoXC1 *
          (21744 +
           skoXC1 *
           (-15264/5 +
            skoXC1 *
            (-5004/5 +
             skoXC1 *
             (48/5 +
              skoXC1 *
              (246/5 +
               skoXC1 *
               (107/10 + skoXC1 * (11/10 + skoXC1 * 1/20)))))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     3456 +
     skoXC1 *
     (5760 +
      skoXC1 *
      (23616/5 +
       skoXC1 *
       (12672/5 +
        skoXC1 *
        (4992/5 +
         skoXC1 *
         (1416/5 +
          skoXC1 *
          (312/5 + skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))),
    inference(resolve, [], [refute_0_77, refute_0_67])).

cnf(refute_0_79, plain,
    (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     5806080 +
     skoXC1 *
     (6773760 +
      skoXC1 *
      (3718656 +
       skoXC1 *
       (1257984 +
        skoXC1 *
        (256608 +
         skoXC1 *
         (21744 +
          skoXC1 *
          (-15264/5 +
           skoXC1 *
           (-5004/5 +
            skoXC1 *
            (48/5 +
             skoXC1 *
             (246/5 +
              skoXC1 *
              (107/10 + skoXC1 * (11/10 + skoXC1 * 1/20))))))))))) <=
     (3456 +
      skoXC1 *
      (5760 +
       skoXC1 *
       (23616/5 +
        skoXC1 *
        (12672/5 +
         skoXC1 *
         (4992/5 +
          skoXC1 *
          (1416/5 +
           skoXC1 *
           (312/5 +
            skoXC1 * (52/5 + skoXC1 * (6/5 + skoXC1 * 1/10))))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(resolve, [], [refute_0_78, refute_0_45])).

cnf(refute_0_80, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-30240 +
         skoXC1 *
         (17136 +
          skoXC1 *
          (-3744 +
           skoXC1 *
           (2004/5 +
            skoXC1 *
            (816/5 +
             skoXC1 *
             (-18 +
              skoXC1 *
              (243/10 +
               skoXC1 *
               (-33/10 + skoXC1 * (17/20 + skoXC1 * -1/10)))))))))))) <=
     0 |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10)))) <
     (30240 +
      skoXC1 *
      (20160 +
       skoXC1 *
       (4248 +
        skoXC1 *
        (-192 +
         skoXC1 *
         (-168 +
          skoXC1 *
          (-9 + skoXC1 * (41/10 + skoXC1 * (4/5 + skoXC1 * 1/20)))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     30240 +
     skoXC1 *
     (20160 +
      skoXC1 *
      (4248 +
       skoXC1 *
       (-192 +
        skoXC1 *
        (-168 +
         skoXC1 *
         (-9 + skoXC1 * (41/10 + skoXC1 * (4/5 + skoXC1 * 1/20))))))) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_82, plain,
    ((18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10)))) |
     18 +
     skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_83, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20)) <
     (-18 +
      skoXC1 * (3 + skoXC1 * (9/10 + skoXC1 * (-1/5 + skoXC1 * -1/20)))) /
     (12 + skoXC1 * (6 + skoXC1)) |
     -18 +
     skoXC1 * (3 + skoXC1 * (9/10 + skoXC1 * (-1/5 + skoXC1 * -1/20))) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (12 + skoXC1 * (6 + skoXC1)) | 12 + skoXC1 * (6 + skoXC1) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_84, plain,
    ((12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) |
     0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_85, plain,
    (exp(X_000071) < X_000072 | ~ lgen(1, exp(X_000071), X_000072)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_86, plain,
    (~ lgen(1,
            (X_000071 ^ 2 + 6 * X_000071 + 12) /
            (X_000071 ^ 2 - 6 * X_000071 + 12), X_000072) | 0 <= X_000071 |
     lgen(1, exp(X_000071), X_000072)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_87, plain,
    (exp(X_000071) < X_000072 |
     ~ lgen(1,
            (X_000071 ^ 2 + 6 * X_000071 + 12) /
            (X_000071 ^ 2 - 6 * X_000071 + 12), X_000072) | 0 <= X_000071),
    inference(resolve, [], [refute_0_86, refute_0_85])).

cnf(refute_0_88, plain,
    (exp(X_000071) < X_000072 |
     X_000072 <=
     (12 + X_000071 * (6 + X_000071)) / (12 + X_000071 * (-6 + X_000071)) |
     0 <= X_000071), inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (exp(skoXC1 * -1) <
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) | 0 <= skoXC1 * -1 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (12 + skoXC1 * -1 * (6 + skoXC1 * -1)) /
     (12 + skoXC1 * -1 * (-6 + skoXC1 * -1))),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (0 <= skoXC1 * -1 |
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (12 + skoXC1 * -1 * (6 + skoXC1 * -1)) /
     (12 + skoXC1 * -1 * (-6 + skoXC1 * -1))),
    inference(resolve, [], [refute_0_31, refute_0_89])).

cnf(refute_0_91, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    ((skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) /
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1))),
    inference(resolve, [], [refute_0_91, refute_0_24])).

cnf(refute_0_93, plain,
    (0 <= -3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20) |
     (12 + skoXC1 * (-6 + skoXC1)) / (12 + skoXC1 * (6 + skoXC1)) *
     (-3/2 + skoXC1 * (-1/2 + skoXC1 * -1/20)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_92, refute_0_84])).

cnf(refute_0_94, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/20) <= -3/2 |
     (-18 +
      skoXC1 * (3 + skoXC1 * (9/10 + skoXC1 * (-1/5 + skoXC1 * -1/20)))) /
     (12 + skoXC1 * (6 + skoXC1)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    ((-18 +
      skoXC1 * (3 + skoXC1 * (9/10 + skoXC1 * (-1/5 + skoXC1 * -1/20)))) /
     (12 + skoXC1 * (6 + skoXC1)) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
     exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))),
    inference(resolve, [], [refute_0_94, refute_0_30])).

cnf(refute_0_96, plain,
    (-18 +
     skoXC1 * (3 + skoXC1 * (9/10 + skoXC1 * (-1/5 + skoXC1 * -1/20))) <=
     (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/10)) +
      exp(skoXC1) * (-3/2 + skoXC1 * (1/2 + skoXC1 * -1/20))) *
     (12 + skoXC1 * (6 + skoXC1)) | 12 + skoXC1 * (6 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_95, refute_0_83])).

cnf(refute_0_97, plain,
    (exp(skoXC1) *
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10)))) |
     skoXC1 * (6 + skoXC1) <= -12),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain, (-12 < skoXC1 * (6 + skoXC1)),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_99, plain,
    (exp(skoXC1) *
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (18 +
     skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))) <=
     0 |
     exp(skoXC1) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(split, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))) <=
     -18 |
     exp(skoXC1) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (-18 <
     skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))),
    inference(decision, [], [refute_0_24])).

cnf(refute_0_103, plain,
    (exp(skoXC1) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(resolve, [], [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 < 0 |
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <
     exp(skoXC1) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_105, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) /
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))))),
    inference(resolve, [], [refute_0_56, refute_0_105])).

cnf(refute_0_107, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (18 +
      skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10)))) |
     18 +
     skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))) <=
     0), inference(resolve, [], [refute_0_106, refute_0_82])).

cnf(refute_0_108, plain,
    (skoXC1 * (3 + skoXC1 * (-9/10 + skoXC1 * (-1/5 + skoXC1 * 1/20))) <=
     -18 |
     (30240 +
      skoXC1 *
      (20160 +
       skoXC1 *
       (4248 +
        skoXC1 *
        (-192 +
         skoXC1 *
         (-168 +
          skoXC1 *
          (-9 + skoXC1 * (41/10 + skoXC1 * (4/5 + skoXC1 * 1/20)))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    ((30240 +
      skoXC1 *
      (20160 +
       skoXC1 *
       (4248 +
        skoXC1 *
        (-192 +
         skoXC1 *
         (-168 +
          skoXC1 *
          (-9 + skoXC1 * (41/10 + skoXC1 * (4/5 + skoXC1 * 1/20)))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     18 +
     skoXC1 *
     (21 +
      skoXC1 *
      (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))),
    inference(resolve, [], [refute_0_108, refute_0_102])).

cnf(refute_0_110, plain,
    (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     30240 +
     skoXC1 *
     (20160 +
      skoXC1 *
      (4248 +
       skoXC1 *
       (-192 +
        skoXC1 *
        (-168 +
         skoXC1 *
         (-9 + skoXC1 * (41/10 + skoXC1 * (4/5 + skoXC1 * 1/20))))))) <=
     (18 +
      skoXC1 *
      (21 +
       skoXC1 *
       (111/10 + skoXC1 * (17/5 + skoXC1 * (13/20 + skoXC1 * 1/10))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(resolve, [], [refute_0_109, refute_0_81])).

cnf(refute_0_111, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-42 +
          skoXC1 *
          (28 + skoXC1 * (-38/5 + skoXC1 * (7/5 + skoXC1 * -1/10)))))))) <=
     0 |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-42 +
          skoXC1 *
          (28 + skoXC1 * (-38/5 + skoXC1 * (7/5 + skoXC1 * -1/10)))))))) <=
     0), inference(decision, [], [refute_0_24])).

cnf(refute_0_113, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (-42 +
          skoXC1 *
          (28 + skoXC1 * (-38/5 + skoXC1 * (7/5 + skoXC1 * -1/10)))))))) <=
     0), inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (-1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-30240 +
         skoXC1 *
         (17136 +
          skoXC1 *
          (-3744 +
           skoXC1 *
           (2004/5 +
            skoXC1 *
            (816/5 +
             skoXC1 *
             (-18 +
              skoXC1 *
              (243/10 +
               skoXC1 *
               (-33/10 + skoXC1 * (17/20 + skoXC1 * -1/10)))))))))))) <=
     0), inference(decision, [], [refute_0_113, refute_0_24])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-30240 +
         skoXC1 *
         (17136 +
          skoXC1 *
          (-3744 +
           skoXC1 *
           (2004/5 +
            skoXC1 *
            (816/5 +
             skoXC1 *
             (-18 +
              skoXC1 *
              (243/10 +
               skoXC1 *
               (-33/10 + skoXC1 * (17/20 + skoXC1 * -1/10)))))))))))) <=
     0), inference(resolve, [], [refute_0_80, refute_0_114])).

cnf(refute_0_116, plain,
    (-1680 < skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))),
    inference(decision, [], [refute_0_115, refute_0_113, refute_0_24])).

cnf(refute_0_117, plain,
    (exp(skoXC1) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))),
    inference(resolve, [], [refute_0_44, refute_0_116])).

cnf(refute_0_118, plain,
    (2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= 0 |
     exp(skoXC1) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(split, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= -2520 |
     exp(skoXC1) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (-2520 <
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 * (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))),
    inference(decision, [], [refute_0_115, refute_0_113, refute_0_24])).

cnf(refute_0_121, plain,
    (exp(skoXC1) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain,
    (1 +
     skoXC1 *
     (1 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (1/6 +
        skoXC1 *
        (1/24 +
         skoXC1 *
         (1/120 +
          skoXC1 *
          (121/87480 +
           skoXC1 *
           (17/87480 +
            skoXC1 *
            (49/2099520 +
             skoXC1 *
             (409/170061120 +
              skoXC1 *
              (361/1700611200 +
               skoXC1 *
               (1/62985600 +
                skoXC1 *
                (181/183666009600 +
                 skoXC1 *
                 (1/20407334400 +
                  skoXC1 *
                  (1/550998028800 +
                   skoXC1 * 1/24794911296000)))))))))))))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(resolve, [], [refute_0_121, refute_0_5])).

cnf(refute_0_123, plain,
    ((1 +
      skoXC1 *
      (1 +
       skoXC1 *
       (1/2 +
        skoXC1 *
        (1/6 +
         skoXC1 *
         (1/24 +
          skoXC1 *
          (1/120 +
           skoXC1 *
           (121/87480 +
            skoXC1 *
            (17/87480 +
             skoXC1 *
             (49/2099520 +
              skoXC1 *
              (409/170061120 +
               skoXC1 *
               (361/1700611200 +
                skoXC1 *
                (1/62985600 +
                 skoXC1 *
                 (181/183666009600 +
                  skoXC1 *
                  (1/20407334400 +
                   skoXC1 *
                   (1/550998028800 +
                    skoXC1 * 1/24794911296000))))))))))))))) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) |
     2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= 0),
    inference(resolve, [], [refute_0_122, refute_0_0])).

cnf(refute_0_124, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-7/486 +
           skoXC1 *
           (787/14580 +
            skoXC1 *
            (1651/58320 +
             skoXC1 *
             (54199/4723920 +
              skoXC1 *
              (48869/14171760 +
               skoXC1 *
               (3347/4428675 +
                skoXC1 *
                (662443/5101833600 +
                 skoXC1 *
                 (279431/15305500800 +
                  skoXC1 *
                  (21937/10203667200 +
                   skoXC1 *
                   (59017/275499014400 +
                    skoXC1 *
                    (298531/16529940864000 +
                     skoXC1 *
                     (659/516560652000 +
                      skoXC1 *
                      (163/2203992115200 +
                       skoXC1 *
                       (67/19835929036800 +
                        skoXC1 *
                        (11/99179645184000 +
                         skoXC1 * 1/495898225920000)))))))))))))))))))) <=
     0 |
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= -2520),
    inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) <
     (4233600 +
      skoXC1 *
      (2822400 +
       skoXC1 *
       (695520 +
        skoXC1 *
        (40320 +
         skoXC1 *
         (-15240 +
          skoXC1 *
          (-2880 +
           skoXC1 *
           (168 +
            skoXC1 *
            (124 +
             skoXC1 * (39/2 + skoXC1 * (3/2 + skoXC1 * 1/20)))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     4233600 +
     skoXC1 *
     (2822400 +
      skoXC1 *
      (695520 +
       skoXC1 *
       (40320 +
        skoXC1 *
        (-15240 +
         skoXC1 *
         (-2880 +
          skoXC1 *
          (168 +
           skoXC1 *
           (124 +
            skoXC1 * (39/2 + skoXC1 * (3/2 + skoXC1 * 1/20))))))))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_126, plain,
    ((2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) |
     2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_127, plain,
    (skoXC1 < 0 |
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <
     exp(skoXC1) |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_128, plain,
    (skoXC1 < 0 |
     (1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(resolve, [], [refute_0_121, refute_0_127])).

cnf(refute_0_129, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) /
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))))),
    inference(resolve, [], [refute_0_56, refute_0_128])).

cnf(refute_0_130, plain,
    ((1680 + skoXC1 * (840 + skoXC1 * (180 + skoXC1 * (20 + skoXC1)))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) *
     (2520 +
      skoXC1 *
      (420 +
       skoXC1 *
       (-66 +
        skoXC1 *
        (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20)))))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10)))))) |
     2520 +
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_126])).

cnf(refute_0_131, plain,
    (skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= -2520 |
     (4233600 +
      skoXC1 *
      (2822400 +
       skoXC1 *
       (695520 +
        skoXC1 *
        (40320 +
         skoXC1 *
         (-15240 +
          skoXC1 *
          (-2880 +
           skoXC1 *
           (168 +
            skoXC1 *
            (124 +
             skoXC1 * (39/2 + skoXC1 * (3/2 + skoXC1 * 1/20)))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))),
    inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    ((4233600 +
      skoXC1 *
      (2822400 +
       skoXC1 *
       (695520 +
        skoXC1 *
        (40320 +
         skoXC1 *
         (-15240 +
          skoXC1 *
          (-2880 +
           skoXC1 *
           (168 +
            skoXC1 *
            (124 +
             skoXC1 * (39/2 + skoXC1 * (3/2 + skoXC1 * 1/20)))))))))) /
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1)))) <=
     2520 +
     skoXC1 *
     (2940 +
      skoXC1 *
      (1614 +
       skoXC1 *
       (546 +
        skoXC1 *
        (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))),
    inference(resolve, [], [refute_0_131, refute_0_120])).

cnf(refute_0_133, plain,
    (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <=
     0 |
     4233600 +
     skoXC1 *
     (2822400 +
      skoXC1 *
      (695520 +
       skoXC1 *
       (40320 +
        skoXC1 *
        (-15240 +
         skoXC1 *
         (-2880 +
          skoXC1 *
          (168 +
           skoXC1 *
           (124 +
            skoXC1 * (39/2 + skoXC1 * (3/2 + skoXC1 * 1/20))))))))) <=
     (2520 +
      skoXC1 *
      (2940 +
       skoXC1 *
       (1614 +
        skoXC1 *
        (546 +
         skoXC1 *
         (249/2 + skoXC1 * (39/2 + skoXC1 * (41/20 + skoXC1 * 1/10))))))) *
     (1680 + skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))))),
    inference(resolve, [], [refute_0_132, refute_0_125])).

cnf(refute_0_134, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (112 +
            skoXC1 * (skoXC1 * (5 + skoXC1 * (skoXC1 * -1/10)))))))))) <=
     0 |
     skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) <= -1680),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (-1680 < skoXC1 * (-840 + skoXC1 * (180 + skoXC1 * (-20 + skoXC1))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (112 +
            skoXC1 * (skoXC1 * (5 + skoXC1 * (skoXC1 * -1/10)))))))))) <=
     0),
    inference(decision, [], [refute_0_115, refute_0_113, refute_0_24])).

cnf(refute_0_136, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (112 +
            skoXC1 * (skoXC1 * (5 + skoXC1 * (skoXC1 * -1/10)))))))))) <=
     0), inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (-7/486 +
           skoXC1 *
           (787/14580 +
            skoXC1 *
            (1651/58320 +
             skoXC1 *
             (54199/4723920 +
              skoXC1 *
              (48869/14171760 +
               skoXC1 *
               (3347/4428675 +
                skoXC1 *
                (662443/5101833600 +
                 skoXC1 *
                 (279431/15305500800 +
                  skoXC1 *
                  (21937/10203667200 +
                   skoXC1 *
                   (59017/275499014400 +
                    skoXC1 *
                    (298531/16529940864000 +
                     skoXC1 *
                     (659/516560652000 +
                      skoXC1 *
                      (163/2203992115200 +
                       skoXC1 *
                       (67/19835929036800 +
                        skoXC1 *
                        (11/99179645184000 +
                         skoXC1 * 1/495898225920000)))))))))))))))))))) |
     skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= -2520),
    inference(decision, [],
              [refute_0_136, refute_0_115, refute_0_113, refute_0_24])).

cnf(refute_0_138, plain,
    (skoXC1 *
     (420 +
      skoXC1 *
      (-66 +
       skoXC1 *
       (-18 + skoXC1 * (1/2 + skoXC1 * (1/2 + skoXC1 * 1/20))))) <= -2520),
    inference(resolve, [], [refute_0_124, refute_0_137])).

cnf(refute_0_139, plain, ($false),
    inference(resolve, [], [refute_0_138, refute_0_120])).
% SZS output end CNFRefutation for exp-hyperbolic-2.tptp
