% SZS output start CNFRefutation for exp-ln-2-sqrt-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_mul_divide_neg, axiom, (Y / Z < X | Y <= X * Z | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_cf3, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120), Y) | lgen(R, exp(X), Y))).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

cnf(sqrt_lower_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R, Y,
            32 * X * (X + 16) * (X ^ 2 + 96 * X + 256) /
            (X ^ 4 + 448 * X ^ 3 + 17920 * X ^ 2 + 114688 * X + 65536)) |
     lgen(R, Y, sqrt(X)))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

fof(exp_ln, conjecture, (exp(8664339757/12500000000 * sqrt(2)) < exp(1))).

fof(subgoal_0, plain, (exp(8664339757/12500000000 * sqrt(2)) < exp(1)),
    inference(strip, [], [exp_ln])).

fof(negate_0_0, plain,
    (~ ~ exp(1) <= exp(8664339757/12500000000 * sqrt(2))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (exp(1) <= exp(8664339757/12500000000 * sqrt(2))),
    inference(canonicalize, [], [negate_0_0])).

cnf(refute_0_0, plain,
    (sqrt(X_000055) < X_000056 | ~ lgen(1, sqrt(X_000055), X_000056)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (X_000055 < 0 | ~ lgen(1, (X_000055 + 1) / 2, X_000056) |
     lgen(1, sqrt(X_000055), X_000056)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_2, plain,
    (sqrt(X_000055) < X_000056 | X_000055 < 0 |
     ~ lgen(1, (X_000055 + 1) / 2, X_000056)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (sqrt(X_000055) < X_000056 | X_000055 < 0 |
     X_000056 <= 1/2 + X_000055 * 1/2),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (2 < 0 | sqrt(2) < 150000000000/8664339757 |
     150000000000/8664339757 <= 1/2 + 2 * 1/2),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (25993019271/625000000 <
     23137/16272 *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     23137/16272 <=
     25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     225212350274192457147/39062500000000000000 +
     sqrt(2) *
     -650438773416031852509423631093/1953125000000000000000000000000 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (~ lgen(0, sqrt(X_000111), X_000112) | sqrt(X_000111) <= X_000112),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (X_000111 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000111 + 17920 * X_000111 ^ 2 +
             448 * X_000111 ^ 3 + X_000111 ^ 4) /
            ((X_000111 + 16) * (256 + 96 * X_000111 + X_000111 ^ 2)),
            X_000112) | lgen(0, sqrt(X_000111), X_000112)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_8, plain,
    (X_000111 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000111 + 17920 * X_000111 ^ 2 +
             448 * X_000111 ^ 3 + X_000111 ^ 4) /
            ((X_000111 + 16) * (256 + 96 * X_000111 + X_000111 ^ 2)),
            X_000112) | sqrt(X_000111) <= X_000112),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    (X_000111 < 1 |
     X_000112 <
     (2048 +
      X_000111 *
      (3584 + X_000111 * (560 + X_000111 * (14 + X_000111 * 1/32)))) /
     (4096 + X_000111 * (1792 + X_000111 * (112 + X_000111))) |
     sqrt(X_000111) <= X_000112), inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (2 < 1 |
     25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2))) |
     sqrt(2) <=
     25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    ((-120 +
      sqrt(2) *
      (-25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        -650438773416031852509423631093/
        1953125000000000000000000000000))) /
     (-120 +
      sqrt(2) *
      (25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        650438773416031852509423631093/1953125000000000000000000000000))) <
     131639193503/48427561125 |
     0 <=
     -120 +
     sqrt(2) *
     (25993019271/625000000 +
      sqrt(2) *
      (-225212350274192457147/39062500000000000000 +
       sqrt(2) *
       650438773416031852509423631093/1953125000000000000000000000000)) |
     -120 +
     sqrt(2) *
     (-25993019271/625000000 +
      sqrt(2) *
      (-225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     131639193503/48427561125 *
     (-120 +
      sqrt(2) *
      (25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        650438773416031852509423631093/1953125000000000000000000000000)))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_12, plain, (exp(1) <= exp(8664339757/12500000000 * sqrt(2))),
    inference(canonicalize, [], [normalize_0_0])).

cnf(refute_0_13, plain, (exp(1) <= exp(sqrt(2) * 8664339757/12500000000)),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (X_000018 < exp(X_000017) | ~ lgen(1, X_000018, exp(X_000017))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_15, plain,
    (~ lgen(1, X_000018,
            (1 + X_000017 / 3 + 1/2 * (X_000017 / 3) ^ 2 +
             1/6 * (X_000017 / 3) ^ 3 + 1/24 * (X_000017 / 3) ^ 4 +
             1/120 * (X_000017 / 3) ^ 5) ^ 3) |
     lgen(1, X_000018, exp(X_000017))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_16, plain,
    (X_000018 < exp(X_000017) |
     ~ lgen(1, X_000018,
            (1 + X_000017 / 3 + 1/2 * (X_000017 / 3) ^ 2 +
             1/6 * (X_000017 / 3) ^ 3 + 1/24 * (X_000017 / 3) ^ 4 +
             1/120 * (X_000017 / 3) ^ 5) ^ 3)),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    (X_000018 < exp(X_000017) |
     1 +
     X_000017 *
     (1 +
      X_000017 *
      (1/2 +
       X_000017 *
       (1/6 +
        X_000017 *
        (1/24 +
         X_000017 *
         (1/120 +
          X_000017 *
          (121/87480 +
           X_000017 *
           (17/87480 +
            X_000017 *
            (49/2099520 +
             X_000017 *
             (409/170061120 +
              X_000017 *
              (361/1700611200 +
               X_000017 *
               (1/62985600 +
                X_000017 *
                (181/183666009600 +
                 X_000017 *
                 (1/20407334400 +
                  X_000017 *
                  (1/550998028800 +
                   X_000017 * 1/24794911296000)))))))))))))) <= X_000018),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (exp(sqrt(2) * 8664339757/12500000000) < exp(1) |
     1 +
     1 *
     (1 +
      1 *
      (1/2 +
       1 *
       (1/6 +
        1 *
        (1/24 +
         1 *
         (1/120 +
          1 *
          (121/87480 +
           1 *
           (17/87480 +
            1 *
            (49/2099520 +
             1 *
             (409/170061120 +
              1 *
              (361/1700611200 +
               1 *
               (1/62985600 +
                1 *
                (181/183666009600 +
                 1 *
                 (1/20407334400 +
                  1 *
                  (1/550998028800 + 1 * 1/24794911296000)))))))))))))) <=
     exp(sqrt(2) * 8664339757/12500000000)),
    inference(subst, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (1 +
     1 *
     (1 +
      1 *
      (1/2 +
       1 *
       (1/6 +
        1 *
        (1/24 +
         1 *
         (1/120 +
          1 *
          (121/87480 +
           1 *
           (17/87480 +
            1 *
            (49/2099520 +
             1 *
             (409/170061120 +
              1 *
              (361/1700611200 +
               1 *
               (1/62985600 +
                1 *
                (181/183666009600 +
                 1 *
                 (1/20407334400 +
                  1 *
                  (1/550998028800 + 1 * 1/24794911296000)))))))))))))) <=
     exp(sqrt(2) * 8664339757/12500000000)),
    inference(resolve, [], [refute_0_13, refute_0_18])).

cnf(refute_0_20, plain,
    (131639193503/48427561125 <= exp(sqrt(2) * 8664339757/12500000000)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (exp(X_000146) < X_000147 | ~ lgen(1, exp(X_000146), X_000147)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_22, plain,
    (4 < X_000146 | X_000146 < 0 |
     ~ lgen(1,
            -(X_000146 ^ 3 + 12 * X_000146 ^ 2 + 60 * X_000146 + 120) /
            (X_000146 ^ 3 - 12 * X_000146 ^ 2 + 60 * X_000146 - 120),
            X_000147) | lgen(1, exp(X_000146), X_000147)),
    inference(subst, [], [exp_upper_bound_cf3])).

cnf(refute_0_23, plain,
    (4 < X_000146 | exp(X_000146) < X_000147 | X_000146 < 0 |
     ~ lgen(1,
            -(X_000146 ^ 3 + 12 * X_000146 ^ 2 + 60 * X_000146 + 120) /
            (X_000146 ^ 3 - 12 * X_000146 ^ 2 + 60 * X_000146 - 120),
            X_000147)),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (4 < X_000146 | exp(X_000146) < X_000147 | X_000146 < 0 |
     X_000147 <=
     (-120 + X_000146 * (-60 + X_000146 * (-12 + X_000146 * -1))) /
     (-120 + X_000146 * (60 + X_000146 * (-12 + X_000146)))),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (sqrt(2) * 8664339757/12500000000 < 0 |
     exp(sqrt(2) * 8664339757/12500000000) < 131639193503/48427561125 |
     4 < sqrt(2) * 8664339757/12500000000 |
     131639193503/48427561125 <=
     (-120 +
      sqrt(2) * 8664339757/12500000000 *
      (-60 +
       sqrt(2) * 8664339757/12500000000 *
       (-12 + sqrt(2) * 8664339757/12500000000 * -1))) /
     (-120 +
      sqrt(2) * 8664339757/12500000000 *
      (60 +
       sqrt(2) * 8664339757/12500000000 *
       (-12 + sqrt(2) * 8664339757/12500000000)))),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (sqrt(2) * 8664339757/12500000000 < 0 |
     4 < sqrt(2) * 8664339757/12500000000 |
     131639193503/48427561125 <=
     (-120 +
      sqrt(2) * 8664339757/12500000000 *
      (-60 +
       sqrt(2) * 8664339757/12500000000 *
       (-12 + sqrt(2) * 8664339757/12500000000 * -1))) /
     (-120 +
      sqrt(2) * 8664339757/12500000000 *
      (60 +
       sqrt(2) * 8664339757/12500000000 *
       (-12 + sqrt(2) * 8664339757/12500000000)))),
    inference(resolve, [], [refute_0_20, refute_0_25])).

cnf(refute_0_27, plain,
    (50000000000/8664339757 < sqrt(2) |
     131639193503/48427561125 <=
     (-120 +
      sqrt(2) *
      (-25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        -650438773416031852509423631093/
        1953125000000000000000000000000))) /
     (-120 +
      sqrt(2) *
      (25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        650438773416031852509423631093/1953125000000000000000000000000)))),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (~ lgen(0, sqrt(X_000055), X_000056) | sqrt(X_000055) <= X_000056),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_29, plain,
    (X_000055 < 0 | ~ lgen(0, (X_000055 + 1) / 2, X_000056) |
     lgen(0, sqrt(X_000055), X_000056)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_30, plain,
    (X_000055 < 0 | ~ lgen(0, (X_000055 + 1) / 2, X_000056) |
     sqrt(X_000055) <= X_000056),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain,
    (X_000055 < 0 | X_000056 < 1/2 + X_000055 * 1/2 |
     sqrt(X_000055) <= X_000056),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (2 < 0 | 50000000000/8664339757 < 1/2 + 2 * 1/2 |
     sqrt(2) <= 50000000000/8664339757),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_33, plain, (sqrt(2) <= 50000000000/8664339757),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (131639193503/48427561125 <=
     (-120 +
      sqrt(2) *
      (-25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        -650438773416031852509423631093/
        1953125000000000000000000000000))) /
     (-120 +
      sqrt(2) *
      (25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        650438773416031852509423631093/1953125000000000000000000000000)))),
    inference(resolve, [], [refute_0_33, refute_0_27])).

cnf(refute_0_35, plain,
    (0 <=
     -120 +
     sqrt(2) *
     (25993019271/625000000 +
      sqrt(2) *
      (-225212350274192457147/39062500000000000000 +
       sqrt(2) *
       650438773416031852509423631093/1953125000000000000000000000000)) |
     -120 +
     sqrt(2) *
     (-25993019271/625000000 +
      sqrt(2) *
      (-225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     131639193503/48427561125 *
     (-120 +
      sqrt(2) *
      (25993019271/625000000 +
       sqrt(2) *
       (-225212350274192457147/39062500000000000000 +
        sqrt(2) *
        650438773416031852509423631093/1953125000000000000000000000000)))),
    inference(resolve, [], [refute_0_34, refute_0_11])).

cnf(refute_0_36, plain,
    (sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     -665693059024/3228504075 |
     sqrt(2) *
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     -120), inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (2 < 0 | sqrt(2) < 3120436214175000000000/390039885259336036349 |
     3120436214175000000000/390039885259336036349 <= 1/2 + 2 * 1/2),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_38, plain,
    (390039885259336036349/2522268808593750000 <
     23137/16272 *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     23137/16272 <=
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     3123381216333578374156113784261/315283601074218750000000000000 +
     sqrt(2) *
     -29280599753310474236811667759432005612101/
     23646270080566406250000000000000000000000 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_39, plain,
    (2 < 1 |
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2))) |
     sqrt(2) <=
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_40, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     -665693059024/3228504075), introduced(tautology, [assume])).

cnf(refute_0_41, plain,
    (-390039885259336036349/2522268808593750000 +
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) < 0 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     0 <
     -390039885259336036349/2522268808593750000 +
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     0 <= -665693059024/3228504075), inference(split, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     -390039885259336036349/2522268808593750000 <
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     390039885259336036349/2522268808593750000 <
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) < 32544/23137 |
     32544/23137 *
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) <=
     561729812402657873006813/41042358053437500000000 |
     72265671202310002842850004626446757/
     5130294756679687500000000000000000 +
     sqrt(2) *
     -677465236492344442417111556949978313847180837/
     384772106750976562500000000000000000000000000 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_44, plain,
    (X_000126 < sqrt(X_000125) | ~ lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_45, plain,
    (X_000125 < 1 |
     ~ lgen(1, X_000126,
            32 * X_000125 * (X_000125 + 16) *
            (X_000125 ^ 2 + 96 * X_000125 + 256) /
            (X_000125 ^ 4 + 448 * X_000125 ^ 3 + 17920 * X_000125 ^ 2 +
             114688 * X_000125 + 65536)) |
     lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [sqrt_lower_bound_2_large])).

cnf(refute_0_46, plain,
    (X_000125 < 1 | X_000126 < sqrt(X_000125) |
     ~ lgen(1, X_000126,
            32 * X_000125 * (X_000125 + 16) *
            (X_000125 ^ 2 + 96 * X_000125 + 256) /
            (X_000125 ^ 4 + 448 * X_000125 ^ 3 + 17920 * X_000125 ^ 2 +
             114688 * X_000125 + 65536))),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain,
    (X_000125 < 1 | X_000126 < sqrt(X_000125) |
     X_000125 *
     (131072 + X_000125 * (57344 + X_000125 * (3584 + X_000125 * 32))) /
     (65536 +
      X_000125 *
      (114688 + X_000125 * (17920 + X_000125 * (448 + X_000125)))) <=
     X_000126), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (2 < 1 |
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) < sqrt(2) |
     2 * (131072 + 2 * (57344 + 2 * (3584 + 2 * 32))) /
     (65536 + 2 * (114688 + 2 * (17920 + 2 * (448 + 2)))) <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (23137/16272 <
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     0 <=
     -390039885259336036349/2522268808593750000 +
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     23137/16272 *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     -665693059024/3228504075),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_50, plain,
    (sqrt(X_000111) < X_000112 | ~ lgen(1, sqrt(X_000111), X_000112)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_51, plain,
    (X_000111 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000111 + 17920 * X_000111 ^ 2 +
             448 * X_000111 ^ 3 + X_000111 ^ 4) /
            ((X_000111 + 16) * (256 + 96 * X_000111 + X_000111 ^ 2)),
            X_000112) | lgen(1, sqrt(X_000111), X_000112)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_52, plain,
    (sqrt(X_000111) < X_000112 | X_000111 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000111 + 17920 * X_000111 ^ 2 +
             448 * X_000111 ^ 3 + X_000111 ^ 4) /
            ((X_000111 + 16) * (256 + 96 * X_000111 + X_000111 ^ 2)),
            X_000112)),
    inference(resolve, [], [refute_0_51, refute_0_50])).

cnf(refute_0_53, plain,
    (sqrt(X_000111) < X_000112 | X_000111 < 1 |
     X_000112 <=
     (2048 +
      X_000111 *
      (3584 + X_000111 * (560 + X_000111 * (14 + X_000111 * 1/32)))) /
     (4096 + X_000111 * (1792 + X_000111 * (112 + X_000111)))),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (2 < 1 |
     sqrt(2) <
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(subst, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     0 <=
     -390039885259336036349/2522268808593750000 +
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <= sqrt(2)),
    inference(split, [], [refute_0_40])).

cnf(refute_0_56, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <= sqrt(2)),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (2 < 1 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(resolve, [], [refute_0_56, refute_0_54])).

cnf(refute_0_58, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     -665693059024/3228504075 /
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <= 23137/16272),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     0 <=
     -390039885259336036349/2522268808593750000 +
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     23137/16272 *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) <=
     -665693059024/3228504075 |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(resolve, [], [refute_0_58, refute_0_49])).

cnf(refute_0_60, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     sqrt(2) *
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) <=
     561729812402657873006813/41042358053437500000000),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     72265671202310002842850004626446757/
     5130294756679687500000000000000000 +
     sqrt(2) *
     -677465236492344442417111556949978313847180837/
     384772106750976562500000000000000000000000000 <= 0 |
     sqrt(2) <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000)),
    inference(split, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= sqrt(2) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     sqrt(2) <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000)),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (sqrt(2) < 3120436214175000000000/390039885259336036349),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_64, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     sqrt(2) <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000)),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain,
    (2 < 1 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     2 * (131072 + 2 * (57344 + 2 * (3584 + 2 * 32))) /
     (65536 + 2 * (114688 + 2 * (17920 + 2 * (448 + 2)))) <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000)),
    inference(resolve, [], [refute_0_64, refute_0_48])).

cnf(refute_0_66, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     32544/23137 <=
     561729812402657873006813/41042358053437500000000 /
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     32544/23137 *
     (72265671202310002842850004626446757/
      5130294756679687500000000000000000 +
      sqrt(2) *
      -677465236492344442417111556949978313847180837/
      384772106750976562500000000000000000000000000) <=
     561729812402657873006813/41042358053437500000000 |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000 |
     72265671202310002842850004626446757/
     5130294756679687500000000000000000 +
     sqrt(2) *
     -677465236492344442417111556949978313847180837/
     384772106750976562500000000000000000000000000 <= 0),
    inference(resolve, [], [refute_0_66, refute_0_43])).

cnf(refute_0_68, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     24555540432834172252097610933585925000000000/
     9926123316372250766279155370447449902502239 <= sqrt(2) |
     3120436214175000000000/390039885259336036349 <= sqrt(2) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (2 < 0 |
     sqrt(2) <
     24555540432834172252097610933585925000000000/
     9926123316372250766279155370447449902502239 |
     24555540432834172252097610933585925000000000/
     9926123316372250766279155370447449902502239 <= 1/2 + 2 * 1/2),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_70, plain,
    (sqrt(2) <
     24555540432834172252097610933585925000000000/
     9926123316372250766279155370447449902502239),
    inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= sqrt(2) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(resolve, [], [refute_0_68, refute_0_70])).

cnf(refute_0_72, plain,
    (2 < 0 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= 1/2 + 2 * 1/2 |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(resolve, [], [refute_0_71, refute_0_37])).

cnf(refute_0_73, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     sqrt(2) *
     (-3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <=
     -390039885259336036349/2522268808593750000),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     390039885259336036349/2522268808593750000 <
     sqrt(2) *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000)),
    inference(resolve, [], [refute_0_73, refute_0_42])).

cnf(refute_0_75, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) < sqrt(2) |
     3123381216333578374156113784261/315283601074218750000000000000 +
     sqrt(2) *
     -29280599753310474236811667759432005612101/
     23646270080566406250000000000000000000000 <= 0),
    inference(split, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) < sqrt(2) |
     3120436214175000000000/390039885259336036349 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) < sqrt(2)),
    inference(resolve, [], [refute_0_76, refute_0_63])).

cnf(refute_0_78, plain,
    (2 < 1 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) <
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(resolve, [], [refute_0_39, refute_0_77])).

cnf(refute_0_79, plain,
    (390039885259336036349/2522268808593750000 /
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) < 23137/16272 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000))),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (390039885259336036349/2522268808593750000 <
     23137/16272 *
     (3123381216333578374156113784261/315283601074218750000000000000 +
      sqrt(2) *
      -29280599753310474236811667759432005612101/
      23646270080566406250000000000000000000000) |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3123381216333578374156113784261/315283601074218750000000000000 +
     sqrt(2) *
     -29280599753310474236811667759432005612101/
     23646270080566406250000000000000000000000 <= 0),
    inference(resolve, [], [refute_0_38, refute_0_79])).

cnf(refute_0_81, plain,
    (sqrt(2) <
     -6241751901804889266858674925000000000/
     78190058964967761064810187007764441 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (~ lgen(0, X_000063, sqrt(X_000062)) | X_000063 <= sqrt(X_000062)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_83, plain,
    (X_000062 < 0 | ~ lgen(0, X_000063, 2 * X_000062 / (X_000062 + 1)) |
     lgen(0, X_000063, sqrt(X_000062))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_84, plain,
    (X_000062 < 0 | ~ lgen(0, X_000063, 2 * X_000062 / (X_000062 + 1)) |
     X_000063 <= sqrt(X_000062)),
    inference(resolve, [], [refute_0_83, refute_0_82])).

cnf(refute_0_85, plain,
    (X_000062 * 2 / (1 + X_000062) < X_000063 | X_000062 < 0 |
     X_000063 <= sqrt(X_000062)),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (2 * 2 / (1 + 2) <
     -6241751901804889266858674925000000000/
     78190058964967761064810187007764441 | 2 < 0 |
     -6241751901804889266858674925000000000/
     78190058964967761064810187007764441 <= sqrt(2)),
    inference(subst, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (-6241751901804889266858674925000000000/
     78190058964967761064810187007764441 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= sqrt(2)),
    inference(resolve, [], [refute_0_87, refute_0_81])).

cnf(refute_0_89, plain,
    (2 < 0 |
     -665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000)) |
     3120436214175000000000/390039885259336036349 <= 1/2 + 2 * 1/2),
    inference(resolve, [], [refute_0_88, refute_0_37])).

cnf(refute_0_90, plain,
    (-665693059024/3228504075 <
     sqrt(2) *
     (-390039885259336036349/2522268808593750000 +
      sqrt(2) *
      (3123381216333578374156113784261/315283601074218750000000000000 +
       sqrt(2) *
       -29280599753310474236811667759432005612101/
       23646270080566406250000000000000000000000))),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (sqrt(2) *
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     -120), inference(resolve, [], [refute_0_36, refute_0_90])).

cnf(refute_0_92, plain,
    (-25993019271/625000000 +
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     0 |
     0 <
     -25993019271/625000000 +
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     0 <= -120), inference(split, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (-25993019271/625000000 <
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) |
     25993019271/625000000 <
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000)),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (2 < 0 |
     sqrt(2) <
     23967073071746292617391650000000000/
     220498744188034798000694610940527 |
     23967073071746292617391650000000000/
     220498744188034798000694610940527 <= 1/2 + 2 * 1/2),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_95, plain,
    (-206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) < 32544/23137 |
     32544/23137 *
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) <=
     -206333171042291/3390000000000 |
     1736912716097996960336713/211875000000000000000000 +
     sqrt(2) *
     -15049201900526728971510534552598741/
     31781250000000000000000000000000000 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_96, plain,
    (2 < 1 |
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) < sqrt(2) |
     2 * (131072 + 2 * (57344 + 2 * (3584 + 2 * 32))) /
     (65536 + 2 * (114688 + 2 * (17920 + 2 * (448 + 2)))) <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_97, plain,
    (23137/16272 <
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) |
     0 <=
     -25993019271/625000000 +
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     23137/16272 *
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     -120), inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_98, plain,
    (2 < 1 |
     sqrt(2) <
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) |
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(subst, [], [refute_0_53])).

cnf(refute_0_99, plain,
    (0 <=
     -25993019271/625000000 +
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     sqrt(2)), inference(split, [], [refute_0_91])).

cnf(refute_0_100, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     sqrt(2)), inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (2 < 1 |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(resolve, [], [refute_0_100, refute_0_98])).

cnf(refute_0_102, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     -120 /
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     23137/16272), inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (0 <=
     -25993019271/625000000 +
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     23137/16272 *
     (-25993019271/625000000 +
      sqrt(2) *
      (225212350274192457147/39062500000000000000 +
       sqrt(2) *
       -650438773416031852509423631093/1953125000000000000000000000000)) <=
     -120 |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000),
    inference(resolve, [], [refute_0_102, refute_0_97])).

cnf(refute_0_104, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     sqrt(2) *
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) <=
     -206333171042291/3390000000000),
    inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     1736912716097996960336713/211875000000000000000000 +
     sqrt(2) *
     -15049201900526728971510534552598741/
     31781250000000000000000000000000000 <= 0 |
     sqrt(2) <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000)),
    inference(split, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (150000000000/8664339757 <= sqrt(2) |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     sqrt(2) <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain, (sqrt(2) < 150000000000/8664339757),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_108, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     sqrt(2) <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000)),
    inference(resolve, [], [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain,
    (2 < 1 |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     2 * (131072 + 2 * (57344 + 2 * (3584 + 2 * 32))) /
     (65536 + 2 * (114688 + 2 * (17920 + 2 * (448 + 2)))) <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000)),
    inference(resolve, [], [refute_0_108, refute_0_96])).

cnf(refute_0_110, plain,
    (32544/23137 <=
     -206333171042291/3390000000000 /
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000), inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (32544/23137 *
     (1736912716097996960336713/211875000000000000000000 +
      sqrt(2) *
      -15049201900526728971510534552598741/
      31781250000000000000000000000000000) <=
     -206333171042291/3390000000000 |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000 |
     1736912716097996960336713/211875000000000000000000 +
     sqrt(2) *
     -15049201900526728971510534552598741/
     31781250000000000000000000000000000 <= 0),
    inference(resolve, [], [refute_0_110, refute_0_95])).

cnf(refute_0_112, plain,
    (150000000000/8664339757 <= sqrt(2) |
     23967073071746292617391650000000000/
     220498744188034798000694610940527 <= sqrt(2) |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000), inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (23967073071746292617391650000000000/
     220498744188034798000694610940527 <= sqrt(2) |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000),
    inference(resolve, [], [refute_0_112, refute_0_107])).

cnf(refute_0_114, plain,
    (2 < 0 |
     23967073071746292617391650000000000/
     220498744188034798000694610940527 <= 1/2 + 2 * 1/2 |
     sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000),
    inference(resolve, [], [refute_0_113, refute_0_94])).

cnf(refute_0_115, plain,
    (sqrt(2) *
     (-225212350274192457147/39062500000000000000 +
      sqrt(2) *
      650438773416031852509423631093/1953125000000000000000000000000) <=
     -25993019271/625000000), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (25993019271/625000000 <
     sqrt(2) *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000)),
    inference(resolve, [], [refute_0_115, refute_0_93])).

cnf(refute_0_117, plain,
    (25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     sqrt(2) |
     225212350274192457147/39062500000000000000 +
     sqrt(2) *
     -650438773416031852509423631093/1953125000000000000000000000000 <= 0),
    inference(split, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     sqrt(2) | 150000000000/8664339757 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     sqrt(2)), inference(resolve, [], [refute_0_118, refute_0_107])).

cnf(refute_0_120, plain,
    (2 < 1 |
     25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     (2048 + 2 * (3584 + 2 * (560 + 2 * (14 + 2 * 1/32)))) /
     (4096 + 2 * (1792 + 2 * (112 + 2)))),
    inference(resolve, [], [refute_0_10, refute_0_119])).

cnf(refute_0_121, plain,
    (25993019271/625000000 /
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) <
     23137/16272), inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (25993019271/625000000 <
     23137/16272 *
     (225212350274192457147/39062500000000000000 +
      sqrt(2) *
      -650438773416031852509423631093/1953125000000000000000000000000) |
     225212350274192457147/39062500000000000000 +
     sqrt(2) *
     -650438773416031852509423631093/1953125000000000000000000000000 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_121])).

cnf(refute_0_123, plain,
    (sqrt(2) < -122479975656343650000000000/1736912716097996960336713 |
     150000000000/8664339757 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (2 * 2 / (1 + 2) <
     -122479975656343650000000000/1736912716097996960336713 | 2 < 0 |
     -122479975656343650000000000/1736912716097996960336713 <= sqrt(2)),
    inference(subst, [], [refute_0_85])).

cnf(refute_0_125, plain,
    (-122479975656343650000000000/1736912716097996960336713 <= sqrt(2)),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain, (150000000000/8664339757 <= sqrt(2)),
    inference(resolve, [], [refute_0_125, refute_0_123])).

cnf(refute_0_127, plain,
    (2 < 0 | 150000000000/8664339757 <= 1/2 + 2 * 1/2),
    inference(resolve, [], [refute_0_126, refute_0_4])).

cnf(refute_0_128, plain, ($false),
    inference(arithmetic, [], [refute_0_127])).
% SZS output end CNFRefutation for exp-ln-2-sqrt-weak.tptp
