% SZS output start CNFRefutation for exp-hyperbolic-1.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_mul_divide_neg, axiom, (Y / Z < X | Y <= X * Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf3, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf3, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120), Y) | lgen(R, exp(X), Y))).

cnf(exp_lower_bound_cf5, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240)) | lgen(R, Y, exp(X)))).

cnf(exp_lower_bound_cf7, axiom,
    (0 <= X |
     ~ lgen(R, Y,
            -(X ^ 7 + 56 * X ^ 6 + 1512 * X ^ 5 + 25200 * X ^ 4 +
             277200 * X ^ 3 + 1995840 * X ^ 2 + 8648640 * X + 17297280) /
            (X ^ 7 - 56 * X ^ 6 + 1512 * X ^ 5 - 25200 * X ^ 4 +
             277200 * X ^ 3 - 1995840 * X ^ 2 + 8648640 * X - 17297280)) |
     lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf5, axiom,
    (X < 0 | 7 < X |
     ~ lgen(R,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240), Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf7, axiom,
    (X < 0 | 19/2 < X |
     ~ lgen(R,
            -(X ^ 7 + 56 * X ^ 6 + 1512 * X ^ 5 + 25200 * X ^ 4 +
             277200 * X ^ 3 + 1995840 * X ^ 2 + 8648640 * X + 17297280) /
            (X ^ 7 - 56 * X ^ 6 + 1512 * X ^ 5 - 25200 * X ^ 4 +
             277200 * X ^ 3 - 1995840 * X ^ 2 + 8648640 * X - 17297280),
            Y) | lgen(R, exp(X), Y))).

fof(exp_hyp, conjecture,
    (! [X] :
       ((0 < X & X < 10357/2000) =>
        (3 + X ^ 2 / 11) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 11) < X))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 10357/2000) =>
        (3 + X ^ 2 / 11) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 11) < X)),
    inference(strip, [], [exp_hyp])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 10357/2000) =>
          (3 + X ^ 2 / 11) * ((exp(X) - exp(-X)) / 2) /
          (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 11) < X)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 10357/2000 &
        X <=
        (3 + X ^ 2 / 11) * ((exp(X) - exp(-X)) / 2) /
        (2 + (exp(X) + exp(-X)) / 2 + X ^ 2 / 11))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 10357/2000 & 0 < skoXC1 &
     skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 11)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 11)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 10357/2000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1)))))) <
     (-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-3843/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (-53/22 + skoXC1 * 1/11))))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (1995840 +
        skoXC1 *
        (277200 +
         skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1))))))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -784604620800 +
     skoXC1 *
     (-523069747200 +
      skoXC1 *
      (-136101772800 +
       skoXC1 *
       (-11481523200 +
        skoXC1 *
        (2396822400 +
         skoXC1 *
         (-725760 +
          skoXC1 *
          (843091200/11 +
           skoXC1 *
           (-64350720/11 +
            skoXC1 *
            (-30275280/11 +
             skoXC1 *
             (6064800/11 +
              skoXC1 *
              (-459438/11 +
               skoXC1 *
               (19017/11 +
                skoXC1 *
                (-3843/22 +
                 skoXC1 *
                 (623/22 + skoXC1 * (-53/22 + skoXC1 * 1/11))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    ((-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <=
     -17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))) |
     -17297280 +
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1)))))) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_2, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-45360 +
      skoXC1 *
      (7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (-3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (123/22 + skoXC1 * (-19/22 + skoXC1 * 1/22))))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) |
     0 <=
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     -45360 +
     skoXC1 *
     (7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (123/22 + skoXC1 * (-19/22 + skoXC1 * 1/22))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_3, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) |
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_4, plain,
    (skoXC1 <=
     (3 + skoXC1 ^ 2 / 11) * ((exp(skoXC1) - exp(-skoXC1)) / 2) /
     (2 + (exp(skoXC1) + exp(-skoXC1)) / 2 + skoXC1 ^ 2 / 11)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_5, plain,
    (skoXC1 <=
     (exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/22)) +
      exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/22))) /
     (2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2)), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/22)) +
      exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/22))) /
     (2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) < skoXC1 |
     skoXC1 *
     (2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) <=
     exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/22)) +
     exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/22)) |
     2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
     exp(skoXC1 * -1) * 1/2 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_7, plain,
    (skoXC1 *
     (2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
      exp(skoXC1 * -1) * 1/2) <=
     exp(skoXC1) * (3/2 + skoXC1 * (skoXC1 * 1/22)) +
     exp(skoXC1 * -1) * (-3/2 + skoXC1 * (skoXC1 * -1/22)) |
     2 + skoXC1 * (skoXC1 * 1/11) + exp(skoXC1) * 1/2 +
     exp(skoXC1 * -1) * 1/2 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_6])).

cnf(refute_0_8, plain,
    (exp(skoXC1 * -1) * (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     exp(skoXC1 * -1) <=
     -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_11, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -2/11) |
     -4 + skoXC1 * (skoXC1 * -2/11) <= exp(skoXC1)),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_15, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 |
     -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (exp(skoXC1) < -4 + skoXC1 * (skoXC1 * -2/11) |
     -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (0 < -4 + skoXC1 * (skoXC1 * -2/11) |
     -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_12, refute_0_17])).

cnf(refute_0_19, plain,
    (4 < skoXC1 * (skoXC1 * -2/11) |
     -4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_21, plain, (skoXC1 < 10357/2000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_22, plain, (skoXC1 * (skoXC1 * -2/11) <= 4),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (-4 + skoXC1 * (skoXC1 * -2/11) + exp(skoXC1) * -1 < exp(skoXC1 * -1)),
    inference(resolve, [], [refute_0_22, refute_0_19])).

cnf(refute_0_24, plain,
    (exp(skoXC1 * -1) * (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_8, refute_0_23])).

cnf(refute_0_25, plain,
    (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0 |
     exp(skoXC1 * -1) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(split, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/22) <= -3/2 |
     exp(skoXC1 * -1) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain, (-3/2 < skoXC1 * (1/2 + skoXC1 * 1/22)),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_28, plain,
    (exp(skoXC1 * -1) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (X_000104 < exp(X_000103) | ~ lgen(1, X_000104, exp(X_000103))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_30, plain,
    (~ lgen(1, X_000104,
            -(X_000103 ^ 5 + 30 * X_000103 ^ 4 + 420 * X_000103 ^ 3 +
             3360 * X_000103 ^ 2 + 15120 * X_000103 + 30240) /
            (X_000103 ^ 5 - 30 * X_000103 ^ 4 + 420 * X_000103 ^ 3 -
             3360 * X_000103 ^ 2 + 15120 * X_000103 - 30240)) |
     0 <= X_000103 | lgen(1, X_000104, exp(X_000103))),
    inference(subst, [], [exp_lower_bound_cf5])).

cnf(refute_0_31, plain,
    (X_000104 < exp(X_000103) |
     ~ lgen(1, X_000104,
            -(X_000103 ^ 5 + 30 * X_000103 ^ 4 + 420 * X_000103 ^ 3 +
             3360 * X_000103 ^ 2 + 15120 * X_000103 + 30240) /
            (X_000103 ^ 5 - 30 * X_000103 ^ 4 + 420 * X_000103 ^ 3 -
             3360 * X_000103 ^ 2 + 15120 * X_000103 - 30240)) |
     0 <= X_000103), inference(resolve, [], [refute_0_30, refute_0_29])).

cnf(refute_0_32, plain,
    (X_000104 < exp(X_000103) | 0 <= X_000103 |
     (-30240 +
      X_000103 *
      (-15120 +
       X_000103 *
       (-3360 + X_000103 * (-420 + X_000103 * (-30 + X_000103 * -1))))) /
     (-30240 +
      X_000103 *
      (15120 +
       X_000103 *
       (-3360 + X_000103 * (420 + X_000103 * (-30 + X_000103))))) <=
     X_000104), inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) < exp(skoXC1 * -1) |
     0 <= skoXC1 * -1 |
     (-30240 +
      skoXC1 * -1 *
      (-15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (-420 + skoXC1 * -1 * (-30 + skoXC1 * -1 * -1))))) /
     (-30240 +
      skoXC1 * -1 *
      (15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (420 + skoXC1 * -1 * (-30 + skoXC1 * -1))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(subst, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (0 <= skoXC1 * -1 |
     (-30240 +
      skoXC1 * -1 *
      (-15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (-420 + skoXC1 * -1 * (-30 + skoXC1 * -1 * -1))))) /
     (-30240 +
      skoXC1 * -1 *
      (15120 +
       skoXC1 * -1 *
       (-3360 +
        skoXC1 * -1 * (420 + skoXC1 * -1 * (-30 + skoXC1 * -1))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_28, refute_0_33])).

cnf(refute_0_35, plain,
    ((-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    ((-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_35, refute_0_20])).

cnf(refute_0_37, plain,
    ((-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_36, refute_0_3])).

cnf(refute_0_38, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/22) <= -3/2 |
     (-45360 +
      skoXC1 *
      (7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (-3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (123/22 + skoXC1 * (-19/22 + skoXC1 * 1/22))))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    ((-45360 +
      skoXC1 *
      (7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (-3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (123/22 + skoXC1 * (-19/22 + skoXC1 * 1/22))))))) /
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_38, refute_0_27])).

cnf(refute_0_40, plain,
    (0 <=
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) <=
     -45360 +
     skoXC1 *
     (7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (123/22 + skoXC1 * (-19/22 + skoXC1 * 1/22))))))),
    inference(resolve, [], [refute_0_39, refute_0_2])).

cnf(refute_0_41, plain,
    (exp(skoXC1) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -45360 +
     skoXC1 *
     (-52920 +
      skoXC1 *
      (-320040/11 +
       skoXC1 *
       (-108150/11 +
        skoXC1 *
        (-24225/11 +
         skoXC1 *
         (-7917/22 +
          skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11))))))) |
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <=
     -30240), inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (-30240 <
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_43, plain,
    (exp(skoXC1) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -45360 +
     skoXC1 *
     (-52920 +
      skoXC1 *
      (-320040/11 +
       skoXC1 *
       (-108150/11 +
        skoXC1 *
        (-24225/11 +
         skoXC1 *
         (-7917/22 +
          skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))),
    inference(resolve, [], [refute_0_41, refute_0_42])).

cnf(refute_0_44, plain,
    (0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     exp(skoXC1)), inference(split, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     exp(skoXC1)), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (-45360 <
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22))))))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_47, plain,
    ((-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     exp(skoXC1)), inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain,
    (exp(X_000165) < X_000166 | ~ lgen(1, exp(X_000165), X_000166)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_49, plain,
    (19/2 < X_000165 | X_000165 < 0 |
     ~ lgen(1,
            -(X_000165 ^ 7 + 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 +
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 +
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 + 17297280) /
            (X_000165 ^ 7 - 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 -
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 -
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 - 17297280),
            X_000166) | lgen(1, exp(X_000165), X_000166)),
    inference(subst, [], [exp_upper_bound_cf7])).

cnf(refute_0_50, plain,
    (19/2 < X_000165 | exp(X_000165) < X_000166 | X_000165 < 0 |
     ~ lgen(1,
            -(X_000165 ^ 7 + 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 +
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 +
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 + 17297280) /
            (X_000165 ^ 7 - 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 -
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 -
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 - 17297280),
            X_000166)),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    (19/2 < X_000165 | exp(X_000165) < X_000166 | X_000165 < 0 |
     X_000166 <=
     (-17297280 +
      X_000165 *
      (-8648640 +
       X_000165 *
       (-1995840 +
        X_000165 *
        (-277200 +
         X_000165 *
         (-25200 +
          X_000165 * (-1512 + X_000165 * (-56 + X_000165 * -1))))))) /
     (-17297280 +
      X_000165 *
      (8648640 +
       X_000165 *
       (-1995840 +
        X_000165 *
        (277200 +
         X_000165 *
         (-25200 + X_000165 * (1512 + X_000165 * (-56 + X_000165)))))))),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     19/2 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(subst, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (skoXC1 < 0 | 19/2 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_47, refute_0_52])).

cnf(refute_0_54, plain,
    (30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <
     (-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-1555/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (51/22 + skoXC1 * 1/11))))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     0 <=
     -25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) |
     (30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -784604620800 +
     skoXC1 *
     (-523069747200 +
      skoXC1 *
      (-136101772800 +
       skoXC1 *
       (-11481523200 +
        skoXC1 *
        (2396822400 +
         skoXC1 *
         (-725760 +
          skoXC1 *
          (843091200/11 +
           skoXC1 *
           (-64350720/11 +
            skoXC1 *
            (-30275280/11 +
             skoXC1 *
             (6064800/11 +
              skoXC1 *
              (-459438/11 +
               skoXC1 *
               (19017/11 +
                skoXC1 *
                (-1555/22 +
                 skoXC1 *
                 (623/22 + skoXC1 * (51/22 + skoXC1 * 1/11))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_55, plain,
    ((-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) <
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) <=
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_56, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-25945920 +
      skoXC1 *
      (4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (-189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 * (929/22 + skoXC1 * (-45/22 + skoXC1 * 1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) |
     0 <=
     -17297280 +
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1)))))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     -25945920 +
     skoXC1 *
     (4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (929/22 + skoXC1 * (-45/22 + skoXC1 * 1/22))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_57, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) |
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_58, plain,
    (X_000125 < exp(X_000124) | ~ lgen(1, X_000125, exp(X_000124))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_59, plain,
    (~ lgen(1, X_000125,
            -(X_000124 ^ 7 + 56 * X_000124 ^ 6 + 1512 * X_000124 ^ 5 +
             25200 * X_000124 ^ 4 + 277200 * X_000124 ^ 3 +
             1995840 * X_000124 ^ 2 + 8648640 * X_000124 + 17297280) /
            (X_000124 ^ 7 - 56 * X_000124 ^ 6 + 1512 * X_000124 ^ 5 -
             25200 * X_000124 ^ 4 + 277200 * X_000124 ^ 3 -
             1995840 * X_000124 ^ 2 + 8648640 * X_000124 - 17297280)) |
     0 <= X_000124 | lgen(1, X_000125, exp(X_000124))),
    inference(subst, [], [exp_lower_bound_cf7])).

cnf(refute_0_60, plain,
    (X_000125 < exp(X_000124) |
     ~ lgen(1, X_000125,
            -(X_000124 ^ 7 + 56 * X_000124 ^ 6 + 1512 * X_000124 ^ 5 +
             25200 * X_000124 ^ 4 + 277200 * X_000124 ^ 3 +
             1995840 * X_000124 ^ 2 + 8648640 * X_000124 + 17297280) /
            (X_000124 ^ 7 - 56 * X_000124 ^ 6 + 1512 * X_000124 ^ 5 -
             25200 * X_000124 ^ 4 + 277200 * X_000124 ^ 3 -
             1995840 * X_000124 ^ 2 + 8648640 * X_000124 - 17297280)) |
     0 <= X_000124), inference(resolve, [], [refute_0_59, refute_0_58])).

cnf(refute_0_61, plain,
    (X_000125 < exp(X_000124) | 0 <= X_000124 |
     (-17297280 +
      X_000124 *
      (-8648640 +
       X_000124 *
       (-1995840 +
        X_000124 *
        (-277200 +
         X_000124 *
         (-25200 +
          X_000124 * (-1512 + X_000124 * (-56 + X_000124 * -1))))))) /
     (-17297280 +
      X_000124 *
      (8648640 +
       X_000124 *
       (-1995840 +
        X_000124 *
        (277200 +
         X_000124 *
         (-25200 + X_000124 * (1512 + X_000124 * (-56 + X_000124))))))) <=
     X_000125), inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) < exp(skoXC1 * -1) |
     0 <= skoXC1 * -1 |
     (-17297280 +
      skoXC1 * -1 *
      (-8648640 +
       skoXC1 * -1 *
       (-1995840 +
        skoXC1 * -1 *
        (-277200 +
         skoXC1 * -1 *
         (-25200 +
          skoXC1 * -1 *
          (-1512 + skoXC1 * -1 * (-56 + skoXC1 * -1 * -1))))))) /
     (-17297280 +
      skoXC1 * -1 *
      (8648640 +
       skoXC1 * -1 *
       (-1995840 +
        skoXC1 * -1 *
        (277200 +
         skoXC1 * -1 *
         (-25200 +
          skoXC1 * -1 * (1512 + skoXC1 * -1 * (-56 + skoXC1 * -1))))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(subst, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (0 <= skoXC1 * -1 |
     (-17297280 +
      skoXC1 * -1 *
      (-8648640 +
       skoXC1 * -1 *
       (-1995840 +
        skoXC1 * -1 *
        (-277200 +
         skoXC1 * -1 *
         (-25200 +
          skoXC1 * -1 *
          (-1512 + skoXC1 * -1 * (-56 + skoXC1 * -1 * -1))))))) /
     (-17297280 +
      skoXC1 * -1 *
      (8648640 +
       skoXC1 * -1 *
       (-1995840 +
        skoXC1 * -1 *
        (277200 +
         skoXC1 * -1 *
         (-25200 +
          skoXC1 * -1 * (1512 + skoXC1 * -1 * (-56 + skoXC1 * -1))))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_28, refute_0_62])).

cnf(refute_0_64, plain,
    ((-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    ((-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_64, refute_0_20])).

cnf(refute_0_66, plain,
    ((-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_65, refute_0_57])).

cnf(refute_0_67, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/22) <= -3/2 |
     (-25945920 +
      skoXC1 *
      (4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (-189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 * (929/22 + skoXC1 * (-45/22 + skoXC1 * 1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    ((-25945920 +
      skoXC1 *
      (4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (-189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 * (929/22 + skoXC1 * (-45/22 + skoXC1 * 1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_67, refute_0_27])).

cnf(refute_0_69, plain,
    (0 <=
     -17297280 +
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1)))))) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) <=
     -25945920 +
     skoXC1 *
     (4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (929/22 + skoXC1 * (-45/22 + skoXC1 * 1/22))))))))),
    inference(resolve, [], [refute_0_68, refute_0_56])).

cnf(refute_0_70, plain,
    (exp(skoXC1) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -25945920 +
     skoXC1 *
     (-30270240 +
      skoXC1 *
      (-16752960 +
       skoXC1 *
       (-5753160 +
        skoXC1 *
        (-1330560 +
         skoXC1 *
         (-229572 +
          skoXC1 *
          (-315672/11 +
           skoXC1 *
           (-51935/22 +
            skoXC1 *
            (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11))))))))) |
     skoXC1 *
     (8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1)))))) <=
     -17297280), inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (-17297280 <
     skoXC1 *
     (8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1))))))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_72, plain,
    (exp(skoXC1) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -25945920 +
     skoXC1 *
     (-30270240 +
      skoXC1 *
      (-16752960 +
       skoXC1 *
       (-5753160 +
        skoXC1 *
        (-1330560 +
         skoXC1 *
         (-229572 +
          skoXC1 *
          (-315672/11 +
           skoXC1 *
           (-51935/22 +
            skoXC1 *
            (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (0 <=
     -25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     exp(skoXC1)), inference(split, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22)))))))) <=
     -25945920 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     exp(skoXC1)), inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (-25945920 <
     skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22))))))))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_76, plain,
    ((-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     exp(skoXC1)), inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (exp(X_000158) < X_000159 | ~ lgen(1, exp(X_000158), X_000159)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_78, plain,
    (7 < X_000158 | X_000158 < 0 |
     ~ lgen(1,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159) |
     lgen(1, exp(X_000158), X_000159)),
    inference(subst, [], [exp_upper_bound_cf5])).

cnf(refute_0_79, plain,
    (7 < X_000158 | exp(X_000158) < X_000159 | X_000158 < 0 |
     ~ lgen(1,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159)),
    inference(resolve, [], [refute_0_78, refute_0_77])).

cnf(refute_0_80, plain,
    (7 < X_000158 | exp(X_000158) < X_000159 | X_000158 < 0 |
     X_000159 <=
     (-30240 +
      X_000158 *
      (-15120 +
       X_000158 *
       (-3360 + X_000158 * (-420 + X_000158 * (-30 + X_000158 * -1))))) /
     (-30240 +
      X_000158 *
      (15120 +
       X_000158 *
       (-3360 + X_000158 * (420 + X_000158 * (-30 + X_000158)))))),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     7 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(subst, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (skoXC1 < 0 | 7 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_76, refute_0_81])).

cnf(refute_0_83, plain,
    (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <
     (-3113510400 +
      skoXC1 *
      (-2075673600 +
       skoXC1 *
       (-505491840 +
        skoXC1 *
        (-22498560 +
         skoXC1 *
         (14757120 +
          skoXC1 *
          (skoXC1 *
           (1287720/11 +
            skoXC1 *
            (156876/11 +
             skoXC1 *
             (108498/11 +
              skoXC1 *
              (90792/11 +
               skoXC1 *
               (1909/2 +
                skoXC1 *
                (1901/22 + skoXC1 * (87/22 + skoXC1 * 1/11))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     0 <=
     -25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -3113510400 +
     skoXC1 *
     (-2075673600 +
      skoXC1 *
      (-505491840 +
       skoXC1 *
       (-22498560 +
        skoXC1 *
        (14757120 +
         skoXC1 *
         (skoXC1 *
          (1287720/11 +
           skoXC1 *
           (156876/11 +
            skoXC1 *
            (108498/11 +
             skoXC1 *
             (90792/11 +
              skoXC1 *
              (1909/2 +
               skoXC1 *
               (1901/22 + skoXC1 * (87/22 + skoXC1 * 1/11))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_84, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_85, plain,
    (exp(X_000138) < X_000139 | ~ lgen(1, exp(X_000138), X_000139)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_86, plain,
    (4 < X_000138 | X_000138 < 0 |
     ~ lgen(1,
            -(X_000138 ^ 3 + 12 * X_000138 ^ 2 + 60 * X_000138 + 120) /
            (X_000138 ^ 3 - 12 * X_000138 ^ 2 + 60 * X_000138 - 120),
            X_000139) | lgen(1, exp(X_000138), X_000139)),
    inference(subst, [], [exp_upper_bound_cf3])).

cnf(refute_0_87, plain,
    (4 < X_000138 | exp(X_000138) < X_000139 | X_000138 < 0 |
     ~ lgen(1,
            -(X_000138 ^ 3 + 12 * X_000138 ^ 2 + 60 * X_000138 + 120) /
            (X_000138 ^ 3 - 12 * X_000138 ^ 2 + 60 * X_000138 - 120),
            X_000139)),
    inference(resolve, [], [refute_0_86, refute_0_85])).

cnf(refute_0_88, plain,
    (4 < X_000138 | exp(X_000138) < X_000139 | X_000138 < 0 |
     X_000139 <=
     (-120 + X_000138 * (-60 + X_000138 * (-12 + X_000138 * -1))) /
     (-120 + X_000138 * (60 + X_000138 * (-12 + X_000138)))),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) |
     4 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (skoXC1 < 0 | 4 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_76, refute_0_89])).

cnf(refute_0_91, plain, (4 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_92, plain,
    (4 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_91, refute_0_90])).

cnf(refute_0_93, plain,
    (4 < skoXC1 | 0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_92, refute_0_84])).

cnf(refute_0_94, plain,
    (4 < skoXC1 | skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120 |
     (-3113510400 +
      skoXC1 *
      (-2075673600 +
       skoXC1 *
       (-505491840 +
        skoXC1 *
        (-22498560 +
         skoXC1 *
         (14757120 +
          skoXC1 *
          (skoXC1 *
           (1287720/11 +
            skoXC1 *
            (156876/11 +
             skoXC1 *
             (108498/11 +
              skoXC1 *
              (90792/11 +
               skoXC1 *
               (1909/2 +
                skoXC1 *
                (1901/22 + skoXC1 * (87/22 + skoXC1 * 1/11))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (-120 < skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) | 4 < skoXC1),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_96, plain,
    (4 < skoXC1 |
     (-3113510400 +
      skoXC1 *
      (-2075673600 +
       skoXC1 *
       (-505491840 +
        skoXC1 *
        (-22498560 +
         skoXC1 *
         (14757120 +
          skoXC1 *
          (skoXC1 *
           (1287720/11 +
            skoXC1 *
            (156876/11 +
             skoXC1 *
             (108498/11 +
              skoXC1 *
              (90792/11 +
               skoXC1 *
               (1909/2 +
                skoXC1 *
                (1901/22 + skoXC1 * (87/22 + skoXC1 * 1/11))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(resolve, [], [refute_0_94, refute_0_95])).

cnf(refute_0_97, plain,
    (4 < skoXC1 |
     0 <=
     -25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -3113510400 +
     skoXC1 *
     (-2075673600 +
      skoXC1 *
      (-505491840 +
       skoXC1 *
       (-22498560 +
        skoXC1 *
        (14757120 +
         skoXC1 *
         (skoXC1 *
          (1287720/11 +
           skoXC1 *
           (156876/11 +
            skoXC1 *
            (108498/11 +
             skoXC1 *
             (90792/11 +
              skoXC1 *
              (1909/2 +
               skoXC1 *
               (1901/22 + skoXC1 * (87/22 + skoXC1 * 1/11))))))))))))),
    inference(resolve, [], [refute_0_96, refute_0_83])).

cnf(refute_0_98, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (3144960 +
          skoXC1 *
          (skoXC1 *
           (-64872 +
            skoXC1 *
            (-20592 +
             skoXC1 *
             (-102984/11 +
              skoXC1 *
              (-1024 +
               skoXC1 * (-89 + skoXC1 * (-4 + skoXC1 * -1/11)))))))))))) <=
     0 |
     skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22)))))))) <=
     -25945920), inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <
     (-5443200 +
      skoXC1 *
      (-3628800 +
       skoXC1 *
       (-9465120/11 +
        skoXC1 *
        (-262080/11 +
         skoXC1 *
         (323640/11 +
          skoXC1 *
          (720/11 +
           skoXC1 *
           (780/11 +
            skoXC1 *
            (273/11 +
             skoXC1 *
             (381/22 +
              skoXC1 *
              (315/22 + skoXC1 * (35/22 + skoXC1 * 1/11))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -5443200 +
     skoXC1 *
     (-3628800 +
      skoXC1 *
      (-9465120/11 +
       skoXC1 *
       (-262080/11 +
        skoXC1 *
        (323640/11 +
         skoXC1 *
         (720/11 +
          skoXC1 *
          (780/11 +
           skoXC1 *
           (273/11 +
            skoXC1 *
            (381/22 +
             skoXC1 * (315/22 + skoXC1 * (35/22 + skoXC1 * 1/11))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_100, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     4 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_102, plain,
    (skoXC1 < 0 | 4 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_47, refute_0_101])).

cnf(refute_0_103, plain,
    (4 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_91, refute_0_102])).

cnf(refute_0_104, plain,
    (4 < skoXC1 | 0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_103, refute_0_100])).

cnf(refute_0_105, plain,
    (4 < skoXC1 | skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120 |
     (-5443200 +
      skoXC1 *
      (-3628800 +
       skoXC1 *
       (-9465120/11 +
        skoXC1 *
        (-262080/11 +
         skoXC1 *
         (323640/11 +
          skoXC1 *
          (720/11 +
           skoXC1 *
           (780/11 +
            skoXC1 *
            (273/11 +
             skoXC1 *
             (381/22 +
              skoXC1 *
              (315/22 + skoXC1 * (35/22 + skoXC1 * 1/11))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (4 < skoXC1 |
     (-5443200 +
      skoXC1 *
      (-3628800 +
       skoXC1 *
       (-9465120/11 +
        skoXC1 *
        (-262080/11 +
         skoXC1 *
         (323640/11 +
          skoXC1 *
          (720/11 +
           skoXC1 *
           (780/11 +
            skoXC1 *
            (273/11 +
             skoXC1 *
             (381/22 +
              skoXC1 *
              (315/22 + skoXC1 * (35/22 + skoXC1 * 1/11))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(resolve, [], [refute_0_105, refute_0_95])).

cnf(refute_0_107, plain,
    (4 < skoXC1 |
     0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -5443200 +
     skoXC1 *
     (-3628800 +
      skoXC1 *
      (-9465120/11 +
       skoXC1 *
       (-262080/11 +
        skoXC1 *
        (323640/11 +
         skoXC1 *
         (720/11 +
          skoXC1 *
          (780/11 +
           skoXC1 *
           (273/11 +
            skoXC1 *
            (381/22 +
             skoXC1 * (315/22 + skoXC1 * (35/22 + skoXC1 * 1/11))))))))))),
    inference(resolve, [], [refute_0_106, refute_0_99])).

cnf(refute_0_108, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (60480/11 +
          skoXC1 *
          (skoXC1 *
           (-1506/11 +
            skoXC1 *
            (-36 +
             skoXC1 *
             (-173/11 + skoXC1 * (-18/11 + skoXC1 * -1/11)))))))))) <= 0 |
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360), inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <
     (-21600 +
      skoXC1 *
      (-14400 +
       skoXC1 *
       (-34920/11 +
        skoXC1 *
        (720/11 +
         skoXC1 *
         (1644/11 +
          skoXC1 *
          (-24/11 +
           skoXC1 *
           (-3/2 +
            skoXC1 * (9/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) |
     0 <=
     -180 +
     skoXC1 *
     (-30 +
      skoXC1 *
      (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22)))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     -21600 +
     skoXC1 *
     (-14400 +
      skoXC1 *
      (-34920/11 +
       skoXC1 *
       (720/11 +
        skoXC1 *
        (1644/11 +
         skoXC1 *
         (-24/11 +
          skoXC1 *
          (-3/2 +
           skoXC1 * (9/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_110, plain,
    ((-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) <
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) |
     0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_111, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-180 +
      skoXC1 *
      (30 +
       skoXC1 *
       (72/11 + skoXC1 * (-39/22 + skoXC1 * (-1/22 + skoXC1 * 1/22))))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) |
     0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     -180 +
     skoXC1 *
     (30 +
      skoXC1 *
      (72/11 + skoXC1 * (-39/22 + skoXC1 * (-1/22 + skoXC1 * 1/22))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_112, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) |
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_113, plain,
    (X_000083 < exp(X_000082) | ~ lgen(1, X_000083, exp(X_000082))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_114, plain,
    (~ lgen(1, X_000083,
            -(X_000082 ^ 3 + 12 * X_000082 ^ 2 + 60 * X_000082 + 120) /
            (X_000082 ^ 3 - 12 * X_000082 ^ 2 + 60 * X_000082 - 120)) |
     0 <= X_000082 | lgen(1, X_000083, exp(X_000082))),
    inference(subst, [], [exp_lower_bound_cf3])).

cnf(refute_0_115, plain,
    (X_000083 < exp(X_000082) |
     ~ lgen(1, X_000083,
            -(X_000082 ^ 3 + 12 * X_000082 ^ 2 + 60 * X_000082 + 120) /
            (X_000082 ^ 3 - 12 * X_000082 ^ 2 + 60 * X_000082 - 120)) |
     0 <= X_000082), inference(resolve, [], [refute_0_114, refute_0_113])).

cnf(refute_0_116, plain,
    (X_000083 < exp(X_000082) | 0 <= X_000082 |
     (-120 + X_000082 * (-60 + X_000082 * (-12 + X_000082 * -1))) /
     (-120 + X_000082 * (60 + X_000082 * (-12 + X_000082))) <= X_000083),
    inference(arithmetic, [], [refute_0_115])).

cnf(refute_0_117, plain,
    ((skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) < exp(skoXC1 * -1) |
     0 <= skoXC1 * -1 |
     (-120 +
      skoXC1 * -1 * (-60 + skoXC1 * -1 * (-12 + skoXC1 * -1 * -1))) /
     (-120 + skoXC1 * -1 * (60 + skoXC1 * -1 * (-12 + skoXC1 * -1))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(subst, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (0 <= skoXC1 * -1 |
     (-120 +
      skoXC1 * -1 * (-60 + skoXC1 * -1 * (-12 + skoXC1 * -1 * -1))) /
     (-120 + skoXC1 * -1 * (60 + skoXC1 * -1 * (-12 + skoXC1 * -1))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_28, refute_0_117])).

cnf(refute_0_119, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_119, refute_0_20])).

cnf(refute_0_121, plain,
    ((-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) *
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_120, refute_0_112])).

cnf(refute_0_122, plain,
    (skoXC1 * (1/2 + skoXC1 * 1/22) <= -3/2 |
     (-180 +
      skoXC1 *
      (30 +
       skoXC1 *
       (72/11 + skoXC1 * (-39/22 + skoXC1 * (-1/22 + skoXC1 * 1/22))))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    ((-180 +
      skoXC1 *
      (30 +
       skoXC1 *
       (72/11 + skoXC1 * (-39/22 + skoXC1 * (-1/22 + skoXC1 * 1/22))))) /
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_122, refute_0_27])).

cnf(refute_0_124, plain,
    (0 <= -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) *
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) <=
     -180 +
     skoXC1 *
     (30 +
      skoXC1 *
      (72/11 + skoXC1 * (-39/22 + skoXC1 * (-1/22 + skoXC1 * 1/22))))),
    inference(resolve, [], [refute_0_123, refute_0_111])).

cnf(refute_0_125, plain,
    (exp(skoXC1) *
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     -180 +
     skoXC1 *
     (-210 +
      skoXC1 *
      (-1248/11 +
       skoXC1 *
       (-807/22 +
        skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11))))) |
     skoXC1 * (60 + skoXC1 * (12 + skoXC1)) <= -120),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain, (-120 < skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_127, plain,
    (exp(skoXC1) *
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     -180 +
     skoXC1 *
     (-210 +
      skoXC1 *
      (-1248/11 +
       skoXC1 *
       (-807/22 +
        skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (0 <=
     -180 +
     skoXC1 *
     (-30 +
      skoXC1 *
      (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22)))) |
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     exp(skoXC1)), inference(split, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 *
     (30 +
      skoXC1 *
      (-72/11 + skoXC1 * (-39/22 + skoXC1 * (1/22 + skoXC1 * 1/22)))) <=
     -180 |
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     exp(skoXC1)), inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (-180 <
     skoXC1 *
     (30 +
      skoXC1 *
      (-72/11 + skoXC1 * (-39/22 + skoXC1 * (1/22 + skoXC1 * 1/22))))),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_131, plain,
    ((-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     exp(skoXC1)), inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) |
     4 < skoXC1 |
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(subst, [], [refute_0_88])).

cnf(refute_0_133, plain,
    (skoXC1 < 0 | 4 < skoXC1 |
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain,
    (4 < skoXC1 |
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     (-120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1))) /
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_91, refute_0_133])).

cnf(refute_0_135, plain,
    (4 < skoXC1 | 0 <= -120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)) |
     -120 + skoXC1 * (-60 + skoXC1 * (-12 + skoXC1 * -1)) <=
     (-180 +
      skoXC1 *
      (-210 +
       skoXC1 *
       (-1248/11 +
        skoXC1 *
        (-807/22 +
         skoXC1 * (-15/2 + skoXC1 * (-23/22 + skoXC1 * -1/11)))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) *
     (-120 + skoXC1 * (60 + skoXC1 * (-12 + skoXC1)))),
    inference(resolve, [], [refute_0_134, refute_0_110])).

cnf(refute_0_136, plain,
    (4 < skoXC1 | skoXC1 * (-60 + skoXC1 * (12 + skoXC1 * -1)) <= -120 |
     (-21600 +
      skoXC1 *
      (-14400 +
       skoXC1 *
       (-34920/11 +
        skoXC1 *
        (720/11 +
         skoXC1 *
         (1644/11 +
          skoXC1 *
          (-24/11 +
           skoXC1 *
           (-3/2 +
            skoXC1 * (9/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain,
    (4 < skoXC1 |
     (-21600 +
      skoXC1 *
      (-14400 +
       skoXC1 *
       (-34920/11 +
        skoXC1 *
        (720/11 +
         skoXC1 *
         (1644/11 +
          skoXC1 *
          (-24/11 +
           skoXC1 *
           (-3/2 +
            skoXC1 * (9/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))) /
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))),
    inference(resolve, [], [refute_0_136, refute_0_95])).

cnf(refute_0_138, plain,
    (4 < skoXC1 |
     0 <=
     -180 +
     skoXC1 *
     (-30 +
      skoXC1 *
      (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22)))) |
     (120 + skoXC1 * (60 + skoXC1 * (12 + skoXC1))) *
     (-180 +
      skoXC1 *
      (-30 +
       skoXC1 *
       (72/11 + skoXC1 * (39/22 + skoXC1 * (-1/22 + skoXC1 * -1/22))))) <=
     -21600 +
     skoXC1 *
     (-14400 +
      skoXC1 *
      (-34920/11 +
       skoXC1 *
       (720/11 +
        skoXC1 *
        (1644/11 +
         skoXC1 *
         (-24/11 +
          skoXC1 *
          (-3/2 +
           skoXC1 * (9/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))),
    inference(resolve, [], [refute_0_137, refute_0_109])).

cnf(refute_0_139, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (240/11 +
          skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))) <= 0 |
     skoXC1 *
     (30 +
      skoXC1 *
      (-72/11 + skoXC1 * (-39/22 + skoXC1 * (1/22 + skoXC1 * 1/22)))) <=
     -180), inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (-180 <
     skoXC1 *
     (30 +
      skoXC1 *
      (-72/11 + skoXC1 * (-39/22 + skoXC1 * (1/22 + skoXC1 * 1/22)))) |
     4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (240/11 +
          skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))) <= 0),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_141, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (240/11 +
          skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))) <= 0),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (240/11 +
          skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))) <= 0 |
     skoXC1 <= 4), inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_143, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (240/11 +
          skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))) <= 0),
    inference(resolve, [], [refute_0_142, refute_0_141])).

cnf(refute_0_144, plain,
    (-45360 <
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) |
     4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (60480/11 +
          skoXC1 *
          (skoXC1 *
           (-1506/11 +
            skoXC1 *
            (-36 +
             skoXC1 *
             (-173/11 + skoXC1 * (-18/11 + skoXC1 * -1/11)))))))))) <= 0),
    inference(decision, [], [refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_145, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (60480/11 +
          skoXC1 *
          (skoXC1 *
           (-1506/11 +
            skoXC1 *
            (-36 +
             skoXC1 *
             (-173/11 + skoXC1 * (-18/11 + skoXC1 * -1/11)))))))))) <= 0),
    inference(resolve, [], [refute_0_108, refute_0_144])).

cnf(refute_0_146, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (60480/11 +
          skoXC1 *
          (skoXC1 *
           (-1506/11 +
            skoXC1 *
            (-36 +
             skoXC1 *
             (-173/11 + skoXC1 * (-18/11 + skoXC1 * -1/11)))))))))) <= 0 |
     skoXC1 <= 4),
    inference(decision, [], [refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_147, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (60480/11 +
          skoXC1 *
          (skoXC1 *
           (-1506/11 +
            skoXC1 *
            (-36 +
             skoXC1 *
             (-173/11 + skoXC1 * (-18/11 + skoXC1 * -1/11)))))))))) <= 0),
    inference(resolve, [], [refute_0_146, refute_0_145])).

cnf(refute_0_148, plain,
    (-25945920 <
     skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22)))))))) |
     4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (3144960 +
          skoXC1 *
          (skoXC1 *
           (-64872 +
            skoXC1 *
            (-20592 +
             skoXC1 *
             (-102984/11 +
              skoXC1 *
              (-1024 +
               skoXC1 * (-89 + skoXC1 * (-4 + skoXC1 * -1/11)))))))))))) <=
     0),
    inference(decision, [],
              [refute_0_147, refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_149, plain,
    (4 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (3144960 +
          skoXC1 *
          (skoXC1 *
           (-64872 +
            skoXC1 *
            (-20592 +
             skoXC1 *
             (-102984/11 +
              skoXC1 *
              (-1024 +
               skoXC1 * (-89 + skoXC1 * (-4 + skoXC1 * -1/11)))))))))))) <=
     0), inference(resolve, [], [refute_0_98, refute_0_148])).

cnf(refute_0_150, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (3144960 +
          skoXC1 *
          (skoXC1 *
           (-64872 +
            skoXC1 *
            (-20592 +
             skoXC1 *
             (-102984/11 +
              skoXC1 *
              (-1024 +
               skoXC1 * (-89 + skoXC1 * (-4 + skoXC1 * -1/11)))))))))))) <=
     0 | skoXC1 <= 4),
    inference(decision, [],
              [refute_0_147, refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_151, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (3144960 +
          skoXC1 *
          (skoXC1 *
           (-64872 +
            skoXC1 *
            (-20592 +
             skoXC1 *
             (-102984/11 +
              skoXC1 *
              (-1024 +
               skoXC1 * (-89 + skoXC1 * (-4 + skoXC1 * -1/11)))))))))))) <=
     0), inference(resolve, [], [refute_0_150, refute_0_149])).

cnf(refute_0_152, plain, (7 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_151, refute_0_147, refute_0_143, refute_0_21,
               refute_0_20])).

cnf(refute_0_153, plain,
    (7 < skoXC1 |
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_152, refute_0_82])).

cnf(refute_0_154, plain, (skoXC1 <= 7),
    inference(decision, [],
              [refute_0_151, refute_0_147, refute_0_143, refute_0_21,
               refute_0_20])).

cnf(refute_0_155, plain,
    ((-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_154, refute_0_153])).

cnf(refute_0_156, plain,
    (0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) <=
     (-25945920 +
      skoXC1 *
      (-30270240 +
       skoXC1 *
       (-16752960 +
        skoXC1 *
        (-5753160 +
         skoXC1 *
         (-1330560 +
          skoXC1 *
          (-229572 +
           skoXC1 *
           (-315672/11 +
            skoXC1 *
            (-51935/22 +
             skoXC1 *
             (-283/2 + skoXC1 * (-111/22 + skoXC1 * -1/11)))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_155, refute_0_55])).

cnf(refute_0_157, plain,
    (skoXC1 *
     (-15120 +
      skoXC1 * (3360 + skoXC1 * (-420 + skoXC1 * (30 + skoXC1 * -1)))) <=
     -30240 |
     (-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-1555/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (51/22 + skoXC1 * 1/11))))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))),
    inference(arithmetic, [], [refute_0_156])).

cnf(refute_0_158, plain,
    (-45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_159, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) *
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22) <= 0 |
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_160, plain,
    (0 * (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0 |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_161, plain,
    (0 <
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) |
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) < exp(skoXC1 * -1)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_162, plain,
    (0 <
     (skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
      exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))) /
     (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_28, refute_0_161])).

cnf(refute_0_163, plain,
    (0 * (3/2 + skoXC1 * (1/2 + skoXC1 * 1/22)) <
     skoXC1 * (-2 + skoXC1 * (skoXC1 * -1/11)) +
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_160, refute_0_162])).

cnf(refute_0_164, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     skoXC1 * (1/2 + skoXC1 * 1/22) <= -3/2),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     exp(skoXC1) * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(resolve, [], [refute_0_164, refute_0_27])).

cnf(refute_0_166, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) < exp(skoXC1) |
     3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22) <= 0),
    inference(split, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) < exp(skoXC1) |
     skoXC1 * (-1/2 + skoXC1 * 1/22) <= -3/2),
    inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain, (-3/2 < skoXC1 * (-1/2 + skoXC1 * 1/22)),
    inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_169, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) < exp(skoXC1)),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (~ lgen(0, exp(X_000158), X_000159) | exp(X_000158) <= X_000159),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_171, plain,
    (7 < X_000158 | X_000158 < 0 |
     ~ lgen(0,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159) |
     lgen(0, exp(X_000158), X_000159)),
    inference(subst, [], [exp_upper_bound_cf5])).

cnf(refute_0_172, plain,
    (7 < X_000158 | X_000158 < 0 |
     ~ lgen(0,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159) |
     exp(X_000158) <= X_000159),
    inference(resolve, [], [refute_0_171, refute_0_170])).

cnf(refute_0_173, plain,
    (7 < X_000158 | X_000158 < 0 |
     X_000159 <
     (-30240 +
      X_000158 *
      (-15120 +
       X_000158 *
       (-3360 + X_000158 * (-420 + X_000158 * (-30 + X_000158 * -1))))) /
     (-30240 +
      X_000158 *
      (15120 +
       X_000158 *
       (-3360 + X_000158 * (420 + X_000158 * (-30 + X_000158))))) |
     exp(X_000158) <= X_000159),
    inference(arithmetic, [], [refute_0_172])).

cnf(refute_0_174, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     7 < skoXC1 |
     exp(skoXC1) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(subst, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     7 < skoXC1), inference(resolve, [], [refute_0_174, refute_0_169])).

cnf(refute_0_176, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     7 < skoXC1), inference(resolve, [], [refute_0_152, refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_154, refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) *
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_159, refute_0_177])).

cnf(refute_0_179, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     skoXC1 * (-1/2 + skoXC1 * 1/22) <= -3/2),
    inference(arithmetic, [], [refute_0_178])).

cnf(refute_0_180, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_179, refute_0_168])).

cnf(refute_0_181, plain,
    (-45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) |
     0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))),
    inference(resolve, [], [refute_0_158, refute_0_180])).

cnf(refute_0_182, plain,
    (-45360 <
     skoXC1 *
     (-52920 +
      skoXC1 *
      (320040/11 +
       skoXC1 *
       (-108150/11 +
        skoXC1 *
        (24225/11 +
         skoXC1 *
         (-7917/22 +
          skoXC1 * (903/22 + skoXC1 * (-59/22 + skoXC1 * 1/11))))))) |
     skoXC1 *
     (-15120 +
      skoXC1 * (3360 + skoXC1 * (-420 + skoXC1 * (30 + skoXC1 * -1)))) <=
     -30240), inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (-45360 <
     skoXC1 *
     (-52920 +
      skoXC1 *
      (320040/11 +
       skoXC1 *
       (-108150/11 +
        skoXC1 *
        (24225/11 +
         skoXC1 *
         (-7917/22 +
          skoXC1 * (903/22 + skoXC1 * (-59/22 + skoXC1 * 1/11))))))) |
     -30240 <
     skoXC1 *
     (-15120 +
      skoXC1 * (3360 + skoXC1 * (-420 + skoXC1 * (30 + skoXC1 * -1))))),
    inference(decision, [],
              [refute_0_151, refute_0_147, refute_0_143, refute_0_21,
               refute_0_20])).

cnf(refute_0_184, plain,
    (-45360 <
     skoXC1 *
     (-52920 +
      skoXC1 *
      (320040/11 +
       skoXC1 *
       (-108150/11 +
        skoXC1 *
        (24225/11 +
         skoXC1 *
         (-7917/22 +
          skoXC1 * (903/22 + skoXC1 * (-59/22 + skoXC1 * 1/11)))))))),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (-30240 <
     skoXC1 *
     (-15120 +
      skoXC1 * (3360 + skoXC1 * (-420 + skoXC1 * (30 + skoXC1 * -1))))),
    inference(decision, [],
              [refute_0_184, refute_0_151, refute_0_147, refute_0_143,
               refute_0_21, refute_0_20])).

cnf(refute_0_186, plain,
    ((-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-1555/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (51/22 + skoXC1 * 1/11))))))))))))))) /
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))),
    inference(resolve, [], [refute_0_157, refute_0_185])).

cnf(refute_0_187, plain,
    (0 <=
     -25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) |
     (30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))) *
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) <=
     -784604620800 +
     skoXC1 *
     (-523069747200 +
      skoXC1 *
      (-136101772800 +
       skoXC1 *
       (-11481523200 +
        skoXC1 *
        (2396822400 +
         skoXC1 *
         (-725760 +
          skoXC1 *
          (843091200/11 +
           skoXC1 *
           (-64350720/11 +
            skoXC1 *
            (-30275280/11 +
             skoXC1 *
             (6064800/11 +
              skoXC1 *
              (-459438/11 +
               skoXC1 *
               (19017/11 +
                skoXC1 *
                (-1555/22 +
                 skoXC1 *
                 (623/22 + skoXC1 * (51/22 + skoXC1 * 1/11))))))))))))))),
    inference(resolve, [], [refute_0_186, refute_0_54])).

cnf(refute_0_188, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (792529920 +
          skoXC1 *
          (skoXC1 *
           (241920 +
            skoXC1 *
            (skoXC1 *
             (-10805760/11 +
              skoXC1 *
              (skoXC1 *
               (-49290/11 +
                skoXC1 *
                (-52 +
                 skoXC1 *
                 (-349/11 +
                  skoXC1 * (-26/11 + skoXC1 * -1/11)))))))))))))) <= 0 |
     skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22)))))))) <=
     -25945920), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1)))) <
     (-1371686400 +
      skoXC1 *
      (-914457600 +
       skoXC1 *
       (-2552860800/11 +
        skoXC1 *
        (-177811200/11 +
         skoXC1 *
         (56851200/11 +
          skoXC1 *
          (645120/11 +
           skoXC1 *
           (1207080/11 +
            skoXC1 *
            (-115920/11 +
             skoXC1 *
             (-57270/11 +
              skoXC1 *
              (11100/11 +
               skoXC1 *
               (-1113/22 +
                skoXC1 *
                (-27/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -1371686400 +
     skoXC1 *
     (-914457600 +
      skoXC1 *
      (-2552860800/11 +
       skoXC1 *
       (-177811200/11 +
        skoXC1 *
        (56851200/11 +
         skoXC1 *
         (645120/11 +
          skoXC1 *
          (1207080/11 +
           skoXC1 *
           (-115920/11 +
            skoXC1 *
            (-57270/11 +
             skoXC1 *
             (11100/11 +
              skoXC1 *
              (-1113/22 +
               skoXC1 *
               (-27/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_190, plain,
    ((-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1))))) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_191, plain,
    (skoXC1 < 0 |
     exp(skoXC1) <
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) |
     7 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(subst, [], [refute_0_80])).

cnf(refute_0_192, plain,
    (skoXC1 < 0 | 7 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_47, refute_0_191])).

cnf(refute_0_193, plain,
    (7 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_152, refute_0_192])).

cnf(refute_0_194, plain,
    ((-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-30240 +
      skoXC1 *
      (-15120 +
       skoXC1 *
       (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1))))) /
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_154, refute_0_193])).

cnf(refute_0_195, plain,
    (0 <=
     -30240 +
     skoXC1 *
     (15120 +
      skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))) |
     -30240 +
     skoXC1 *
     (-15120 +
      skoXC1 * (-3360 + skoXC1 * (-420 + skoXC1 * (-30 + skoXC1 * -1)))) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (-3360 + skoXC1 * (420 + skoXC1 * (-30 + skoXC1)))))),
    inference(resolve, [], [refute_0_194, refute_0_190])).

cnf(refute_0_196, plain,
    (skoXC1 *
     (-15120 +
      skoXC1 * (3360 + skoXC1 * (-420 + skoXC1 * (30 + skoXC1 * -1)))) <=
     -30240 |
     (-1371686400 +
      skoXC1 *
      (-914457600 +
       skoXC1 *
       (-2552860800/11 +
        skoXC1 *
        (-177811200/11 +
         skoXC1 *
         (56851200/11 +
          skoXC1 *
          (645120/11 +
           skoXC1 *
           (1207080/11 +
            skoXC1 *
            (-115920/11 +
             skoXC1 *
             (-57270/11 +
              skoXC1 *
              (11100/11 +
               skoXC1 *
               (-1113/22 +
                skoXC1 *
                (-27/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))),
    inference(arithmetic, [], [refute_0_195])).

cnf(refute_0_197, plain,
    ((-1371686400 +
      skoXC1 *
      (-914457600 +
       skoXC1 *
       (-2552860800/11 +
        skoXC1 *
        (-177811200/11 +
         skoXC1 *
         (56851200/11 +
          skoXC1 *
          (645120/11 +
           skoXC1 *
           (1207080/11 +
            skoXC1 *
            (-115920/11 +
             skoXC1 *
             (-57270/11 +
              skoXC1 *
              (11100/11 +
               skoXC1 *
               (-1113/22 +
                skoXC1 *
                (-27/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     30240 +
     skoXC1 *
     (15120 + skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))),
    inference(resolve, [], [refute_0_196, refute_0_185])).

cnf(refute_0_198, plain,
    (0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (30240 +
      skoXC1 *
      (15120 +
       skoXC1 * (3360 + skoXC1 * (420 + skoXC1 * (30 + skoXC1))))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -1371686400 +
     skoXC1 *
     (-914457600 +
      skoXC1 *
      (-2552860800/11 +
       skoXC1 *
       (-177811200/11 +
        skoXC1 *
        (56851200/11 +
         skoXC1 *
         (645120/11 +
          skoXC1 *
          (1207080/11 +
           skoXC1 *
           (-115920/11 +
            skoXC1 *
            (-57270/11 +
             skoXC1 *
             (11100/11 +
              skoXC1 *
              (-1113/22 +
               skoXC1 *
               (-27/22 + skoXC1 * (-1/22 + skoXC1 * 1/11))))))))))))),
    inference(resolve, [], [refute_0_197, refute_0_189])).

cnf(refute_0_199, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (15240960/11 +
          skoXC1 *
          (skoXC1 *
           (-60480/11 +
            skoXC1 *
            (skoXC1 *
             (-1680 +
              skoXC1 *
              (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))))))) <= 0 |
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360), inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    (-45360 <
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (15240960/11 +
          skoXC1 *
          (skoXC1 *
           (-60480/11 +
            skoXC1 *
            (skoXC1 *
             (-1680 +
              skoXC1 *
              (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))))))) <= 0),
    inference(decision, [],
              [refute_0_184, refute_0_151, refute_0_147, refute_0_143,
               refute_0_21, refute_0_20])).

cnf(refute_0_201, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (15240960/11 +
          skoXC1 *
          (skoXC1 *
           (-60480/11 +
            skoXC1 *
            (skoXC1 *
             (-1680 +
              skoXC1 *
              (skoXC1 * (-1 + skoXC1 * (skoXC1 * -1/11)))))))))))) <= 0),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (-25945920 <
     skoXC1 *
     (4324320 +
      skoXC1 *
      (-544320 +
       skoXC1 *
       (-189000 +
        skoXC1 *
        (-10080 +
         skoXC1 *
         (2268 +
          skoXC1 *
          (5208/11 +
           skoXC1 * (929/22 + skoXC1 * (45/22 + skoXC1 * 1/22)))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (792529920 +
          skoXC1 *
          (skoXC1 *
           (241920 +
            skoXC1 *
            (skoXC1 *
             (-10805760/11 +
              skoXC1 *
              (skoXC1 *
               (-49290/11 +
                skoXC1 *
                (-52 +
                 skoXC1 *
                 (-349/11 +
                  skoXC1 * (-26/11 + skoXC1 * -1/11)))))))))))))) <= 0),
    inference(decision, [],
              [refute_0_201, refute_0_184, refute_0_151, refute_0_147,
               refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_203, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (792529920 +
          skoXC1 *
          (skoXC1 *
           (241920 +
            skoXC1 *
            (skoXC1 *
             (-10805760/11 +
              skoXC1 *
              (skoXC1 *
               (-49290/11 +
                skoXC1 *
                (-52 +
                 skoXC1 *
                 (-349/11 +
                  skoXC1 * (-26/11 + skoXC1 * -1/11)))))))))))))) <= 0),
    inference(resolve, [], [refute_0_188, refute_0_202])).

cnf(refute_0_204, plain, (19/2 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_203, refute_0_201, refute_0_184, refute_0_151,
               refute_0_147, refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_205, plain,
    (19/2 < skoXC1 |
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_204, refute_0_53])).

cnf(refute_0_206, plain, (skoXC1 <= 19/2),
    inference(decision, [],
              [refute_0_203, refute_0_201, refute_0_184, refute_0_151,
               refute_0_147, refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_207, plain,
    ((-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_206, refute_0_205])).

cnf(refute_0_208, plain,
    (0 <=
     -17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))) |
     -17297280 +
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1)))))) <=
     (-45360 +
      skoXC1 *
      (-52920 +
       skoXC1 *
       (-320040/11 +
        skoXC1 *
        (-108150/11 +
         skoXC1 *
         (-24225/11 +
          skoXC1 *
          (-7917/22 +
           skoXC1 * (-903/22 + skoXC1 * (-59/22 + skoXC1 * -1/11)))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) *
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_207, refute_0_1])).

cnf(refute_0_209, plain,
    (skoXC1 *
     (-8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (25200 + skoXC1 * (-1512 + skoXC1 * (56 + skoXC1 * -1)))))) <=
     -17297280 |
     (-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-3843/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (-53/22 + skoXC1 * 1/11))))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (-25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) *
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     0 <=
     -17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))) |
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_211, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) *
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22) <= 0 |
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_212, plain,
    (~ lgen(0, exp(X_000165), X_000166) | exp(X_000165) <= X_000166),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_213, plain,
    (19/2 < X_000165 | X_000165 < 0 |
     ~ lgen(0,
            -(X_000165 ^ 7 + 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 +
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 +
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 + 17297280) /
            (X_000165 ^ 7 - 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 -
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 -
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 - 17297280),
            X_000166) | lgen(0, exp(X_000165), X_000166)),
    inference(subst, [], [exp_upper_bound_cf7])).

cnf(refute_0_214, plain,
    (19/2 < X_000165 | X_000165 < 0 |
     ~ lgen(0,
            -(X_000165 ^ 7 + 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 +
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 +
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 + 17297280) /
            (X_000165 ^ 7 - 56 * X_000165 ^ 6 + 1512 * X_000165 ^ 5 -
             25200 * X_000165 ^ 4 + 277200 * X_000165 ^ 3 -
             1995840 * X_000165 ^ 2 + 8648640 * X_000165 - 17297280),
            X_000166) | exp(X_000165) <= X_000166),
    inference(resolve, [], [refute_0_213, refute_0_212])).

cnf(refute_0_215, plain,
    (19/2 < X_000165 | X_000165 < 0 |
     X_000166 <
     (-17297280 +
      X_000165 *
      (-8648640 +
       X_000165 *
       (-1995840 +
        X_000165 *
        (-277200 +
         X_000165 *
         (-25200 +
          X_000165 * (-1512 + X_000165 * (-56 + X_000165 * -1))))))) /
     (-17297280 +
      X_000165 *
      (8648640 +
       X_000165 *
       (-1995840 +
        X_000165 *
        (277200 +
         X_000165 *
         (-25200 + X_000165 * (1512 + X_000165 * (-56 + X_000165))))))) |
     exp(X_000165) <= X_000166),
    inference(arithmetic, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     19/2 < skoXC1 |
     exp(skoXC1) <=
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22))),
    inference(subst, [], [refute_0_215])).

cnf(refute_0_217, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     19/2 < skoXC1), inference(resolve, [], [refute_0_216, refute_0_169])).

cnf(refute_0_218, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     19/2 < skoXC1), inference(resolve, [], [refute_0_204, refute_0_217])).

cnf(refute_0_219, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) /
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_206, refute_0_218])).

cnf(refute_0_220, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-17297280 +
      skoXC1 *
      (-8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (-277200 +
         skoXC1 *
         (-25200 + skoXC1 * (-1512 + skoXC1 * (-56 + skoXC1 * -1))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) *
     (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22)) |
     3/2 + skoXC1 * (-1/2 + skoXC1 * 1/22) <= 0),
    inference(resolve, [], [refute_0_211, refute_0_219])).

cnf(refute_0_221, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     skoXC1 * (-1/2 + skoXC1 * 1/22) <= -3/2),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) <
     (-25945920 +
      skoXC1 *
      (-4324320 +
       skoXC1 *
       (544320 +
        skoXC1 *
        (189000 +
         skoXC1 *
         (10080 +
          skoXC1 *
          (-2268 +
           skoXC1 *
           (-5208/11 +
            skoXC1 *
            (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22))))))))) /
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1)))))))),
    inference(resolve, [], [refute_0_221, refute_0_168])).

cnf(refute_0_223, plain,
    (-25945920 +
     skoXC1 *
     (-4324320 +
      skoXC1 *
      (544320 +
       skoXC1 *
       (189000 +
        skoXC1 *
        (10080 +
         skoXC1 *
         (-2268 +
          skoXC1 *
          (-5208/11 +
           skoXC1 * (-929/22 + skoXC1 * (-45/22 + skoXC1 * -1/22)))))))) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * 1/11)) *
     (-17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (-1995840 +
        skoXC1 *
        (277200 +
         skoXC1 *
         (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))) |
     0 <=
     -17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (-1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (-25200 + skoXC1 * (1512 + skoXC1 * (-56 + skoXC1))))))),
    inference(resolve, [], [refute_0_210, refute_0_222])).

cnf(refute_0_224, plain,
    (-25945920 <
     skoXC1 *
     (-30270240 +
      skoXC1 *
      (16752960 +
       skoXC1 *
       (-5753160 +
        skoXC1 *
        (1330560 +
         skoXC1 *
         (-229572 +
          skoXC1 *
          (315672/11 +
           skoXC1 *
           (-51935/22 +
            skoXC1 * (283/2 + skoXC1 * (-111/22 + skoXC1 * 1/11))))))))) |
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (25200 + skoXC1 * (-1512 + skoXC1 * (56 + skoXC1 * -1)))))) <=
     -17297280), inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (-25945920 <
     skoXC1 *
     (-30270240 +
      skoXC1 *
      (16752960 +
       skoXC1 *
       (-5753160 +
        skoXC1 *
        (1330560 +
         skoXC1 *
         (-229572 +
          skoXC1 *
          (315672/11 +
           skoXC1 *
           (-51935/22 +
            skoXC1 * (283/2 + skoXC1 * (-111/22 + skoXC1 * 1/11))))))))) |
     -17297280 <
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (25200 + skoXC1 * (-1512 + skoXC1 * (56 + skoXC1 * -1))))))),
    inference(decision, [],
              [refute_0_203, refute_0_201, refute_0_184, refute_0_151,
               refute_0_147, refute_0_143, refute_0_21, refute_0_20])).

cnf(refute_0_226, plain,
    (-25945920 <
     skoXC1 *
     (-30270240 +
      skoXC1 *
      (16752960 +
       skoXC1 *
       (-5753160 +
        skoXC1 *
        (1330560 +
         skoXC1 *
         (-229572 +
          skoXC1 *
          (315672/11 +
           skoXC1 *
           (-51935/22 +
            skoXC1 * (283/2 + skoXC1 * (-111/22 + skoXC1 * 1/11)))))))))),
    inference(resolve, [], [refute_0_224, refute_0_225])).

cnf(refute_0_227, plain,
    (-17297280 <
     skoXC1 *
     (-8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (-277200 +
        skoXC1 *
        (25200 + skoXC1 * (-1512 + skoXC1 * (56 + skoXC1 * -1))))))),
    inference(decision, [],
              [refute_0_226, refute_0_203, refute_0_201, refute_0_184,
               refute_0_151, refute_0_147, refute_0_143, refute_0_21,
               refute_0_20])).

cnf(refute_0_228, plain,
    ((-784604620800 +
      skoXC1 *
      (-523069747200 +
       skoXC1 *
       (-136101772800 +
        skoXC1 *
        (-11481523200 +
         skoXC1 *
         (2396822400 +
          skoXC1 *
          (-725760 +
           skoXC1 *
           (843091200/11 +
            skoXC1 *
            (-64350720/11 +
             skoXC1 *
             (-30275280/11 +
              skoXC1 *
              (6064800/11 +
               skoXC1 *
               (-459438/11 +
                skoXC1 *
                (19017/11 +
                 skoXC1 *
                 (-3843/22 +
                  skoXC1 *
                  (623/22 +
                   skoXC1 * (-53/22 + skoXC1 * 1/11))))))))))))))) /
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     17297280 +
     skoXC1 *
     (8648640 +
      skoXC1 *
      (1995840 +
       skoXC1 *
       (277200 +
        skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1))))))),
    inference(resolve, [], [refute_0_209, refute_0_227])).

cnf(refute_0_229, plain,
    (0 <=
     -45360 +
     skoXC1 *
     (-7560 +
      skoXC1 *
      (12600/11 +
       skoXC1 *
       (3990/11 +
        skoXC1 *
        (135/11 +
         skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22)))))) |
     (17297280 +
      skoXC1 *
      (8648640 +
       skoXC1 *
       (1995840 +
        skoXC1 *
        (277200 +
         skoXC1 * (25200 + skoXC1 * (1512 + skoXC1 * (56 + skoXC1))))))) *
     (-45360 +
      skoXC1 *
      (-7560 +
       skoXC1 *
       (12600/11 +
        skoXC1 *
        (3990/11 +
         skoXC1 *
         (135/11 +
          skoXC1 * (-123/22 + skoXC1 * (-19/22 + skoXC1 * -1/22))))))) <=
     -784604620800 +
     skoXC1 *
     (-523069747200 +
      skoXC1 *
      (-136101772800 +
       skoXC1 *
       (-11481523200 +
        skoXC1 *
        (2396822400 +
         skoXC1 *
         (-725760 +
          skoXC1 *
          (843091200/11 +
           skoXC1 *
           (-64350720/11 +
            skoXC1 *
            (-30275280/11 +
             skoXC1 *
             (6064800/11 +
              skoXC1 *
              (-459438/11 +
               skoXC1 *
               (19017/11 +
                skoXC1 *
                (-3843/22 +
                 skoXC1 *
                 (623/22 + skoXC1 * (-53/22 + skoXC1 * 1/11))))))))))))))),
    inference(resolve, [], [refute_0_228, refute_0_0])).

cnf(refute_0_230, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (792529920 +
          skoXC1 *
          (skoXC1 *
           (241920 +
            skoXC1 *
            (skoXC1 *
             (-10805760/11 +
              skoXC1 *
              (skoXC1 *
               (-49290/11 +
                skoXC1 *
                (52 +
                 skoXC1 *
                 (-349/11 +
                  skoXC1 * (26/11 + skoXC1 * -1/11)))))))))))))) <= 0 |
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360), inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (792529920 +
          skoXC1 *
          (skoXC1 *
           (241920 +
            skoXC1 *
            (skoXC1 *
             (-10805760/11 +
              skoXC1 *
              (skoXC1 *
               (-49290/11 +
                skoXC1 *
                (52 +
                 skoXC1 *
                 (-349/11 +
                  skoXC1 * (26/11 + skoXC1 * -1/11)))))))))))))) |
     skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360),
    inference(decision, [],
              [refute_0_226, refute_0_203, refute_0_201, refute_0_184,
               refute_0_151, refute_0_147, refute_0_143, refute_0_21,
               refute_0_20])).

cnf(refute_0_232, plain,
    (skoXC1 *
     (7560 +
      skoXC1 *
      (-12600/11 +
       skoXC1 *
       (-3990/11 +
        skoXC1 *
        (-135/11 +
         skoXC1 * (123/22 + skoXC1 * (19/22 + skoXC1 * 1/22)))))) <=
     -45360), inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain, ($false),
    inference(resolve, [], [refute_0_232, refute_0_46])).
% SZS output end CNFRefutation for exp-hyperbolic-1.tptp
