% SZS output start CNFRefutation for HEATING.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

fof('HEATING', conjecture,
    (! [X] :
       (0 <= X =>
        14 <
        92/5 - 333/100 * exp(-9/10 * X) + 62/25 * exp(-73/200 * X) +
        62/25 * exp(-183/100 * X) + 280 * exp(147/200 * X) -
        280 * exp(-147/200 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        14 <
        92/5 - 333/100 * exp(-9/10 * X) + 62/25 * exp(-73/200 * X) +
        62/25 * exp(-183/100 * X) + 280 * exp(147/200 * X) -
        280 * exp(-147/200 * X))), inference(strip, [], ['HEATING'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          14 <
          92/5 - 333/100 * exp(-9/10 * X) + 62/25 * exp(-73/200 * X) +
          62/25 * exp(-183/100 * X) + 280 * exp(147/200 * X) -
          280 * exp(-147/200 * X))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        92/5 - 333/100 * exp(-9/10 * X) + 62/25 * exp(-73/200 * X) +
        62/25 * exp(-183/100 * X) + 280 * exp(147/200 * X) -
        280 * exp(-147/200 * X) <= 14)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     92/5 - 333/100 * exp(-9/10 * skoXC1) + 62/25 * exp(-73/200 * skoXC1) +
     62/25 * exp(-183/100 * skoXC1) + 280 * exp(147/200 * skoXC1) -
     280 * exp(-147/200 * skoXC1) <= 14),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (92/5 - 333/100 * exp(-9/10 * skoXC1) + 62/25 * exp(-73/200 * skoXC1) +
     62/25 * exp(-183/100 * skoXC1) + 280 * exp(147/200 * skoXC1) -
     280 * exp(-147/200 * skoXC1) <= 14),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (17750016/875 +
     skoXC1 *
     (172667199744/21875 +
      skoXC1 *
      (6244813536048/546875 +
       skoXC1 *
       (116123348431494/13671875 +
        skoXC1 *
        (26833592366204409/6250000000 +
         skoXC1 *
         (14145144714771254523/8750000000000 +
          skoXC1 *
          (3320510166399609903627/7000000000000000 +
           skoXC1 *
           (433495521529517853/3906250000000 +
            skoXC1 *
            (104078199258611929953/5000000000000000 +
             skoXC1 *
             (311563955546600917293/100000000000000000 +
              skoXC1 *
              (5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 788272543602726609483/
                 12800000000000000000000)))))))))))) <
     0 *
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     0 <=
     -41140224/875 +
     skoXC1 *
     (-1681606656/21875 +
      skoXC1 *
      (-34367836032/546875 +
       skoXC1 *
       (-468261765936/13671875 +
        skoXC1 *
        (-10755901913877/781250000 +
         skoXC1 *
         (-4678198438943979/1093750000000 +
          skoXC1 *
          (-920555450522252271/875000000000000 +
           skoXC1 *
           (-258012459325890297/1250000000000000 +
            skoXC1 *
            (-803763671664639483/25000000000000000 +
             skoXC1 *
             (-1951834656139927083/500000000000000000 +
              skoXC1 *
              (-28312175689314893397/80000000000000000000 +
               skoXC1 *
               (-17608886012003022639/800000000000000000000 +
                skoXC1 *
                -23747763704260957137/32000000000000000000000))))))))))) |
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) <=
     exp(skoXC1 * -183/100)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) *
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) <
     55240704/875 +
     skoXC1 *
     (145006848/3125 +
      skoXC1 *
      (1332250416/78125 +
       skoXC1 *
       (8160033798/1953125 +
        skoXC1 *
        (4198337389071/6250000000 +
         skoXC1 *
         (88165085170491/1250000000000 +
          skoXC1 * 4320089173354059/1000000000000000))))) |
     2304 +
     skoXC1 *
     (10368/5 +
      skoXC1 *
      (23328/25 +
       skoXC1 *
       (34992/125 +
        skoXC1 *
        (137781/2500 +
         skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))) <= 0 |
     (55240704/875 +
      skoXC1 *
      (145006848/3125 +
       skoXC1 *
       (1332250416/78125 +
        skoXC1 *
        (8160033798/1953125 +
         skoXC1 *
         (4198337389071/6250000000 +
          skoXC1 *
          (88165085170491/1250000000000 +
           skoXC1 * 4320089173354059/1000000000000000)))))) /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) <=
     6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_6, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) *
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) |
     23976/875 +
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) <=
     (6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_8, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_8, refute_0_7])).

cnf(refute_0_10, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (0 < skoXC1 * -9/10 |
     (6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 * -9/10 *
      (-2304 +
       skoXC1 * -9/10 *
       (1152 +
        skoXC1 * -9/10 *
        (-384 +
         skoXC1 * -9/10 *
         (84 + skoXC1 * -9/10 * (-12 + skoXC1 * -9/10)))))) |
     exp(skoXC1 * -9/10) <=
     (6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000))))))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_12, plain,
    ((711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
      exp(skoXC1 * -9/10) * -333/28000) *
     (2304 +
      skoXC1 *
      (42336/25 +
       skoXC1 *
       (388962/625 +
        skoXC1 *
        (9529569/62500 +
         skoXC1 *
         (9805926501/400000000 +
          skoXC1 *
          (205924456521/80000000000 +
           skoXC1 * 10090298369529/64000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (42336/25 +
      skoXC1 *
      (388962/625 +
       skoXC1 *
       (9529569/62500 +
        skoXC1 *
        (9805926501/400000000 +
         skoXC1 *
         (205924456521/80000000000 +
          skoXC1 * 10090298369529/64000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (42336/25 +
       skoXC1 *
       (388962/625 +
        skoXC1 *
        (9529569/62500 +
         skoXC1 *
         (9805926501/400000000 +
          skoXC1 *
          (205924456521/80000000000 +
           skoXC1 * 10090298369529/64000000000000)))))) <=
     711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_13, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_15, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (0 <
     -3555/31 + skoXC1 * -5145/62 + exp(skoXC1 * -183/100) * -1 +
     exp(skoXC1 * -9/10) * 333/248 + exp(skoXC1 * -147/200) * 3500/31 |
     -3555/31 + skoXC1 * -5145/62 + exp(skoXC1 * -183/100) * -1 +
     exp(skoXC1 * -9/10) * 333/248 + exp(skoXC1 * -147/200) * 3500/31 <
     exp(skoXC1 * -73/200)), inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (92/5 - 333/100 * exp(-9/10 * skoXC1) + 62/25 * exp(-73/200 * skoXC1) +
     62/25 * exp(-183/100 * skoXC1) + 280 * exp(147/200 * skoXC1) -
     280 * exp(-147/200 * skoXC1) <= 14),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * 147/200) <=
     -11/700 + exp(skoXC1 * -183/100) * -31/3500 +
     exp(skoXC1 * -9/10) * 333/28000 + exp(skoXC1 * -147/200) +
     exp(skoXC1 * -73/200) * -31/3500),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (X_000048 < exp(X_000047) | ~ lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_20, plain,
    (X_000047 < -1 | ~ lgen(1, X_000048, 1 + X_000047) |
     lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_21, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) |
     ~ lgen(1, X_000048, 1 + X_000047)),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) | 1 + X_000047 <= X_000048),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 * 147/200 < -1 |
     -11/700 + exp(skoXC1 * -183/100) * -31/3500 +
     exp(skoXC1 * -9/10) * 333/28000 + exp(skoXC1 * -147/200) +
     exp(skoXC1 * -73/200) * -31/3500 < exp(skoXC1 * 147/200) |
     1 + skoXC1 * 147/200 <=
     -11/700 + exp(skoXC1 * -183/100) * -31/3500 +
     exp(skoXC1 * -9/10) * 333/28000 + exp(skoXC1 * -147/200) +
     exp(skoXC1 * -73/200) * -31/3500),
    inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * 147/200 < -1 |
     1 + skoXC1 * 147/200 <=
     -11/700 + exp(skoXC1 * -183/100) * -31/3500 +
     exp(skoXC1 * -9/10) * 333/28000 + exp(skoXC1 * -147/200) +
     exp(skoXC1 * -73/200) * -31/3500),
    inference(resolve, [], [refute_0_18, refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 < -200/147 |
     exp(skoXC1 * -73/200) <=
     -3555/31 + skoXC1 * -5145/62 + exp(skoXC1 * -183/100) * -1 +
     exp(skoXC1 * -9/10) * 333/248 + exp(skoXC1 * -147/200) * 3500/31),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (-200/147 <= skoXC1),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (exp(skoXC1 * -73/200) <=
     -3555/31 + skoXC1 * -5145/62 + exp(skoXC1 * -183/100) * -1 +
     exp(skoXC1 * -9/10) * 333/248 + exp(skoXC1 * -147/200) * 3500/31),
    inference(resolve, [], [refute_0_27, refute_0_25])).

cnf(refute_0_29, plain,
    (0 <
     -3555/31 + skoXC1 * -5145/62 + exp(skoXC1 * -183/100) * -1 +
     exp(skoXC1 * -9/10) * 333/248 + exp(skoXC1 * -147/200) * 3500/31),
    inference(resolve, [], [refute_0_28, refute_0_16])).

cnf(refute_0_30, plain,
    (711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000 < exp(skoXC1 * -147/200)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (0 < skoXC1 * -147/200 |
     711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000 <
     2304 /
     (2304 +
      skoXC1 * -147/200 *
      (-2304 +
       skoXC1 * -147/200 *
       (1152 +
        skoXC1 * -147/200 *
        (-384 +
         skoXC1 * -147/200 *
         (84 + skoXC1 * -147/200 * (-12 + skoXC1 * -147/200)))))) |
     exp(skoXC1 * -147/200) <=
     711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_32, plain,
    (0 < skoXC1 * -147/200 |
     711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000 <
     2304 /
     (2304 +
      skoXC1 * -147/200 *
      (-2304 +
       skoXC1 * -147/200 *
       (1152 +
        skoXC1 * -147/200 *
        (-384 +
         skoXC1 * -147/200 *
         (84 + skoXC1 * -147/200 * (-12 + skoXC1 * -147/200))))))),
    inference(resolve, [], [refute_0_31, refute_0_30])).

cnf(refute_0_33, plain,
    (skoXC1 < 0 |
     711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000 <
     2304 /
     (2304 +
      skoXC1 *
      (42336/25 +
       skoXC1 *
       (388962/625 +
        skoXC1 *
        (9529569/62500 +
         skoXC1 *
         (9805926501/400000000 +
          skoXC1 *
          (205924456521/80000000000 +
           skoXC1 * 10090298369529/64000000000000))))))),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
     exp(skoXC1 * -9/10) * -333/28000 <
     2304 /
     (2304 +
      skoXC1 *
      (42336/25 +
       skoXC1 *
       (388962/625 +
        skoXC1 *
        (9529569/62500 +
         skoXC1 *
         (9805926501/400000000 +
          skoXC1 *
          (205924456521/80000000000 +
           skoXC1 * 10090298369529/64000000000000))))))),
    inference(resolve, [], [refute_0_26, refute_0_33])).

cnf(refute_0_35, plain,
    ((711/700 + skoXC1 * 147/200 + exp(skoXC1 * -183/100) * 31/3500 +
      exp(skoXC1 * -9/10) * -333/28000) *
     (2304 +
      skoXC1 *
      (42336/25 +
       skoXC1 *
       (388962/625 +
        skoXC1 *
        (9529569/62500 +
         skoXC1 *
         (9805926501/400000000 +
          skoXC1 *
          (205924456521/80000000000 +
           skoXC1 * 10090298369529/64000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (42336/25 +
      skoXC1 *
      (388962/625 +
       skoXC1 *
       (9529569/62500 +
        skoXC1 *
        (9805926501/400000000 +
         skoXC1 *
         (205924456521/80000000000 +
          skoXC1 * 10090298369529/64000000000000))))) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_34])).

cnf(refute_0_36, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     exp(skoXC1 * -9/10) *
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) |
     skoXC1 *
     (42336/25 +
      skoXC1 *
      (388962/625 +
       skoXC1 *
       (9529569/62500 +
        skoXC1 *
        (9805926501/400000000 +
         skoXC1 *
         (205924456521/80000000000 +
          skoXC1 * 10090298369529/64000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (-2304 <
     skoXC1 *
     (42336/25 +
      skoXC1 *
      (388962/625 +
       skoXC1 *
       (9529569/62500 +
        skoXC1 *
        (9805926501/400000000 +
         skoXC1 *
         (205924456521/80000000000 +
          skoXC1 * 10090298369529/64000000000000)))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_38, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     exp(skoXC1 * -9/10) *
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000))))))),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     exp(skoXC1 * -9/10) |
     23976/875 +
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000))))) <= 0),
    inference(split, [], [refute_0_38])).

cnf(refute_0_40, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     exp(skoXC1 * -9/10) |
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000))))) <= -23976/875),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (-23976/875 <
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000)))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_42, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     exp(skoXC1 * -9/10)),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (0 < skoXC1 * -9/10 |
     (6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 * -9/10 *
      (-2304 +
       skoXC1 * -9/10 *
       (1152 +
        skoXC1 * -9/10 *
        (-384 +
         skoXC1 * -9/10 *
         (84 + skoXC1 * -9/10 * (-12 + skoXC1 * -9/10))))))),
    inference(resolve, [], [refute_0_11, refute_0_42])).

cnf(refute_0_44, plain,
    (skoXC1 < 0 |
     (6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))))),
    inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) /
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))))),
    inference(resolve, [], [refute_0_26, refute_0_44])).

cnf(refute_0_46, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     2304 /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) *
     (23976/875 +
      skoXC1 *
      (62937/3125 +
       skoXC1 *
       (9251739/1250000 +
        skoXC1 *
        (453335211/250000000 +
         skoXC1 *
         (466481932119/1600000000000 +
          skoXC1 *
          (9796120574499/320000000000000 +
           skoXC1 * 480009908150451/256000000000000000)))))) |
     23976/875 +
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_45])).

cnf(refute_0_47, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     (55240704/875 +
      skoXC1 *
      (145006848/3125 +
       skoXC1 *
       (1332250416/78125 +
        skoXC1 *
        (8160033798/1953125 +
         skoXC1 *
         (4198337389071/6250000000 +
          skoXC1 *
          (88165085170491/1250000000000 +
           skoXC1 * 4320089173354059/1000000000000000)))))) /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) |
     skoXC1 *
     (62937/3125 +
      skoXC1 *
      (9251739/1250000 +
       skoXC1 *
       (453335211/250000000 +
        skoXC1 *
        (466481932119/1600000000000 +
         skoXC1 *
         (9796120574499/320000000000000 +
          skoXC1 * 480009908150451/256000000000000000))))) <= -23976/875),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (6336/175 +
     skoXC1 *
     (2133432/625 +
      skoXC1 *
      (58649913/31250 +
       skoXC1 *
       (3826802637/6250000 +
        skoXC1 *
        (5478711220773/40000000000 +
         skoXC1 *
         (165063160791333/8000000000000 +
          skoXC1 *
          (13133244063539817/6400000000000000 +
           skoXC1 * 1483273860320763/12800000000000000)))))) +
     exp(skoXC1 * -183/100) *
     (17856/875 +
      skoXC1 *
      (46872/3125 +
       skoXC1 *
       (861273/156250 +
        skoXC1 *
        (42202377/31250000 +
         skoXC1 *
         (43426245933/200000000000 +
          skoXC1 *
          (911951164593/40000000000000 +
           skoXC1 * 44685607065057/32000000000000000)))))) <
     (55240704/875 +
      skoXC1 *
      (145006848/3125 +
       skoXC1 *
       (1332250416/78125 +
        skoXC1 *
        (8160033798/1953125 +
         skoXC1 *
         (4198337389071/6250000000 +
          skoXC1 *
          (88165085170491/1250000000000 +
           skoXC1 * 4320089173354059/1000000000000000)))))) /
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))))),
    inference(resolve, [], [refute_0_47, refute_0_41])).

cnf(refute_0_49, plain,
    ((6336/175 +
      skoXC1 *
      (2133432/625 +
       skoXC1 *
       (58649913/31250 +
        skoXC1 *
        (3826802637/6250000 +
         skoXC1 *
         (5478711220773/40000000000 +
          skoXC1 *
          (165063160791333/8000000000000 +
           skoXC1 *
           (13133244063539817/6400000000000000 +
            skoXC1 * 1483273860320763/12800000000000000)))))) +
      exp(skoXC1 * -183/100) *
      (17856/875 +
       skoXC1 *
       (46872/3125 +
        skoXC1 *
        (861273/156250 +
         skoXC1 *
         (42202377/31250000 +
          skoXC1 *
          (43426245933/200000000000 +
           skoXC1 *
           (911951164593/40000000000000 +
            skoXC1 * 44685607065057/32000000000000000))))))) *
     (2304 +
      skoXC1 *
      (10368/5 +
       skoXC1 *
       (23328/25 +
        skoXC1 *
        (34992/125 +
         skoXC1 *
         (137781/2500 +
          skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))) <
     55240704/875 +
     skoXC1 *
     (145006848/3125 +
      skoXC1 *
      (1332250416/78125 +
       skoXC1 *
       (8160033798/1953125 +
        skoXC1 *
        (4198337389071/6250000000 +
         skoXC1 *
         (88165085170491/1250000000000 +
          skoXC1 * 4320089173354059/1000000000000000))))) |
     2304 +
     skoXC1 *
     (10368/5 +
      skoXC1 *
      (23328/25 +
       skoXC1 *
       (34992/125 +
        skoXC1 *
        (137781/2500 +
         skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))) <= 0),
    inference(resolve, [], [refute_0_5, refute_0_48])).

cnf(refute_0_50, plain,
    (17750016/875 +
     skoXC1 *
     (172667199744/21875 +
      skoXC1 *
      (6244813536048/546875 +
       skoXC1 *
       (116123348431494/13671875 +
        skoXC1 *
        (26833592366204409/6250000000 +
         skoXC1 *
         (14145144714771254523/8750000000000 +
          skoXC1 *
          (3320510166399609903627/7000000000000000 +
           skoXC1 *
           (433495521529517853/3906250000000 +
            skoXC1 *
            (104078199258611929953/5000000000000000 +
             skoXC1 *
             (311563955546600917293/100000000000000000 +
              skoXC1 *
              (5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 788272543602726609483/
                 12800000000000000000000)))))))))))) <
     exp(skoXC1 * -183/100) *
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     skoXC1 *
     (10368/5 +
      skoXC1 *
      (23328/25 +
       skoXC1 *
       (34992/125 +
        skoXC1 *
        (137781/2500 +
         skoXC1 * (177147/25000 + skoXC1 * 531441/1000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (-2304 <
     skoXC1 *
     (10368/5 +
      skoXC1 *
      (23328/25 +
       skoXC1 *
       (34992/125 +
        skoXC1 *
        (137781/2500 +
         skoXC1 * (177147/25000 + skoXC1 * 531441/1000000)))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_52, plain,
    (17750016/875 +
     skoXC1 *
     (172667199744/21875 +
      skoXC1 *
      (6244813536048/546875 +
       skoXC1 *
       (116123348431494/13671875 +
        skoXC1 *
        (26833592366204409/6250000000 +
         skoXC1 *
         (14145144714771254523/8750000000000 +
          skoXC1 *
          (3320510166399609903627/7000000000000000 +
           skoXC1 *
           (433495521529517853/3906250000000 +
            skoXC1 *
            (104078199258611929953/5000000000000000 +
             skoXC1 *
             (311563955546600917293/100000000000000000 +
              skoXC1 *
              (5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 788272543602726609483/
                 12800000000000000000000)))))))))))) <
     exp(skoXC1 * -183/100) *
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    (exp(skoXC1 * -183/100) <
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     0 <=
     -41140224/875 +
     skoXC1 *
     (-1681606656/21875 +
      skoXC1 *
      (-34367836032/546875 +
       skoXC1 *
       (-468261765936/13671875 +
        skoXC1 *
        (-10755901913877/781250000 +
         skoXC1 *
         (-4678198438943979/1093750000000 +
          skoXC1 *
          (-920555450522252271/875000000000000 +
           skoXC1 *
           (-258012459325890297/1250000000000000 +
            skoXC1 *
            (-803763671664639483/25000000000000000 +
             skoXC1 *
             (-1951834656139927083/500000000000000000 +
              skoXC1 *
              (-28312175689314893397/80000000000000000000 +
               skoXC1 *
               (-17608886012003022639/800000000000000000000 +
                skoXC1 *
                -23747763704260957137/32000000000000000000000)))))))))))),
    inference(split, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (exp(skoXC1 * -183/100) <
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     skoXC1 *
     (1681606656/21875 +
      skoXC1 *
      (34367836032/546875 +
       skoXC1 *
       (468261765936/13671875 +
        skoXC1 *
        (10755901913877/781250000 +
         skoXC1 *
         (4678198438943979/1093750000000 +
          skoXC1 *
          (920555450522252271/875000000000000 +
           skoXC1 *
           (258012459325890297/1250000000000000 +
            skoXC1 *
            (803763671664639483/25000000000000000 +
             skoXC1 *
             (1951834656139927083/500000000000000000 +
              skoXC1 *
              (28312175689314893397/80000000000000000000 +
               skoXC1 *
               (17608886012003022639/800000000000000000000 +
                skoXC1 *
                23747763704260957137/32000000000000000000000))))))))))) <=
     -41140224/875), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-41140224/875 <
     skoXC1 *
     (1681606656/21875 +
      skoXC1 *
      (34367836032/546875 +
       skoXC1 *
       (468261765936/13671875 +
        skoXC1 *
        (10755901913877/781250000 +
         skoXC1 *
         (4678198438943979/1093750000000 +
          skoXC1 *
          (920555450522252271/875000000000000 +
           skoXC1 *
           (258012459325890297/1250000000000000 +
            skoXC1 *
            (803763671664639483/25000000000000000 +
             skoXC1 *
             (1951834656139927083/500000000000000000 +
              skoXC1 *
              (28312175689314893397/80000000000000000000 +
               skoXC1 *
               (17608886012003022639/800000000000000000000 +
                skoXC1 *
                23747763704260957137/32000000000000000000000)))))))))))),
    inference(decision, [], [refute_0_26])).

cnf(refute_0_56, plain,
    (exp(skoXC1 * -183/100) <
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (0 <
     (17750016/875 +
      skoXC1 *
      (172667199744/21875 +
       skoXC1 *
       (6244813536048/546875 +
        skoXC1 *
        (116123348431494/13671875 +
         skoXC1 *
         (26833592366204409/6250000000 +
          skoXC1 *
          (14145144714771254523/8750000000000 +
           skoXC1 *
           (3320510166399609903627/7000000000000000 +
            skoXC1 *
            (433495521529517853/3906250000000 +
             skoXC1 *
             (104078199258611929953/5000000000000000 +
              skoXC1 *
              (311563955546600917293/100000000000000000 +
               skoXC1 *
               (5832418897049731828317/16000000000000000000 +
                skoXC1 *
                (2551388094775149606837/80000000000000000000 +
                 skoXC1 *
                 (12234694649056507949517/6400000000000000000000 +
                  skoXC1 *
                  788272543602726609483/
                  12800000000000000000000))))))))))))) /
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_4, refute_0_56])).

cnf(refute_0_58, plain,
    (17750016/875 +
     skoXC1 *
     (172667199744/21875 +
      skoXC1 *
      (6244813536048/546875 +
       skoXC1 *
       (116123348431494/13671875 +
        skoXC1 *
        (26833592366204409/6250000000 +
         skoXC1 *
         (14145144714771254523/8750000000000 +
          skoXC1 *
          (3320510166399609903627/7000000000000000 +
           skoXC1 *
           (433495521529517853/3906250000000 +
            skoXC1 *
            (104078199258611929953/5000000000000000 +
             skoXC1 *
             (311563955546600917293/100000000000000000 +
              skoXC1 *
              (5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 788272543602726609483/
                 12800000000000000000000)))))))))))) <
     0 *
     (-41140224/875 +
      skoXC1 *
      (-1681606656/21875 +
       skoXC1 *
       (-34367836032/546875 +
        skoXC1 *
        (-468261765936/13671875 +
         skoXC1 *
         (-10755901913877/781250000 +
          skoXC1 *
          (-4678198438943979/1093750000000 +
           skoXC1 *
           (-920555450522252271/875000000000000 +
            skoXC1 *
            (-258012459325890297/1250000000000000 +
             skoXC1 *
             (-803763671664639483/25000000000000000 +
              skoXC1 *
              (-1951834656139927083/500000000000000000 +
               skoXC1 *
               (-28312175689314893397/80000000000000000000 +
                skoXC1 *
                (-17608886012003022639/800000000000000000000 +
                 skoXC1 *
                 -23747763704260957137/
                 32000000000000000000000)))))))))))) |
     0 <=
     -41140224/875 +
     skoXC1 *
     (-1681606656/21875 +
      skoXC1 *
      (-34367836032/546875 +
       skoXC1 *
       (-468261765936/13671875 +
        skoXC1 *
        (-10755901913877/781250000 +
         skoXC1 *
         (-4678198438943979/1093750000000 +
          skoXC1 *
          (-920555450522252271/875000000000000 +
           skoXC1 *
           (-258012459325890297/1250000000000000 +
            skoXC1 *
            (-803763671664639483/25000000000000000 +
             skoXC1 *
             (-1951834656139927083/500000000000000000 +
              skoXC1 *
              (-28312175689314893397/80000000000000000000 +
               skoXC1 *
               (-17608886012003022639/800000000000000000000 +
                skoXC1 *
                -23747763704260957137/32000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_0, refute_0_57])).

cnf(refute_0_59, plain,
    (17750016/875 <
     skoXC1 *
     (-172667199744/21875 +
      skoXC1 *
      (-6244813536048/546875 +
       skoXC1 *
       (-116123348431494/13671875 +
        skoXC1 *
        (-26833592366204409/6250000000 +
         skoXC1 *
         (-14145144714771254523/8750000000000 +
          skoXC1 *
          (-3320510166399609903627/7000000000000000 +
           skoXC1 *
           (-433495521529517853/3906250000000 +
            skoXC1 *
            (-104078199258611929953/5000000000000000 +
             skoXC1 *
             (-311563955546600917293/100000000000000000 +
              skoXC1 *
              (-5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (-2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (-12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 -788272543602726609483/
                 12800000000000000000000)))))))))))) |
     skoXC1 *
     (1681606656/21875 +
      skoXC1 *
      (34367836032/546875 +
       skoXC1 *
       (468261765936/13671875 +
        skoXC1 *
        (10755901913877/781250000 +
         skoXC1 *
         (4678198438943979/1093750000000 +
          skoXC1 *
          (920555450522252271/875000000000000 +
           skoXC1 *
           (258012459325890297/1250000000000000 +
            skoXC1 *
            (803763671664639483/25000000000000000 +
             skoXC1 *
             (1951834656139927083/500000000000000000 +
              skoXC1 *
              (28312175689314893397/80000000000000000000 +
               skoXC1 *
               (17608886012003022639/800000000000000000000 +
                skoXC1 *
                23747763704260957137/32000000000000000000000))))))))))) <=
     -41140224/875), inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 *
     (-172667199744/21875 +
      skoXC1 *
      (-6244813536048/546875 +
       skoXC1 *
       (-116123348431494/13671875 +
        skoXC1 *
        (-26833592366204409/6250000000 +
         skoXC1 *
         (-14145144714771254523/8750000000000 +
          skoXC1 *
          (-3320510166399609903627/7000000000000000 +
           skoXC1 *
           (-433495521529517853/3906250000000 +
            skoXC1 *
            (-104078199258611929953/5000000000000000 +
             skoXC1 *
             (-311563955546600917293/100000000000000000 +
              skoXC1 *
              (-5832418897049731828317/16000000000000000000 +
               skoXC1 *
               (-2551388094775149606837/80000000000000000000 +
                skoXC1 *
                (-12234694649056507949517/6400000000000000000000 +
                 skoXC1 *
                 -788272543602726609483/
                 12800000000000000000000)))))))))))) <= 17750016/875 |
     skoXC1 *
     (1681606656/21875 +
      skoXC1 *
      (34367836032/546875 +
       skoXC1 *
       (468261765936/13671875 +
        skoXC1 *
        (10755901913877/781250000 +
         skoXC1 *
         (4678198438943979/1093750000000 +
          skoXC1 *
          (920555450522252271/875000000000000 +
           skoXC1 *
           (258012459325890297/1250000000000000 +
            skoXC1 *
            (803763671664639483/25000000000000000 +
             skoXC1 *
             (1951834656139927083/500000000000000000 +
              skoXC1 *
              (28312175689314893397/80000000000000000000 +
               skoXC1 *
               (17608886012003022639/800000000000000000000 +
                skoXC1 *
                23747763704260957137/32000000000000000000000))))))))))) <=
     -41140224/875), inference(decision, [], [refute_0_26])).

cnf(refute_0_61, plain,
    (skoXC1 *
     (1681606656/21875 +
      skoXC1 *
      (34367836032/546875 +
       skoXC1 *
       (468261765936/13671875 +
        skoXC1 *
        (10755901913877/781250000 +
         skoXC1 *
         (4678198438943979/1093750000000 +
          skoXC1 *
          (920555450522252271/875000000000000 +
           skoXC1 *
           (258012459325890297/1250000000000000 +
            skoXC1 *
            (803763671664639483/25000000000000000 +
             skoXC1 *
             (1951834656139927083/500000000000000000 +
              skoXC1 *
              (28312175689314893397/80000000000000000000 +
               skoXC1 *
               (17608886012003022639/800000000000000000000 +
                skoXC1 *
                23747763704260957137/32000000000000000000000))))))))))) <=
     -41140224/875), inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain, ($false),
    inference(resolve, [], [refute_0_61, refute_0_55])).
% SZS output end CNFRefutation for HEATING.tptp
