% SZS output start CNFRefutation for KeYmaera_damped_pendulum_general.tptp
cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, Y, X) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_const_1, axiom,
    (~ lgen(R, 1, Y) | lgen(R, cos(X), Y))).

cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

fof('KeYmaera', conjecture,
    (! [X2, T, A] :
       ((0 <= T & X2 <= 1 & 0 <= X2) =>
        exp(-1 * T) * X2 * (A * A * cos(A * T) + A * sin(A * T)) <=
        A * A))).

fof(subgoal_0, plain,
    (! [X2, T, A] :
       ((0 <= T & X2 <= 1 & 0 <= X2) =>
        exp(-1 * T) * X2 * (A * A * cos(A * T) + A * sin(A * T)) <=
        A * A)), inference(strip, [], ['KeYmaera'])).

fof(negate_0_0, plain,
    (~ ! [X2, T, A] :
         ((0 <= T & X2 <= 1 & 0 <= X2) =>
          exp(-1 * T) * X2 * (A * A * cos(A * T) + A * sin(A * T)) <=
          A * A)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [T, X2] :
       (X2 <= 1 & 0 <= T & 0 <= X2 &
        ? [A] :
          A * A <
          exp(-1 * T) * X2 * (A * A * cos(A * T) + A * sin(A * T)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoTC1 & 0 <= skoX2C1 & skoX2C1 <= 1 &
     ? [A] :
       A * A <
       exp(-1 * skoTC1) * skoX2C1 *
       (A * A * cos(A * skoTC1) + A * sin(A * skoTC1))),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (? [A] :
       A * A <
       exp(-1 * skoTC1) * skoX2C1 *
       (A * A * cos(A * skoTC1) + A * sin(A * skoTC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain,
    (skoAC1 * skoAC1 <
     exp(-1 * skoTC1) * skoX2C1 *
     (skoAC1 * skoAC1 * cos(skoAC1 * skoTC1) +
      skoAC1 * sin(skoAC1 * skoTC1))),
    inference(skolemize, [], [normalize_0_2])).

fof(normalize_0_4, plain, (0 <= skoTC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoX2C1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoX2C1 <= 1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoAC1 * skoAC1 <
     exp(-1 * skoTC1) * skoX2C1 *
     (skoAC1 * skoAC1 * cos(skoAC1 * skoTC1) +
      skoAC1 * sin(skoAC1 * skoTC1))),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_1, plain,
    (skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) <
     sin(skoTC1 * skoAC1) * (exp(skoTC1 * -1) * (skoX2C1 * skoAC1))),
    inference(arithmetic, [], [refute_0_0])).

cnf(refute_0_2, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) < 0 |
     skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) < 0 |
     0 < exp(skoTC1 * -1) * (skoX2C1 * skoAC1)),
    inference(split, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (skoAC1 * skoAC1 <
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * skoAC1)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 * (skoTC1 * 2304) |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0 |
     skoX2C1 * (skoTC1 * 2304) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 +
        skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <= 1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_5, plain,
    (1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * skoTC1) | skoX2C1 * skoTC1 <= 0 |
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 +
        skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <=
     1 / (skoX2C1 * skoTC1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (~ lgen(0, exp(X_000116), X_000117) | exp(X_000116) <= X_000117),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (0 < X_000116 |
     ~ lgen(0,
            2304 /
            (-(X_000116 ^ 3) + 6 * X_000116 ^ 2 - 24 * X_000116 + 48) ^ 2,
            X_000117) | lgen(0, exp(X_000116), X_000117)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_8, plain,
    (0 < X_000116 |
     ~ lgen(0,
            2304 /
            (-(X_000116 ^ 3) + 6 * X_000116 ^ 2 - 24 * X_000116 + 48) ^ 2,
            X_000117) | exp(X_000116) <= X_000117),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    (0 < X_000116 |
     X_000117 <
     2304 /
     (2304 +
      X_000116 *
      (-2304 +
       X_000116 *
       (1152 +
        X_000116 *
        (-384 + X_000116 * (84 + X_000116 * (-12 + X_000116)))))) |
     exp(X_000116) <= X_000117), inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (0 < skoTC1 * -1 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) <= 1 / (skoX2C1 * skoTC1)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) <
     skoTC1 * skoAC1 * (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 * skoAC1 <=
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain,
    ((skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) < sin(skoTC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_1])).

cnf(refute_0_13, plain,
    (~ lgen(0, sin(X_000075), X_000076) | sin(X_000075) <= X_000076),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_14, plain,
    (X_000075 < 0 | ~ lgen(0, X_000075, X_000076) |
     lgen(0, sin(X_000075), X_000076)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_15, plain,
    (X_000075 < 0 | ~ lgen(0, X_000075, X_000076) |
     sin(X_000075) <= X_000076),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (X_000076 < X_000075 | X_000075 < 0 | sin(X_000075) <= X_000076),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (skoTC1 * skoAC1 < 0 |
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) < skoTC1 * skoAC1 |
     sin(skoTC1 * skoAC1) <=
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1))),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (skoTC1 * skoAC1 < 0 |
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) < skoTC1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_17, refute_0_12])).

cnf(refute_0_19, plain,
    (0 < skoTC1 * (skoAC1 * -1) |
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) < skoTC1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (0 <= skoTC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_21, plain, (0 <= skoX2C1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_22, plain, (skoX2C1 <= 1),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_23, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_24, plain,
    (0 < 0 | skoX2C1 * skoAC1 < 0 | 0 < skoX2C1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (0 < skoX2C1 * (skoAC1 * -1) | 0 < skoX2C1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (exp(skoTC1 * -1) < 0 / (skoX2C1 * skoAC1) | 0 <= skoX2C1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_23])).

cnf(refute_0_27, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoAC1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (0 < skoX2C1 * skoAC1 | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_27, refute_0_25])).

cnf(refute_0_29, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 * (skoAC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_30, plain,
    ((skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) < skoTC1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_29, refute_0_19])).

cnf(refute_0_31, plain,
    (skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) <
     skoTC1 * skoAC1 * (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_11, refute_0_30])).

cnf(refute_0_32, plain,
    (skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) < 0 |
     skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     0 | 0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoAC1 * skoAC1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0 |
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 +
        skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <=
     skoAC1 * skoAC1), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_36, plain,
    (skoAC1 * skoAC1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0 |
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 +
        skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <=
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_37, plain,
    (skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     1 * (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     1 <=
     (skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_38, plain,
    (~ lgen(0, cos(X_000016), X_000017) | cos(X_000016) <= X_000017),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_39, plain,
    (~ lgen(0, 1, X_000017) | lgen(0, cos(X_000016), X_000017)),
    inference(subst, [], [cos_upper_bound_const_1])).

cnf(refute_0_40, plain,
    (~ lgen(0, 1, X_000017) | cos(X_000016) <= X_000017),
    inference(resolve, [], [refute_0_39, refute_0_38])).

cnf(refute_0_41, plain, (X_000017 < 1 | cos(X_000016) <= X_000017),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    ((skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) < 1 |
     cos(skoTC1 * skoAC1) <=
     (skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)))),
    inference(subst, [], [refute_0_41])).

cnf(refute_0_43, plain,
    ((skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) <
     cos(skoTC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_32])).

cnf(refute_0_44, plain,
    ((skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) < 1 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     1 * (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_37, refute_0_44])).

cnf(refute_0_46, plain,
    (skoAC1 * skoAC1 <
     exp(skoTC1 * -1) *
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (0 < skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(decision, [],
              [refute_0_28, refute_0_22, refute_0_21, refute_0_20])).

cnf(refute_0_49, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (0 < skoTC1 * -1 |
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) <=
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)))),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_51, plain,
    (0 < skoTC1 * -1 |
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (skoTC1 < 0 |
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (0 <= skoTC1 | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_54, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_53, refute_0_52])).

cnf(refute_0_55, plain,
    (skoAC1 * skoAC1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_36, refute_0_54])).

cnf(refute_0_56, plain,
    (skoAC1 * skoAC1 <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoAC1 * skoAC1 <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_56, refute_0_48])).

cnf(refute_0_58, plain,
    (skoAC1 * skoAC1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0), inference(resolve, [], [refute_0_35, refute_0_57])).

cnf(refute_0_59, plain,
    (skoAC1 * (skoAC1 * 2304) +
     skoTC1 *
     (skoAC1 * (skoAC1 * 2304) +
      skoTC1 *
      (skoAC1 * (skoAC1 * 1152) +
       skoTC1 *
       (skoAC1 * (skoAC1 * 384) +
        skoTC1 *
        (skoAC1 * (skoAC1 * 84) +
         skoTC1 *
         (skoAC1 * (skoAC1 * 12) + skoTC1 * (skoAC1 * skoAC1)))))) <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304 |
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) <=
     skoAC1 * (skoAC1 * 2304) +
     skoTC1 *
     (skoAC1 * (skoAC1 * 2304) +
      skoTC1 *
      (skoAC1 * (skoAC1 * 1152) +
       skoTC1 *
       (skoAC1 * (skoAC1 * 384) +
        skoTC1 *
        (skoAC1 * (skoAC1 * 84) +
         skoTC1 *
         (skoAC1 * (skoAC1 * 12) + skoTC1 * (skoAC1 * skoAC1))))))),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_61, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (-2304 <
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_63, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_63, refute_0_34])).

cnf(refute_0_65, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_66, plain,
    (0 < 0 | skoX2C1 * (skoAC1 * (skoAC1 * -1)) < 0 |
     0 < skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(split, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (0 < skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     0 < skoX2C1 * (skoAC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (exp(skoTC1 * -1) < 0 / (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 <= skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(split, [], [refute_0_65])).

cnf(refute_0_69, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (0 < skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(resolve, [], [refute_0_69, refute_0_67])).

cnf(refute_0_71, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoAC1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_28, refute_0_22, refute_0_21, refute_0_20])).

cnf(refute_0_72, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_71, refute_0_70])).

cnf(refute_0_73, plain,
    (skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_72, refute_0_64])).

cnf(refute_0_74, plain,
    (skoAC1 * skoAC1 / (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (1 / (skoX2C1 * skoTC1) < exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0 | skoAC1 = 0),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoAC1 != 0 | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_77, plain,
    (1 / (skoX2C1 * skoTC1) < exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (0 < skoTC1 * -1 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_77])).

cnf(refute_0_79, plain,
    (skoTC1 < 0 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (skoAC1 * skoAC1 < 0 | skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) < 0 |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(split, [], [refute_0_73])).

cnf(refute_0_81, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_83, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_82, refute_0_81])).

cnf(refute_0_84, plain,
    (0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoAC1 * (skoAC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_28])).

cnf(refute_0_85, plain,
    (0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_84, refute_0_83])).

cnf(refute_0_86, plain,
    (skoTC1 < 0 | 0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(decision, [],
              [refute_0_85, refute_0_20, refute_0_21, refute_0_22,
               refute_0_28])).

cnf(refute_0_87, plain,
    (skoTC1 < 0 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_79, refute_0_86])).

cnf(refute_0_88, plain,
    (1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_53, refute_0_87])).

cnf(refute_0_89, plain,
    (1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * skoTC1) | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoX2C1 * skoTC1 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_88])).

cnf(refute_0_90, plain,
    (1 <
     skoX2C1 * (skoTC1 * 2304) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 | skoX2C1 * skoTC1 <= 0),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (0 < skoX2C1 * skoTC1 | exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(decision, [],
              [refute_0_85, refute_0_20, refute_0_21, refute_0_22,
               refute_0_28])).

cnf(refute_0_92, plain,
    (1 <
     skoX2C1 * (skoTC1 * 2304) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain,
    (1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 * (skoTC1 * 2304) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0), inference(resolve, [], [refute_0_4, refute_0_92])).

cnf(refute_0_94, plain,
    (2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <
     skoX2C1 * (skoTC1 * 2304) |
     exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304 |
     skoX2C1 * (skoTC1 * 2304) <=
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))),
    inference(decision, [],
              [refute_0_85, refute_0_20, refute_0_21, refute_0_22,
               refute_0_28])).

cnf(refute_0_96, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(resolve, [], [refute_0_95, refute_0_94])).

cnf(refute_0_97, plain, (exp(skoTC1 * -1) * (skoX2C1 * skoAC1) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_62])).

cnf(refute_0_98, plain,
    (skoAC1 * skoAC1 <
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1))),
    inference(resolve, [], [refute_0_97, refute_0_3])).

cnf(refute_0_99, plain,
    (skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) <
     skoTC1 * skoAC1 * (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     0 <= exp(skoTC1 * -1) * (skoX2C1 * skoAC1) |
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) <= skoTC1 * skoAC1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_100, plain,
    (~ lgen(0, X_000083, sin(X_000082)) | X_000083 <= sin(X_000082)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_101, plain,
    (0 < X_000082 | ~ lgen(0, X_000083, X_000082) |
     lgen(0, X_000083, sin(X_000082))),
    inference(subst, [], [sin_lower_bound_0_neg])).

cnf(refute_0_102, plain,
    (0 < X_000082 | ~ lgen(0, X_000083, X_000082) |
     X_000083 <= sin(X_000082)),
    inference(resolve, [], [refute_0_101, refute_0_100])).

cnf(refute_0_103, plain,
    (0 < X_000082 | X_000082 < X_000083 | X_000083 <= sin(X_000082)),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (0 < skoTC1 * skoAC1 |
     skoTC1 * skoAC1 <
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) <= sin(skoTC1 * skoAC1)),
    inference(subst, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (sin(skoTC1 * skoAC1) <
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     0 <= exp(skoTC1 * -1) * (skoX2C1 * skoAC1)),
    inference(split, [], [refute_0_1])).

cnf(refute_0_106, plain,
    (sin(skoTC1 * skoAC1) <
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (0 < skoTC1 * skoAC1 |
     skoTC1 * skoAC1 <
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_104, refute_0_106])).

cnf(refute_0_108, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_109, plain,
    (exp(skoTC1 * -1) < 0 / (skoX2C1 * (skoAC1 * -1)) |
     0 <= skoX2C1 * (skoAC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(split, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * skoAC1 <= 0), inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (0 < 0 | skoX2C1 * (skoAC1 * -1) < 0 | 0 < skoX2C1 * (skoAC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(split, [], [refute_0_108])).

cnf(refute_0_112, plain,
    (0 < skoX2C1 * (skoAC1 * -1) | 0 < skoX2C1 * skoAC1 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (0 < skoX2C1 * (skoAC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_110, refute_0_112])).

cnf(refute_0_114, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoTC1 * skoAC1 <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_110,
               refute_0_113])).

cnf(refute_0_115, plain,
    (skoTC1 * skoAC1 <
     (skoAC1 * skoAC1 +
      cos(skoTC1 * skoAC1) *
      (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_114, refute_0_107])).

cnf(refute_0_116, plain,
    (skoAC1 * skoAC1 +
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))) <
     skoTC1 * skoAC1 * (exp(skoTC1 * -1) * (skoX2C1 * skoAC1)) |
     0 <= exp(skoTC1 * -1) * (skoX2C1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_99, refute_0_115])).

cnf(refute_0_117, plain,
    (skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) < 0 |
     skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     0 | 0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(split, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    ((skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) <
     cos(skoTC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_117])).

cnf(refute_0_121, plain,
    ((skoAC1 * skoAC1 +
      exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))))) /
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) < 1 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_42, refute_0_120])).

cnf(refute_0_122, plain,
    (skoAC1 * skoAC1 +
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1)))) <
     1 * (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_37, refute_0_121])).

cnf(refute_0_123, plain,
    (skoAC1 * skoAC1 <
     exp(skoTC1 * -1) *
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) | exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (0 < skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_126, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) | exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (0 < skoTC1 * -1 |
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_50, refute_0_126])).

cnf(refute_0_128, plain,
    (skoTC1 < 0 |
     skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (0 <= skoTC1 | exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_130, plain,
    (skoAC1 * skoAC1 /
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_128])).

cnf(refute_0_131, plain,
    (skoAC1 * skoAC1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_36, refute_0_130])).

cnf(refute_0_132, plain,
    (skoAC1 * skoAC1 <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoX2C1 * (skoAC1 * skoAC1 + skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (skoAC1 * skoAC1 <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_132, refute_0_125])).

cnf(refute_0_134, plain,
    (skoAC1 * skoAC1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0), inference(resolve, [], [refute_0_35, refute_0_133])).

cnf(refute_0_135, plain,
    (skoAC1 * (skoAC1 * 2304) +
     skoTC1 *
     (skoAC1 * (skoAC1 * 2304) +
      skoTC1 *
      (skoAC1 * (skoAC1 * 1152) +
       skoTC1 *
       (skoAC1 * (skoAC1 * 384) +
        skoTC1 *
        (skoAC1 * (skoAC1 * 84) +
         skoTC1 *
         (skoAC1 * (skoAC1 * 12) + skoTC1 * (skoAC1 * skoAC1)))))) <
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_134])).

cnf(refute_0_136, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304 |
     skoX2C1 *
     (skoAC1 * (skoAC1 * 2304) + skoTC1 * (skoAC1 * (skoAC1 * 2304))) <=
     skoAC1 * (skoAC1 * 2304) +
     skoTC1 *
     (skoAC1 * (skoAC1 * 2304) +
      skoTC1 *
      (skoAC1 * (skoAC1 * 1152) +
       skoTC1 *
       (skoAC1 * (skoAC1 * 384) +
        skoTC1 *
        (skoAC1 * (skoAC1 * 84) +
         skoTC1 *
         (skoAC1 * (skoAC1 * 12) + skoTC1 * (skoAC1 * skoAC1))))))),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_137, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(resolve, [], [refute_0_136, refute_0_135])).

cnf(refute_0_138, plain,
    (-2304 <
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_139, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_139, refute_0_119])).

cnf(refute_0_141, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoAC1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_20, refute_0_21, refute_0_22, refute_0_110,
               refute_0_113])).

cnf(refute_0_142, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(resolve, [], [refute_0_141, refute_0_70])).

cnf(refute_0_143, plain,
    (skoAC1 * skoAC1 <
     exp(skoTC1 * -1) * (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_142, refute_0_140])).

cnf(refute_0_144, plain,
    (skoAC1 * skoAC1 / (skoX2C1 * (skoTC1 * (skoAC1 * skoAC1))) <
     exp(skoTC1 * -1) | exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (1 / (skoX2C1 * skoTC1) < exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0 | skoAC1 = 0),
    inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain,
    (skoAC1 != 0 | exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_147, plain,
    (1 / (skoX2C1 * skoTC1) < exp(skoTC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    (0 < skoTC1 * -1 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 * -1 *
      (-2304 +
       skoTC1 * -1 *
       (1152 +
        skoTC1 * -1 *
        (-384 +
         skoTC1 * -1 * (84 + skoTC1 * -1 * (-12 + skoTC1 * -1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_147])).

cnf(refute_0_149, plain,
    (skoTC1 < 0 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (skoAC1 * skoAC1 < 0 | skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) < 0 |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(split, [], [refute_0_143])).

cnf(refute_0_151, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * (skoTC1 * (skoAC1 * (skoAC1 * -1))) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_153, plain,
    (0 < skoAC1 * (skoAC1 * -1) |
     0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_152, refute_0_151])).

cnf(refute_0_154, plain,
    (0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoAC1 * (skoAC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_113, refute_0_110, refute_0_22, refute_0_21,
               refute_0_20])).

cnf(refute_0_155, plain,
    (0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_154, refute_0_153])).

cnf(refute_0_156, plain,
    (skoTC1 < 0 | 0 < skoX2C1 * (skoTC1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_155, refute_0_113, refute_0_110, refute_0_22,
               refute_0_21, refute_0_20])).

cnf(refute_0_157, plain,
    (skoTC1 < 0 |
     1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_149, refute_0_156])).

cnf(refute_0_158, plain,
    (1 / (skoX2C1 * skoTC1) <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_157])).

cnf(refute_0_159, plain,
    (1 <
     2304 /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) *
     (skoX2C1 * skoTC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * skoTC1 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_158])).

cnf(refute_0_160, plain,
    (1 <
     skoX2C1 * (skoTC1 * 2304) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoX2C1 * skoTC1 <= 0), inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (0 < skoX2C1 * skoTC1 |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_155, refute_0_113, refute_0_110, refute_0_22,
               refute_0_21, refute_0_20])).

cnf(refute_0_162, plain,
    (1 <
     skoX2C1 * (skoTC1 * 2304) /
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (1 *
     (2304 +
      skoTC1 *
      (2304 +
       skoTC1 *
       (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))) <
     skoX2C1 * (skoTC1 * 2304) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     0), inference(resolve, [], [refute_0_4, refute_0_162])).

cnf(refute_0_164, plain,
    (2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <
     skoX2C1 * (skoTC1 * 2304) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304 |
     skoX2C1 * (skoTC1 * 2304) <=
     2304 +
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1)))))),
    inference(decision, [],
              [refute_0_155, refute_0_113, refute_0_110, refute_0_22,
               refute_0_21, refute_0_20])).

cnf(refute_0_166, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0 |
     skoTC1 *
     (2304 +
      skoTC1 *
      (1152 + skoTC1 * (384 + skoTC1 * (84 + skoTC1 * (12 + skoTC1))))) <=
     -2304), inference(resolve, [], [refute_0_165, refute_0_164])).

cnf(refute_0_167, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_166, refute_0_138])).

cnf(refute_0_168, plain,
    (skoAC1 * skoAC1 <
     cos(skoTC1 * skoAC1) *
     (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)))),
    inference(resolve, [], [refute_0_167, refute_0_98])).

cnf(refute_0_169, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) < 0 |
     skoAC1 * skoAC1 < 0 |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))),
    inference(split, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     0 < skoAC1 * (skoAC1 * -1)),
    inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain, (skoAC1 * (skoAC1 * -1) <= 0),
    inference(decision, [], [refute_0_20, refute_0_21, refute_0_22])).

cnf(refute_0_172, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1))),
    inference(resolve, [], [refute_0_171, refute_0_170])).

cnf(refute_0_173, plain,
    (skoX2C1 * (skoAC1 * -1) <= 0 |
     exp(skoTC1 * -1) <= 0 / (skoX2C1 * (skoAC1 * -1))),
    inference(split, [], [refute_0_167])).

cnf(refute_0_174, plain, (skoX2C1 * (skoAC1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_176, plain,
    (0 < 0 | skoX2C1 * (skoAC1 * skoAC1) < 0 |
     0 < skoX2C1 * (skoAC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_175])).

cnf(refute_0_177, plain,
    (0 < skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     0 < skoX2C1 * (skoAC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (exp(skoTC1 * -1) < 0 / (skoX2C1 * (skoAC1 * skoAC1)) |
     0 <= skoX2C1 * (skoAC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(split, [], [refute_0_175])).

cnf(refute_0_179, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0 |
     skoX2C1 * (skoAC1 * (skoAC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_178])).

cnf(refute_0_180, plain,
    (0 < skoX2C1 * (skoAC1 * skoAC1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_179, refute_0_177])).

cnf(refute_0_181, plain,
    (skoX2C1 * skoAC1 <= 0 | exp(skoTC1 * -1) <= 0 / (skoX2C1 * skoAC1)),
    inference(split, [], [refute_0_97])).

cnf(refute_0_182, plain, (skoX2C1 * skoAC1 <= 0),
    inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (0 < skoX2C1 * (skoAC1 * -1) |
     exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(decision, [],
              [refute_0_180, refute_0_182, refute_0_20, refute_0_21,
               refute_0_22])).

cnf(refute_0_184, plain,
    (exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * skoAC1)) <= 0),
    inference(resolve, [], [refute_0_174, refute_0_183])).

cnf(refute_0_185, plain,
    (0 < exp(skoTC1 * -1) * (skoX2C1 * (skoAC1 * (skoAC1 * -1)))),
    inference(resolve, [], [refute_0_184, refute_0_172])).

cnf(refute_0_186, plain,
    (0 < 0 | skoX2C1 * (skoAC1 * (skoAC1 * -1)) < 0 |
     0 < skoX2C1 * (skoAC1 * (skoAC1 * -1))),
    inference(split, [], [refute_0_185])).

cnf(refute_0_187, plain,
    (0 < skoX2C1 * (skoAC1 * (skoAC1 * -1)) |
     0 < skoX2C1 * (skoAC1 * skoAC1)),
    inference(arithmetic, [], [refute_0_186])).

cnf(refute_0_188, plain,
    (exp(skoTC1 * -1) < 0 / (skoX2C1 * (skoAC1 * (skoAC1 * -1))) |
     0 <= skoX2C1 * (skoAC1 * (skoAC1 * -1))),
    inference(split, [], [refute_0_185])).

cnf(refute_0_189, plain, (skoX2C1 * (skoAC1 * skoAC1) <= 0),
    inference(arithmetic, [], [refute_0_188])).

cnf(refute_0_190, plain, (0 < skoX2C1 * (skoAC1 * (skoAC1 * -1))),
    inference(resolve, [], [refute_0_189, refute_0_187])).

cnf(refute_0_191, plain, (skoX2C1 * (skoAC1 * (skoAC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_22, refute_0_21, refute_0_20, refute_0_182])).

cnf(refute_0_192, plain, ($false),
    inference(resolve, [], [refute_0_191, refute_0_190])).
% SZS output end CNFRefutation for KeYmaera_damped_pendulum_general.tptp
