% SZS output start CNFRefutation for log-fun-ineq-fneg2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_lower_bound_cf6, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            (X ^ 6 + 42 * X ^ 5 + 840 * X ^ 4 + 10080 * X ^ 3 +
             75600 * X ^ 2 + 332640 * X + 665280) /
            (X ^ 6 - 42 * X ^ 5 + 840 * X ^ 4 - 10080 * X ^ 3 +
             75600 * X ^ 2 - 332640 * X + 665280)) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf5, axiom,
    (X < 0 | 7 < X |
     ~ lgen(R,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240), Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf7, axiom,
    (X < 0 | 19/2 < X |
     ~ lgen(R,
            -(X ^ 7 + 56 * X ^ 6 + 1512 * X ^ 5 + 25200 * X ^ 4 +
             277200 * X ^ 3 + 1995840 * X ^ 2 + 8648640 * X + 17297280) /
            (X ^ 7 - 56 * X ^ 6 + 1512 * X ^ 5 - 25200 * X ^ 4 +
             277200 * X ^ 3 - 1995840 * X ^ 2 + 8648640 * X - 17297280),
            Y) | lgen(R, exp(X), Y))).

cnf(ln_lower_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (142 * X ^ 5 + 2272 * X ^ 4 + 6397 * X ^ 3 + 4397 * X ^ 2 +
             647 * X + 5) * (X - 1) /
            (X *
             (X ^ 5 + 30 * X ^ 4 + 150 * X ^ 3 + 200 * X ^ 2 + 75 * X +
              6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/70 *
            (10 + 1921 * X + 20149 * X ^ 2 + 50774 * X ^ 3 +
             38524 * X ^ 4 + 8389 * X ^ 5 + 353 * X ^ 6) * (X - 1) /
            (X *
             (7 + 126 * X + 525 * X ^ 2 + 700 * X ^ 3 + 315 * X ^ 4 +
              42 * X ^ 5 + X ^ 6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf15, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/280 *
            (35 + 9411 * X + 141123 * X ^ 2 + 547235 * X ^ 3 +
             718735 * X ^ 4 + 334575 * X ^ 5 + 49199 * X ^ 6 +
             1487 * X ^ 7) * (X - 1) /
            (X *
             (8 + 196 * X + 1176 * X ^ 2 + 2450 * X ^ 3 + 1960 * X ^ 4 +
              588 * X ^ 5 + 56 * X ^ 6 + X ^ 7))) | lgen(R, Y, ln(X)))).

fof(lfi_f, conjecture,
    (! [X] : (0 < X => -ln(X) <= 10 ^ -13 + 1 / (exp(1) * 2 * X ^ 2)))).

fof(subgoal_0, plain,
    (! [X] : (0 < X => -ln(X) <= 10 ^ -13 + 1 / (exp(1) * 2 * X ^ 2))),
    inference(strip, [], [lfi_f])).

fof(negate_0_0, plain,
    (~ ! [X] : (0 < X => -ln(X) <= 10 ^ -13 + 1 / (exp(1) * 2 * X ^ 2))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (0 < X & 10 ^ -13 + 1 / (exp(1) * 2 * X ^ 2) < -ln(X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (10 ^ -13 + 1 / (exp(1) * 2 * skoXC1 ^ 2) < -ln(skoXC1) & 0 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (10 ^ -13 + 1 / (exp(1) * 2 * skoXC1 ^ 2) < -ln(skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     28245729/10391023 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     28245729/10391023 <=
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (1/625000000000 +
        skoXC1 *
        (49/1250000000000 +
         skoXC1 *
         (147/625000000000 +
          skoXC1 *
          (49/100000000000 +
           skoXC1 *
           (49/125000000000 +
            skoXC1 *
            (147/1250000000000 +
             skoXC1 * (7/625000000000 + skoXC1 * 1/5000000000000))))))))) <
     (-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2) |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (1/625000000000 +
         skoXC1 *
         (49/1250000000000 +
          skoXC1 *
          (147/625000000000 +
           skoXC1 *
           (49/100000000000 +
            skoXC1 *
            (49/125000000000 +
             skoXC1 *
             (147/1250000000000 +
              skoXC1 *
              (7/625000000000 + skoXC1 * 1/5000000000000)))))))))) /
     (exp(1) * (skoXC1 * -2)) <=
     -1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (-1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))))) |
     skoXC1 *
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) <=
     0 |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <=
     (-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) /
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_3, plain,
    (10 ^ -13 + 1 / (exp(1) * 2 * skoXC1 ^ 2) < -ln(skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_4, plain,
    (1 / (exp(1) * (skoXC1 * (skoXC1 * 2))) <
     -1/10000000000000 + ln(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (1 <
     (-1/10000000000000 + ln(skoXC1) * -1) *
     (exp(1) * (skoXC1 * (skoXC1 * 2))) |
     exp(1) * (skoXC1 * (skoXC1 * 2)) <= 0 |
     -1/10000000000000 + ln(skoXC1) * -1 <=
     1 / (exp(1) * (skoXC1 * (skoXC1 * 2)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (1 <
     (-1/10000000000000 + ln(skoXC1) * -1) *
     (exp(1) * (skoXC1 * (skoXC1 * 2))) |
     exp(1) * (skoXC1 * (skoXC1 * 2)) <= 0),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000)) <
     ln(skoXC1) * (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     exp(1) * (skoXC1 * (skoXC1 * 2)) <= 0),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (0 < exp(1) * (skoXC1 * (skoXC1 * 2)) |
     exp(1) * (skoXC1 * (skoXC1 * 2)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_9, plain,
    (0 < exp(1) * (skoXC1 * (skoXC1 * 2)) | skoXC1 * (skoXC1 * 2) <= 0 |
     exp(1) <= 0 / (skoXC1 * (skoXC1 * 2))),
    inference(split, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (0 < exp(1) * (skoXC1 * (skoXC1 * 2)) | skoXC1 * (skoXC1 * 2) <= 0),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_12, plain, (0 < skoXC1 * (skoXC1 * 2)),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_13, plain, (0 < exp(1) * (skoXC1 * (skoXC1 * 2))),
    inference(resolve, [], [refute_0_10, refute_0_12])).

cnf(refute_0_14, plain,
    (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000)) <
     ln(skoXC1) * (exp(1) * (skoXC1 * (skoXC1 * -2)))),
    inference(resolve, [], [refute_0_7, refute_0_13])).

cnf(refute_0_15, plain,
    (ln(skoXC1) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     0 <= exp(1) * (skoXC1 * (skoXC1 * -2))),
    inference(split, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (ln(skoXC1) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     exp(1) * (skoXC1 * (skoXC1 * 2)) <= 0),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (ln(skoXC1) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2)))),
    inference(resolve, [], [refute_0_16, refute_0_13])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000111, ln(X_000110)) | X_000111 <= ln(X_000110)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000111,
            1/280 *
            (35 + 9411 * X_000110 + 141123 * X_000110 ^ 2 +
             547235 * X_000110 ^ 3 + 718735 * X_000110 ^ 4 +
             334575 * X_000110 ^ 5 + 49199 * X_000110 ^ 6 +
             1487 * X_000110 ^ 7) * (X_000110 - 1) /
            (X_000110 *
             (8 + 196 * X_000110 + 1176 * X_000110 ^ 2 +
              2450 * X_000110 ^ 3 + 1960 * X_000110 ^ 4 +
              588 * X_000110 ^ 5 + 56 * X_000110 ^ 6 + X_000110 ^ 7))) |
     X_000110 <= 0 | lgen(0, X_000111, ln(X_000110))),
    inference(subst, [], [ln_lower_bound_cf15])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000111,
            1/280 *
            (35 + 9411 * X_000110 + 141123 * X_000110 ^ 2 +
             547235 * X_000110 ^ 3 + 718735 * X_000110 ^ 4 +
             334575 * X_000110 ^ 5 + 49199 * X_000110 ^ 6 +
             1487 * X_000110 ^ 7) * (X_000110 - 1) /
            (X_000110 *
             (8 + 196 * X_000110 + 1176 * X_000110 ^ 2 +
              2450 * X_000110 ^ 3 + 1960 * X_000110 ^ 4 +
              588 * X_000110 ^ 5 + 56 * X_000110 ^ 6 + X_000110 ^ 7))) |
     X_000110 <= 0 | X_000111 <= ln(X_000110)),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    ((-1/8 +
      X_000110 *
      (-1172/35 +
       X_000110 *
       (-2352/5 +
        X_000110 *
        (-7252/5 +
         X_000110 *
         (-1225/2 +
          X_000110 *
          (1372 +
           X_000110 *
           (5096/5 + X_000110 * (852/5 + X_000110 * 1487/280)))))))) /
     (X_000110 *
      (8 +
       X_000110 *
       (196 +
        X_000110 *
        (1176 +
         X_000110 *
         (2450 +
          X_000110 *
          (1960 + X_000110 * (588 + X_000110 * (56 + X_000110)))))))) <
     X_000111 | X_000110 <= 0 | X_000111 <= ln(X_000110)),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    ((-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) /
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    ((-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) /
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_22, refute_0_17])).

cnf(refute_0_24, plain,
    ((-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) /
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2)))),
    inference(resolve, [], [refute_0_23, refute_0_11])).

cnf(refute_0_25, plain,
    (-1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (8 +
       skoXC1 *
       (196 +
        skoXC1 *
        (1176 +
         skoXC1 *
         (2450 +
          skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))))) |
     skoXC1 *
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) <=
     0), inference(resolve, [], [refute_0_2, refute_0_24])).

cnf(refute_0_26, plain,
    (-1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280))))))) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (1/625000000000 +
         skoXC1 *
         (49/1250000000000 +
          skoXC1 *
          (147/625000000000 +
           skoXC1 *
           (49/100000000000 +
            skoXC1 *
            (49/125000000000 +
             skoXC1 *
             (147/1250000000000 +
              skoXC1 *
              (7/625000000000 + skoXC1 * 1/5000000000000)))))))))) /
     (exp(1) * (skoXC1 * -2)) |
     skoXC1 *
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) <=
     0 | skoXC1 = 0), inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     1084483/398959 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     0 <=
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) |
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) <=
     1084483/398959), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_28, plain,
    (~ lgen(0, X_000062, exp(X_000061)) | X_000062 <= exp(X_000061)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_29, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062,
            (X_000061 ^ 6 + 42 * X_000061 ^ 5 + 840 * X_000061 ^ 4 +
             10080 * X_000061 ^ 3 + 75600 * X_000061 ^ 2 +
             332640 * X_000061 + 665280) /
            (X_000061 ^ 6 - 42 * X_000061 ^ 5 + 840 * X_000061 ^ 4 -
             10080 * X_000061 ^ 3 + 75600 * X_000061 ^ 2 -
             332640 * X_000061 + 665280)) |
     lgen(0, X_000062, exp(X_000061))),
    inference(subst, [], [exp_lower_bound_cf6])).

cnf(refute_0_30, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062,
            (X_000061 ^ 6 + 42 * X_000061 ^ 5 + 840 * X_000061 ^ 4 +
             10080 * X_000061 ^ 3 + 75600 * X_000061 ^ 2 +
             332640 * X_000061 + 665280) /
            (X_000061 ^ 6 - 42 * X_000061 ^ 5 + 840 * X_000061 ^ 4 -
             10080 * X_000061 ^ 3 + 75600 * X_000061 ^ 2 -
             332640 * X_000061 + 665280)) | X_000062 <= exp(X_000061)),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain,
    ((665280 +
      X_000061 *
      (332640 +
       X_000061 *
       (75600 +
        X_000061 *
        (10080 + X_000061 * (840 + X_000061 * (42 + X_000061)))))) /
     (665280 +
      X_000061 *
      (-332640 +
       X_000061 *
       (75600 +
        X_000061 *
        (-10080 + X_000061 * (840 + X_000061 * (-42 + X_000061)))))) <
     X_000062 | X_000061 < 0 | X_000062 <= exp(X_000061)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) <= exp(1)),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (3/2500000000000 +
        skoXC1 *
        (3/200000000000 +
         skoXC1 *
         (1/25000000000 +
          skoXC1 *
          (3/100000000000 +
           skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000))))))) <
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2) |
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (3/2500000000000 +
         skoXC1 *
         (3/200000000000 +
          skoXC1 *
          (1/25000000000 +
           skoXC1 *
           (3/100000000000 +
            skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000)))))))) /
     (exp(1) * (skoXC1 * -2)) <=
     -1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_34, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0 |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <=
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_35, plain,
    (~ lgen(0, X_000076, ln(X_000075)) | X_000076 <= ln(X_000075)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000076,
            1/30 *
            (142 * X_000075 ^ 5 + 2272 * X_000075 ^ 4 +
             6397 * X_000075 ^ 3 + 4397 * X_000075 ^ 2 + 647 * X_000075 +
             5) * (X_000075 - 1) /
            (X_000075 *
             (X_000075 ^ 5 + 30 * X_000075 ^ 4 + 150 * X_000075 ^ 3 +
              200 * X_000075 ^ 2 + 75 * X_000075 + 6))) | X_000075 <= 0 |
     lgen(0, X_000076, ln(X_000075))),
    inference(subst, [], [ln_lower_bound_cf11])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000076,
            1/30 *
            (142 * X_000075 ^ 5 + 2272 * X_000075 ^ 4 +
             6397 * X_000075 ^ 3 + 4397 * X_000075 ^ 2 + 647 * X_000075 +
             5) * (X_000075 - 1) /
            (X_000075 *
             (X_000075 ^ 5 + 30 * X_000075 ^ 4 + 150 * X_000075 ^ 3 +
              200 * X_000075 ^ 2 + 75 * X_000075 + 6))) | X_000075 <= 0 |
     X_000076 <= ln(X_000075)),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    ((-1/6 +
      X_000075 *
      (-107/5 +
       X_000075 *
       (-125 +
        X_000075 *
        (-200/3 +
         X_000075 * (275/2 + X_000075 * (71 + X_000075 * 71/15)))))) /
     (X_000075 *
      (6 +
       X_000075 *
       (75 +
        X_000075 *
        (200 + X_000075 * (150 + X_000075 * (30 + X_000075)))))) <
     X_000076 | X_000075 <= 0 | X_000076 <= ln(X_000075)),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_40, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_39, refute_0_17])).

cnf(refute_0_41, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2)))),
    inference(resolve, [], [refute_0_40, refute_0_11])).

cnf(refute_0_42, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_34, refute_0_41])).

cnf(refute_0_43, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (3/2500000000000 +
         skoXC1 *
         (3/200000000000 +
          skoXC1 *
          (1/25000000000 +
           skoXC1 *
           (3/100000000000 +
            skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000)))))))) /
     (exp(1) * (skoXC1 * -2)) |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0 | skoXC1 = 0), inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (0 <
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) |
     skoXC1 = 0), inference(decision, [], [refute_0_11])).

cnf(refute_0_45, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (3/2500000000000 +
         skoXC1 *
         (3/200000000000 +
          skoXC1 *
          (1/25000000000 +
           skoXC1 *
           (3/100000000000 +
            skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000)))))))) /
     (exp(1) * (skoXC1 * -2)) | skoXC1 = 0),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain, (skoXC1 != 0),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_47, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (3/2500000000000 +
         skoXC1 *
         (3/200000000000 +
          skoXC1 *
          (1/25000000000 +
           skoXC1 *
           (3/100000000000 +
            skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000)))))))) /
     (exp(1) * (skoXC1 * -2))),
    inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (3/2500000000000 +
        skoXC1 *
        (3/200000000000 +
         skoXC1 *
         (1/25000000000 +
          skoXC1 *
          (3/100000000000 +
           skoXC1 * (3/500000000000 + skoXC1 * 1/5000000000000))))))) <
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2)),
    inference(resolve, [], [refute_0_33, refute_0_47])).

cnf(refute_0_49, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     exp(1) *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     exp(1) * (skoXC1 * 2) <= 0),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 < exp(1) * (skoXC1 * 2) | exp(1) * (skoXC1 * 2) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_51, plain,
    (0 < exp(1) * (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     exp(1) <= 0 / (skoXC1 * 2)), inference(split, [], [refute_0_50])).

cnf(refute_0_52, plain, (0 < exp(1) * (skoXC1 * 2) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain, (0 < exp(1) * (skoXC1 * 2)),
    inference(resolve, [], [refute_0_52, refute_0_11])).

cnf(refute_0_54, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     exp(1) *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000)))))))),
    inference(resolve, [], [refute_0_49, refute_0_53])).

cnf(refute_0_55, plain,
    (exp(1) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     0 <=
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000))))))),
    inference(split, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (exp(1) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(resolve, [], [refute_0_32, refute_0_56])).

cnf(refute_0_58, plain,
    (1084483/398959 <
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     1084483/398959 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     0 <=
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(resolve, [], [refute_0_27, refute_0_58])).

cnf(refute_0_60, plain,
    (6 <
     skoXC1 *
     (-88681292/1196877 +
      skoXC1 *
      (-83439819000003253449/997397500000000000 +
       skoXC1 *
       (42255379999996746551/79791800000000000 +
        skoXC1 *
        (9947172249996746551/29921925000000000 +
         skoXC1 *
         (-29863178400003253449/39895900000000000 +
          skoXC1 *
          (-76998293000003253449/199479500000000000 +
           skoXC1 * -153996586000003253449/5984385000000000000)))))) |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 *
     (-88681292/1196877 +
      skoXC1 *
      (-83439819000003253449/997397500000000000 +
       skoXC1 *
       (42255379999996746551/79791800000000000 +
        skoXC1 *
        (9947172249996746551/29921925000000000 +
         skoXC1 *
         (-29863178400003253449/39895900000000000 +
          skoXC1 *
          (-76998293000003253449/199479500000000000 +
           skoXC1 * -153996586000003253449/5984385000000000000)))))) <= 6 |
     skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(decision, [], [refute_0_11])).

cnf(refute_0_62, plain,
    (skoXC1 *
     (-1/3 +
      skoXC1 *
      (-106999999999997/2500000000000 +
       skoXC1 *
       (-49999999999997/200000000000 +
        skoXC1 *
        (-9999999999997/75000000000 +
         skoXC1 *
         (27500000000003/100000000000 +
          skoXC1 *
          (71000000000003/500000000000 +
           skoXC1 * 142000000000003/15000000000000)))))) <= 0),
    inference(resolve, [], [refute_0_61, refute_0_60])).

cnf(refute_0_63, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     1084483/398959 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     0 <=
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <=
     1084483/398959), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_64, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <= exp(1)),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_65, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (7/5000000000000 +
        skoXC1 *
        (63/2500000000000 +
         skoXC1 *
         (21/200000000000 +
          skoXC1 *
          (7/50000000000 +
           skoXC1 *
           (63/1000000000000 +
            skoXC1 * (21/2500000000000 + skoXC1 * 1/5000000000000)))))))) <
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2) |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (7/5000000000000 +
         skoXC1 *
         (63/2500000000000 +
          skoXC1 *
          (21/200000000000 +
           skoXC1 *
           (7/50000000000 +
            skoXC1 *
            (63/1000000000000 +
             skoXC1 *
             (21/2500000000000 + skoXC1 * 1/5000000000000))))))))) /
     (exp(1) * (skoXC1 * -2)) <=
     -1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_66, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0 |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <=
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_67, plain,
    (~ lgen(0, X_000090, ln(X_000089)) | X_000090 <= ln(X_000089)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_68, plain,
    (~ lgen(0, X_000090,
            1/70 *
            (10 + 1921 * X_000089 + 20149 * X_000089 ^ 2 +
             50774 * X_000089 ^ 3 + 38524 * X_000089 ^ 4 +
             8389 * X_000089 ^ 5 + 353 * X_000089 ^ 6) * (X_000089 - 1) /
            (X_000089 *
             (7 + 126 * X_000089 + 525 * X_000089 ^ 2 +
              700 * X_000089 ^ 3 + 315 * X_000089 ^ 4 + 42 * X_000089 ^ 5 +
              X_000089 ^ 6))) | X_000089 <= 0 |
     lgen(0, X_000090, ln(X_000089))),
    inference(subst, [], [ln_lower_bound_cf13])).

cnf(refute_0_69, plain,
    (~ lgen(0, X_000090,
            1/70 *
            (10 + 1921 * X_000089 + 20149 * X_000089 ^ 2 +
             50774 * X_000089 ^ 3 + 38524 * X_000089 ^ 4 +
             8389 * X_000089 ^ 5 + 353 * X_000089 ^ 6) * (X_000089 - 1) /
            (X_000089 *
             (7 + 126 * X_000089 + 525 * X_000089 ^ 2 +
              700 * X_000089 ^ 3 + 315 * X_000089 ^ 4 + 42 * X_000089 ^ 5 +
              X_000089 ^ 6))) | X_000089 <= 0 | X_000090 <= ln(X_000089)),
    inference(resolve, [], [refute_0_68, refute_0_67])).

cnf(refute_0_70, plain,
    ((-1/7 +
      X_000089 *
      (-273/10 +
       X_000089 *
       (-1302/5 +
        X_000089 *
        (-875/2 +
         X_000089 *
         (175 +
          X_000089 *
          (861/2 + X_000089 * (574/5 + X_000089 * 353/70))))))) /
     (X_000089 *
      (7 +
       X_000089 *
       (126 +
        X_000089 *
        (525 +
         X_000089 *
         (700 + X_000089 * (315 + X_000089 * (42 + X_000089))))))) <
     X_000090 | X_000089 <= 0 | X_000090 <= ln(X_000089)),
    inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) |
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_70])).

cnf(refute_0_72, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_71, refute_0_17])).

cnf(refute_0_73, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2)))),
    inference(resolve, [], [refute_0_72, refute_0_11])).

cnf(refute_0_74, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (1 + exp(1) * (skoXC1 * (skoXC1 * 1/5000000000000))) /
     (exp(1) * (skoXC1 * (skoXC1 * -2))) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_66, refute_0_73])).

cnf(refute_0_75, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (7/5000000000000 +
         skoXC1 *
         (63/2500000000000 +
          skoXC1 *
          (21/200000000000 +
           skoXC1 *
           (7/50000000000 +
            skoXC1 *
            (63/1000000000000 +
             skoXC1 *
             (21/2500000000000 + skoXC1 * 1/5000000000000))))))))) /
     (exp(1) * (skoXC1 * -2)) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (0 <
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) |
     skoXC1 = 0), inference(decision, [], [refute_0_11, refute_0_62])).

cnf(refute_0_77, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (7/5000000000000 +
         skoXC1 *
         (63/2500000000000 +
          skoXC1 *
          (21/200000000000 +
           skoXC1 *
           (7/50000000000 +
            skoXC1 *
            (63/1000000000000 +
             skoXC1 *
             (21/2500000000000 + skoXC1 * 1/5000000000000))))))))) /
     (exp(1) * (skoXC1 * -2)) | skoXC1 = 0),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (7/5000000000000 +
         skoXC1 *
         (63/2500000000000 +
          skoXC1 *
          (21/200000000000 +
           skoXC1 *
           (7/50000000000 +
            skoXC1 *
            (63/1000000000000 +
             skoXC1 *
             (21/2500000000000 + skoXC1 * 1/5000000000000))))))))) /
     (exp(1) * (skoXC1 * -2))),
    inference(resolve, [], [refute_0_77, refute_0_46])).

cnf(refute_0_79, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (7/5000000000000 +
        skoXC1 *
        (63/2500000000000 +
         skoXC1 *
         (21/200000000000 +
          skoXC1 *
          (7/50000000000 +
           skoXC1 *
           (63/1000000000000 +
            skoXC1 * (21/2500000000000 + skoXC1 * 1/5000000000000)))))))) <
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2)),
    inference(resolve, [], [refute_0_65, refute_0_78])).

cnf(refute_0_80, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     exp(1) *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     exp(1) * (skoXC1 * 2) <= 0),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     exp(1) *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000))))))))),
    inference(resolve, [], [refute_0_80, refute_0_53])).

cnf(refute_0_82, plain,
    (exp(1) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     0 <=
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000)))))))),
    inference(split, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (exp(1) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_64, refute_0_83])).

cnf(refute_0_85, plain,
    (1084483/398959 <
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     1084483/398959 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     0 <=
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_63, refute_0_85])).

cnf(refute_0_87, plain,
    (7 <
     skoXC1 *
     (-349712872/2792713 +
      skoXC1 *
      (-751203516000007591381/1994795000000000000 +
       skoXC1 *
       (713818615999931677571/997397500000000000 +
        skoXC1 *
        (14968191636361565987/7253800000000000 +
         skoXC1 *
         (-19816266400007591381/19947950000000000 +
          skoXC1 *
          (-84921711090915302039/36269000000000000 +
           skoXC1 *
           (-622493242000022774143/997397500000000000 +
            skoXC1 * -382822499000007591381/13963565000000000000))))))) |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 *
     (-349712872/2792713 +
      skoXC1 *
      (-751203516000007591381/1994795000000000000 +
       skoXC1 *
       (713818615999931677571/997397500000000000 +
        skoXC1 *
        (14968191636361565987/7253800000000000 +
         skoXC1 *
         (-19816266400007591381/19947950000000000 +
          skoXC1 *
          (-84921711090915302039/36269000000000000 +
           skoXC1 *
           (-622493242000022774143/997397500000000000 +
            skoXC1 * -382822499000007591381/13963565000000000000))))))) <=
     7 |
     skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(decision, [], [refute_0_62, refute_0_11])).

cnf(refute_0_89, plain,
    (skoXC1 *
     (-2/7 +
      skoXC1 *
      (-272999999999993/5000000000000 +
       skoXC1 *
       (-1301999999999937/2500000000000 +
        skoXC1 *
        (-174999999999979/200000000000 +
         skoXC1 *
         (17500000000007/50000000000 +
          skoXC1 *
          (861000000000063/1000000000000 +
           skoXC1 *
           (574000000000021/2500000000000 +
            skoXC1 * 353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_87])).

cnf(refute_0_90, plain,
    (0 <
     skoXC1 *
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) |
     skoXC1 = 0),
    inference(decision, [], [refute_0_11, refute_0_62, refute_0_89])).

cnf(refute_0_91, plain,
    (-1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280))))))) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (1/625000000000 +
         skoXC1 *
         (49/1250000000000 +
          skoXC1 *
          (147/625000000000 +
           skoXC1 *
           (49/100000000000 +
            skoXC1 *
            (49/125000000000 +
             skoXC1 *
             (147/1250000000000 +
              skoXC1 *
              (7/625000000000 + skoXC1 * 1/5000000000000)))))))))) /
     (exp(1) * (skoXC1 * -2)) | skoXC1 = 0),
    inference(resolve, [], [refute_0_26, refute_0_90])).

cnf(refute_0_92, plain,
    (-1/8 +
     skoXC1 *
     (-1172/35 +
      skoXC1 *
      (-2352/5 +
       skoXC1 *
       (-7252/5 +
        skoXC1 *
        (-1225/2 +
         skoXC1 *
         (1372 +
          skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280))))))) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
      exp(1) *
      (skoXC1 *
       (skoXC1 *
        (1/625000000000 +
         skoXC1 *
         (49/1250000000000 +
          skoXC1 *
          (147/625000000000 +
           skoXC1 *
           (49/100000000000 +
            skoXC1 *
            (49/125000000000 +
             skoXC1 *
             (147/1250000000000 +
              skoXC1 *
              (7/625000000000 + skoXC1 * 1/5000000000000)))))))))) /
     (exp(1) * (skoXC1 * -2))),
    inference(resolve, [], [refute_0_91, refute_0_46])).

cnf(refute_0_93, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) +
     exp(1) *
     (skoXC1 *
      (skoXC1 *
       (1/625000000000 +
        skoXC1 *
        (49/1250000000000 +
         skoXC1 *
         (147/625000000000 +
          skoXC1 *
          (49/100000000000 +
           skoXC1 *
           (49/125000000000 +
            skoXC1 *
            (147/1250000000000 +
             skoXC1 * (7/625000000000 + skoXC1 * 1/5000000000000))))))))) <
     (-1/8 +
      skoXC1 *
      (-1172/35 +
       skoXC1 *
       (-2352/5 +
        skoXC1 *
        (-7252/5 +
         skoXC1 *
         (-1225/2 +
          skoXC1 *
          (1372 +
           skoXC1 * (5096/5 + skoXC1 * (852/5 + skoXC1 * 1487/280)))))))) *
     (exp(1) * (skoXC1 * -2)) | 0 <= exp(1) * (skoXC1 * -2)),
    inference(resolve, [], [refute_0_1, refute_0_92])).

cnf(refute_0_94, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     exp(1) *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     exp(1) * (skoXC1 * 2) <= 0),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     exp(1) *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000)))))))))),
    inference(resolve, [], [refute_0_94, refute_0_53])).

cnf(refute_0_96, plain,
    ((8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) < exp(1) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(split, [], [refute_0_95])).

cnf(refute_0_97, plain,
    (~ lgen(0, exp(X_000134), X_000135) | exp(X_000134) <= X_000135),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_98, plain,
    (19/2 < X_000134 | X_000134 < 0 |
     ~ lgen(0,
            -(X_000134 ^ 7 + 56 * X_000134 ^ 6 + 1512 * X_000134 ^ 5 +
             25200 * X_000134 ^ 4 + 277200 * X_000134 ^ 3 +
             1995840 * X_000134 ^ 2 + 8648640 * X_000134 + 17297280) /
            (X_000134 ^ 7 - 56 * X_000134 ^ 6 + 1512 * X_000134 ^ 5 -
             25200 * X_000134 ^ 4 + 277200 * X_000134 ^ 3 -
             1995840 * X_000134 ^ 2 + 8648640 * X_000134 - 17297280),
            X_000135) | lgen(0, exp(X_000134), X_000135)),
    inference(subst, [], [exp_upper_bound_cf7])).

cnf(refute_0_99, plain,
    (19/2 < X_000134 | X_000134 < 0 |
     ~ lgen(0,
            -(X_000134 ^ 7 + 56 * X_000134 ^ 6 + 1512 * X_000134 ^ 5 +
             25200 * X_000134 ^ 4 + 277200 * X_000134 ^ 3 +
             1995840 * X_000134 ^ 2 + 8648640 * X_000134 + 17297280) /
            (X_000134 ^ 7 - 56 * X_000134 ^ 6 + 1512 * X_000134 ^ 5 -
             25200 * X_000134 ^ 4 + 277200 * X_000134 ^ 3 -
             1995840 * X_000134 ^ 2 + 8648640 * X_000134 - 17297280),
            X_000135) | exp(X_000134) <= X_000135),
    inference(resolve, [], [refute_0_98, refute_0_97])).

cnf(refute_0_100, plain,
    (19/2 < X_000134 | X_000134 < 0 |
     X_000135 <
     (-17297280 +
      X_000134 *
      (-8648640 +
       X_000134 *
       (-1995840 +
        X_000134 *
        (-277200 +
         X_000134 *
         (-25200 +
          X_000134 * (-1512 + X_000134 * (-56 + X_000134 * -1))))))) /
     (-17297280 +
      X_000134 *
      (8648640 +
       X_000134 *
       (-1995840 +
        X_000134 *
        (277200 +
         X_000134 *
         (-25200 + X_000134 * (1512 + X_000134 * (-56 + X_000134))))))) |
     exp(X_000134) <= X_000135), inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (1 < 0 | 19/2 < 1 |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     exp(1) <=
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000)))))))))),
    inference(subst, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (1 < 0 | 19/2 < 1 |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_101, refute_0_96])).

cnf(refute_0_103, plain,
    ((8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     28245729/10391023 |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     49171/18089 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     49171/18089 <=
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_105, plain,
    (~ lgen(0, exp(X_000127), X_000128) | exp(X_000127) <= X_000128),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_106, plain,
    (7 < X_000127 | X_000127 < 0 |
     ~ lgen(0,
            -(X_000127 ^ 5 + 30 * X_000127 ^ 4 + 420 * X_000127 ^ 3 +
             3360 * X_000127 ^ 2 + 15120 * X_000127 + 30240) /
            (X_000127 ^ 5 - 30 * X_000127 ^ 4 + 420 * X_000127 ^ 3 -
             3360 * X_000127 ^ 2 + 15120 * X_000127 - 30240), X_000128) |
     lgen(0, exp(X_000127), X_000128)),
    inference(subst, [], [exp_upper_bound_cf5])).

cnf(refute_0_107, plain,
    (7 < X_000127 | X_000127 < 0 |
     ~ lgen(0,
            -(X_000127 ^ 5 + 30 * X_000127 ^ 4 + 420 * X_000127 ^ 3 +
             3360 * X_000127 ^ 2 + 15120 * X_000127 + 30240) /
            (X_000127 ^ 5 - 30 * X_000127 ^ 4 + 420 * X_000127 ^ 3 -
             3360 * X_000127 ^ 2 + 15120 * X_000127 - 30240), X_000128) |
     exp(X_000127) <= X_000128),
    inference(resolve, [], [refute_0_106, refute_0_105])).

cnf(refute_0_108, plain,
    (7 < X_000127 | X_000127 < 0 |
     X_000128 <
     (-30240 +
      X_000127 *
      (-15120 +
       X_000127 *
       (-3360 + X_000127 * (-420 + X_000127 * (-30 + X_000127 * -1))))) /
     (-30240 +
      X_000127 *
      (15120 +
       X_000127 *
       (-3360 + X_000127 * (420 + X_000127 * (-30 + X_000127))))) |
     exp(X_000127) <= X_000128),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (1 < 0 | 7 < 1 |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000)))))))))),
    inference(subst, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (1 < 0 | 7 < 1 |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_109, refute_0_96])).

cnf(refute_0_111, plain,
    ((8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     49171/18089 |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     49171/18089 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_104, refute_0_111])).

cnf(refute_0_113, plain,
    (8 <
     skoXC1 *
     (-14132605/72356 +
      skoXC1 *
      (-78660802000000344197/79139375000000000 +
       skoXC1 *
       (2427533499997590621/22611250000000000 +
        skoXC1 *
        (66987997999992771863/11305625000000000 +
         skoXC1 *
         (4959814299997590621/1808900000000000 +
          skoXC1 *
          (-16992276000002409379/2261125000000000 +
           skoXC1 *
           (-125310319250007228137/22611250000000000 +
            skoXC1 *
            (-10473423000000344197/11305625000000000 +
             skoXC1 * -18279319250000344197/633115000000000000)))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     49171/18089 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     49171/18089 <=
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_115, plain,
    ((7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) < exp(1) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(split, [], [refute_0_81])).

cnf(refute_0_116, plain,
    (1 < 0 | 7 < 1 |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000))))))))),
    inference(subst, [], [refute_0_108])).

cnf(refute_0_117, plain,
    (1 < 0 | 7 < 1 |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_116, refute_0_115])).

cnf(refute_0_118, plain,
    ((7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     49171/18089 |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     1084483/398959 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     0 <=
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <=
     1084483/398959), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_120, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <= exp(1)),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_121, plain,
    (exp(1) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     0 <=
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000))))))))),
    inference(split, [], [refute_0_95])).

cnf(refute_0_122, plain,
    (exp(1) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (1 < 0 |
     (665280 +
      1 *
      (332640 + 1 * (75600 + 1 * (10080 + 1 * (840 + 1 * (42 + 1)))))) /
     (665280 +
      1 *
      (-332640 + 1 * (75600 + 1 * (-10080 + 1 * (840 + 1 * (-42 + 1)))))) <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_120, refute_0_122])).

cnf(refute_0_124, plain,
    (1084483/398959 <
     (8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     1084483/398959 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     0 <=
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_119, refute_0_124])).

cnf(refute_0_126, plain,
    (8 <
     skoXC1 *
     (-311699373/1595836 +
      skoXC1 *
      (-157717321454546144671/158676875000000000 +
       skoXC1 *
       (53540070499946860333/498698750000000000 +
        skoXC1 *
        (1477442903999840580999/249349375000000000 +
         skoXC1 *
         (109390378299946860333/39895900000000000 +
          skoXC1 *
          (-374770382000053139667/49869875000000000 +
           skoXC1 *
           (-251251034795469038091/45336250000000000 +
            skoXC1 *
            (-20999534454546144671/22668125000000000 +
             skoXC1 * -403156555250007591381/13963565000000000000)))))))) |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (skoXC1 *
     (-311699373/1595836 +
      skoXC1 *
      (-157717321454546144671/158676875000000000 +
       skoXC1 *
       (53540070499946860333/498698750000000000 +
        skoXC1 *
        (1477442903999840580999/249349375000000000 +
         skoXC1 *
         (109390378299946860333/39895900000000000 +
          skoXC1 *
          (-374770382000053139667/49869875000000000 +
           skoXC1 *
           (-251251034795469038091/45336250000000000 +
            skoXC1 *
            (-20999534454546144671/22668125000000000 +
             skoXC1 *
             -403156555250007591381/13963565000000000000)))))))) <= 8 |
     skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(decision, [], [refute_0_89, refute_0_62, refute_0_11])).

cnf(refute_0_128, plain,
    (skoXC1 *
     (-1/4 +
      skoXC1 *
      (-292999999999993/4375000000000 +
       skoXC1 *
       (-1175999999999951/1250000000000 +
        skoXC1 *
        (-1812999999999853/625000000000 +
         skoXC1 *
         (-122499999999951/100000000000 +
          skoXC1 *
          (343000000000049/125000000000 +
           skoXC1 *
           (2548000000000147/1250000000000 +
            skoXC1 *
            (213000000000007/625000000000 +
             skoXC1 * 371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_127, refute_0_126])).

cnf(refute_0_129, plain,
    (0 <
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000)))))))),
    inference(decision, [],
              [refute_0_11, refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_130, plain,
    ((7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     49171/18089), inference(resolve, [], [refute_0_118, refute_0_129])).

cnf(refute_0_131, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     49171/18089 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_114, refute_0_130])).

cnf(refute_0_132, plain,
    (7 <
     skoXC1 *
     (-15856156/126623 +
      skoXC1 *
      (-34059942000000344197/90445000000000000 +
       skoXC1 *
       (32364891999996902227/45222500000000000 +
        skoXC1 *
        (7465317999998967409/3617800000000000 +
         skoXC1 *
         (-898479400000344197/904450000000000 +
          skoXC1 *
          (-42354320000003097773/18089000000000000 +
           skoXC1 *
           (-28224154000001032591/45222500000000000 +
            skoXC1 * -17357363000000344197/633115000000000000))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     49171/18089 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     49171/18089 <=
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_134, plain,
    ((6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) < exp(1) |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(split, [], [refute_0_54])).

cnf(refute_0_135, plain,
    (1 < 0 | 7 < 1 |
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     exp(1) <=
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000)))))))),
    inference(subst, [], [refute_0_108])).

cnf(refute_0_136, plain,
    (1 < 0 | 7 < 1 |
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) <
     (-30240 +
      1 * (-15120 + 1 * (-3360 + 1 * (-420 + 1 * (-30 + 1 * -1))))) /
     (-30240 + 1 * (15120 + 1 * (-3360 + 1 * (420 + 1 * (-30 + 1))))) |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(resolve, [], [refute_0_135, refute_0_134])).

cnf(refute_0_137, plain,
    ((6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) < 49171/18089 |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain,
    (0 <
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000))))))),
    inference(decision, [],
              [refute_0_11, refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_139, plain,
    ((6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) /
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) < 49171/18089),
    inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain,
    (6 +
     skoXC1 *
     (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))) <
     49171/18089 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (106999999999997/2500000000000 +
        skoXC1 *
        (49999999999997/200000000000 +
         skoXC1 *
         (9999999999997/75000000000 +
          skoXC1 *
          (-27500000000003/100000000000 +
           skoXC1 *
           (-71000000000003/500000000000 +
            skoXC1 * -142000000000003/15000000000000))))))) |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(resolve, [], [refute_0_133, refute_0_139])).

cnf(refute_0_141, plain,
    (6 <
     skoXC1 *
     (-4020854/54267 +
      skoXC1 *
      (-3783203000000147513/45222500000000000 +
       skoXC1 *
       (1915879999999852487/3617800000000000 +
        skoXC1 *
        (451009749999852487/1356675000000000 +
         skoXC1 *
         (-1354011400000147513/1808900000000000 +
          skoXC1 *
          (-3491141000000147513/9044500000000000 +
           skoXC1 * -6982282000000147513/271335000000000000)))))) |
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000)))))) <= 0),
    inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (6 <
     skoXC1 *
     (-4020854/54267 +
      skoXC1 *
      (-3783203000000147513/45222500000000000 +
       skoXC1 *
       (1915879999999852487/3617800000000000 +
        skoXC1 *
        (451009749999852487/1356675000000000 +
         skoXC1 *
         (-1354011400000147513/1808900000000000 +
          skoXC1 *
          (-3491141000000147513/9044500000000000 +
           skoXC1 * -6982282000000147513/271335000000000000)))))) |
     0 <
     skoXC1 *
     (1/3 +
      skoXC1 *
      (106999999999997/2500000000000 +
       skoXC1 *
       (49999999999997/200000000000 +
        skoXC1 *
        (9999999999997/75000000000 +
         skoXC1 *
         (-27500000000003/100000000000 +
          skoXC1 *
          (-71000000000003/500000000000 +
           skoXC1 * -142000000000003/15000000000000))))))),
    inference(decision, [],
              [refute_0_11, refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_143, plain,
    (6 <
     skoXC1 *
     (-4020854/54267 +
      skoXC1 *
      (-3783203000000147513/45222500000000000 +
       skoXC1 *
       (1915879999999852487/3617800000000000 +
        skoXC1 *
        (451009749999852487/1356675000000000 +
         skoXC1 *
         (-1354011400000147513/1808900000000000 +
          skoXC1 *
          (-3491141000000147513/9044500000000000 +
           skoXC1 * -6982282000000147513/271335000000000000))))))),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (7 <
     skoXC1 *
     (-15856156/126623 +
      skoXC1 *
      (-34059942000000344197/90445000000000000 +
       skoXC1 *
       (32364891999996902227/45222500000000000 +
        skoXC1 *
        (7465317999998967409/3617800000000000 +
         skoXC1 *
         (-898479400000344197/904450000000000 +
          skoXC1 *
          (-42354320000003097773/18089000000000000 +
           skoXC1 *
           (-28224154000001032591/45222500000000000 +
            skoXC1 * -17357363000000344197/633115000000000000))))))) |
     0 <
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000)))))))),
    inference(decision, [],
              [refute_0_143, refute_0_11, refute_0_62, refute_0_89,
               refute_0_128])).

cnf(refute_0_145, plain,
    (7 <
     skoXC1 *
     (-15856156/126623 +
      skoXC1 *
      (-34059942000000344197/90445000000000000 +
       skoXC1 *
       (32364891999996902227/45222500000000000 +
        skoXC1 *
        (7465317999998967409/3617800000000000 +
         skoXC1 *
         (-898479400000344197/904450000000000 +
          skoXC1 *
          (-42354320000003097773/18089000000000000 +
           skoXC1 *
           (-28224154000001032591/45222500000000000 +
            skoXC1 * -17357363000000344197/633115000000000000)))))))),
    inference(resolve, [], [refute_0_132, refute_0_144])).

cnf(refute_0_146, plain,
    (8 <
     skoXC1 *
     (-14132605/72356 +
      skoXC1 *
      (-78660802000000344197/79139375000000000 +
       skoXC1 *
       (2427533499997590621/22611250000000000 +
        skoXC1 *
        (66987997999992771863/11305625000000000 +
         skoXC1 *
         (4959814299997590621/1808900000000000 +
          skoXC1 *
          (-16992276000002409379/2261125000000000 +
           skoXC1 *
           (-125310319250007228137/22611250000000000 +
            skoXC1 *
            (-10473423000000344197/11305625000000000 +
             skoXC1 * -18279319250000344197/633115000000000000)))))))) |
     0 <
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000))))))))),
    inference(decision, [],
              [refute_0_145, refute_0_143, refute_0_11, refute_0_62,
               refute_0_89, refute_0_128])).

cnf(refute_0_147, plain,
    (8 <
     skoXC1 *
     (-14132605/72356 +
      skoXC1 *
      (-78660802000000344197/79139375000000000 +
       skoXC1 *
       (2427533499997590621/22611250000000000 +
        skoXC1 *
        (66987997999992771863/11305625000000000 +
         skoXC1 *
         (4959814299997590621/1808900000000000 +
          skoXC1 *
          (-16992276000002409379/2261125000000000 +
           skoXC1 *
           (-125310319250007228137/22611250000000000 +
            skoXC1 *
            (-10473423000000344197/11305625000000000 +
             skoXC1 * -18279319250000344197/633115000000000000))))))))),
    inference(resolve, [], [refute_0_113, refute_0_146])).

cnf(refute_0_148, plain,
    (0 <
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000))))))))),
    inference(decision, [],
              [refute_0_147, refute_0_145, refute_0_143, refute_0_11,
               refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_149, plain,
    ((8 +
      skoXC1 *
      (196 +
       skoXC1 *
       (1176 +
        skoXC1 *
        (2450 +
         skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1))))))) /
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) <
     28245729/10391023),
    inference(resolve, [], [refute_0_103, refute_0_148])).

cnf(refute_0_150, plain,
    (8 +
     skoXC1 *
     (196 +
      skoXC1 *
      (1176 +
       skoXC1 *
       (2450 +
        skoXC1 * (1960 + skoXC1 * (588 + skoXC1 * (56 + skoXC1)))))) <
     28245729/10391023 *
     (skoXC1 *
      (1/4 +
       skoXC1 *
       (292999999999993/4375000000000 +
        skoXC1 *
        (1175999999999951/1250000000000 +
         skoXC1 *
         (1812999999999853/625000000000 +
          skoXC1 *
          (122499999999951/100000000000 +
           skoXC1 *
           (-343000000000049/125000000000 +
            skoXC1 *
            (-2548000000000147/1250000000000 +
             skoXC1 *
             (-213000000000007/625000000000 +
              skoXC1 * -371750000000007/35000000000000))))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_149])).

cnf(refute_0_151, plain,
    (8 <
     skoXC1 *
     (-8118316303/41564092 +
      skoXC1 *
      (-45185814738000197720103/45460725625000000000 +
       skoXC1 *
       (1394469366498615959279/12988778750000000000 +
        skoXC1 *
        (38480503501995847877837/6494389375000000000 +
         skoXC1 *
         (2849109650098615959279/1039102300000000000 +
          skoXC1 *
          (-9761022208001384040721/1298877875000000000 +
           skoXC1 *
           (-71983106270754152122163/12988778750000000000 +
            skoXC1 *
            (-6016340277000197720103/6494389375000000000 +
             skoXC1 *
             -10500349755750197720103/363685805000000000000)))))))) |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     28245729/10391023 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     28245729/10391023 <=
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_153, plain,
    (1 < 0 | 19/2 < 1 |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     exp(1) <=
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000))))))))),
    inference(subst, [], [refute_0_100])).

cnf(refute_0_154, plain,
    (1 < 0 | 19/2 < 1 |
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     (-17297280 +
      1 *
      (-8648640 +
       1 *
       (-1995840 +
        1 *
        (-277200 + 1 * (-25200 + 1 * (-1512 + 1 * (-56 + 1 * -1))))))) /
     (-17297280 +
      1 *
      (8648640 +
       1 *
       (-1995840 +
        1 * (277200 + 1 * (-25200 + 1 * (1512 + 1 * (-56 + 1))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_153, refute_0_115])).

cnf(refute_0_155, plain,
    ((7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     28245729/10391023 |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    ((7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) /
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) <
     28245729/10391023),
    inference(resolve, [], [refute_0_155, refute_0_129])).

cnf(refute_0_157, plain,
    (7 +
     skoXC1 *
     (126 +
      skoXC1 *
      (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))) <
     28245729/10391023 *
     (skoXC1 *
      (2/7 +
       skoXC1 *
       (272999999999993/5000000000000 +
        skoXC1 *
        (1301999999999937/2500000000000 +
         skoXC1 *
         (174999999999979/200000000000 +
          skoXC1 *
          (-17500000000007/50000000000 +
           skoXC1 *
           (-861000000000063/1000000000000 +
            skoXC1 *
            (-574000000000021/2500000000000 +
             skoXC1 * -353000000000007/35000000000000)))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(resolve, [], [refute_0_152, refute_0_156])).

cnf(refute_0_158, plain,
    (7 <
     skoXC1 *
     (-9108390828/72737161 +
      skoXC1 *
      (-19565351358000197720103/51955115000000000000 +
       skoXC1 *
       (18591648907998220519073/25977557500000000000 +
        skoXC1 *
        (4288368125999406839691/2078204600000000000 +
         skoXC1 *
         (-516121405800197720103/519551150000000000 +
          skoXC1 *
          (-24329963692001779480927/10391023000000000000 +
           skoXC1 *
           (-16213048446000593160309/25977557500000000000 +
            skoXC1 * -9970742337000197720103/363685805000000000000))))))) |
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000))))))) <= 0),
    inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    (7 <
     skoXC1 *
     (-9108390828/72737161 +
      skoXC1 *
      (-19565351358000197720103/51955115000000000000 +
       skoXC1 *
       (18591648907998220519073/25977557500000000000 +
        skoXC1 *
        (4288368125999406839691/2078204600000000000 +
         skoXC1 *
         (-516121405800197720103/519551150000000000 +
          skoXC1 *
          (-24329963692001779480927/10391023000000000000 +
           skoXC1 *
           (-16213048446000593160309/25977557500000000000 +
            skoXC1 * -9970742337000197720103/363685805000000000000))))))) |
     0 <
     skoXC1 *
     (2/7 +
      skoXC1 *
      (272999999999993/5000000000000 +
       skoXC1 *
       (1301999999999937/2500000000000 +
        skoXC1 *
        (174999999999979/200000000000 +
         skoXC1 *
         (-17500000000007/50000000000 +
          skoXC1 *
          (-861000000000063/1000000000000 +
           skoXC1 *
           (-574000000000021/2500000000000 +
            skoXC1 * -353000000000007/35000000000000)))))))),
    inference(decision, [],
              [refute_0_147, refute_0_145, refute_0_143, refute_0_11,
               refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_160, plain,
    (7 <
     skoXC1 *
     (-9108390828/72737161 +
      skoXC1 *
      (-19565351358000197720103/51955115000000000000 +
       skoXC1 *
       (18591648907998220519073/25977557500000000000 +
        skoXC1 *
        (4288368125999406839691/2078204600000000000 +
         skoXC1 *
         (-516121405800197720103/519551150000000000 +
          skoXC1 *
          (-24329963692001779480927/10391023000000000000 +
           skoXC1 *
           (-16213048446000593160309/25977557500000000000 +
            skoXC1 * -9970742337000197720103/363685805000000000000)))))))),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (skoXC1 *
     (-8118316303/41564092 +
      skoXC1 *
      (-45185814738000197720103/45460725625000000000 +
       skoXC1 *
       (1394469366498615959279/12988778750000000000 +
        skoXC1 *
        (38480503501995847877837/6494389375000000000 +
         skoXC1 *
         (2849109650098615959279/1039102300000000000 +
          skoXC1 *
          (-9761022208001384040721/1298877875000000000 +
           skoXC1 *
           (-71983106270754152122163/12988778750000000000 +
            skoXC1 *
            (-6016340277000197720103/6494389375000000000 +
             skoXC1 *
             -10500349755750197720103/363685805000000000000)))))))) <= 8 |
     skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(decision, [],
              [refute_0_160, refute_0_147, refute_0_145, refute_0_143,
               refute_0_11, refute_0_62, refute_0_89, refute_0_128])).

cnf(refute_0_162, plain,
    (skoXC1 *
     (1/4 +
      skoXC1 *
      (292999999999993/4375000000000 +
       skoXC1 *
       (1175999999999951/1250000000000 +
        skoXC1 *
        (1812999999999853/625000000000 +
         skoXC1 *
         (122499999999951/100000000000 +
          skoXC1 *
          (-343000000000049/125000000000 +
           skoXC1 *
           (-2548000000000147/1250000000000 +
            skoXC1 *
            (-213000000000007/625000000000 +
             skoXC1 * -371750000000007/35000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_161, refute_0_151])).

cnf(refute_0_163, plain, ($false),
    inference(resolve, [], [refute_0_162, refute_0_148])).
% SZS output end CNFRefutation for log-fun-ineq-fneg2.tptp
