% SZS output start CNFRefutation for log-upper-21-sqrt-weak2.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/12 * (3 * X ^ 3 + 131 * X ^ 2 + 239 * X + 47) * (X - 1) /
            (4 * X ^ 3 + 18 * X ^ 2 + 12 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (6 * X ^ 4 + 481 * X ^ 3 + 1881 * X ^ 2 + 1281 * X + 131) *
            (X - 1) / (5 * X ^ 4 + 40 * X ^ 3 + 60 * X ^ 2 + 20 * X + 1),
            Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/70 *
            (353 + 1921 * X ^ 5 + 8389 * X + 38524 * X ^ 2 +
             20149 * X ^ 4 + 50774 * X ^ 3 + 10 * X ^ 6) * (X - 1) /
            (1 + 126 * X ^ 5 + 42 * X + 315 * X ^ 2 + 525 * X ^ 4 +
             700 * X ^ 3 + 7 * X ^ 6), Y) | lgen(R, ln(X), Y))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_small, axiom,
    (X < 0 | 1 <= X |
     ~ lgen(R,
            1/32 *
            (65536 * X ^ 4 + 114688 * X ^ 3 + 17920 * X ^ 2 + 448 * X +
             1) / ((16 * X + 1) * (256 * X ^ 2 + 96 * X + 1)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

fof(log_upper_bound_problem_21, conjecture,
    (! [X] :
       ((1/1000 < X & X < 8) => ln(1 + 1 / X) < 1 / sqrt(X ^ 2 + X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((1/1000 < X & X < 8) => ln(1 + 1 / X) < 1 / sqrt(X ^ 2 + X))),
    inference(strip, [], [log_upper_bound_problem_21])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((1/1000 < X & X < 8) => ln(1 + 1 / X) < 1 / sqrt(X ^ 2 + X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (1/1000 < X & X < 8 & 1 / sqrt(X ^ 2 + X) <= ln(1 + 1 / X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 8 & 1/1000 < skoXC1 &
     1 / sqrt(skoXC1 ^ 2 + skoXC1) <= ln(1 + 1 / skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1 / sqrt(skoXC1 ^ 2 + skoXC1) <= ln(1 + 1 / skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (1/1000 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 8),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (-1 < 1 / skoXC1 | 1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    (1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-952 +
       skoXC1 *
       (-33404 +
        skoXC1 *
        (-518952 +
         skoXC1 *
         (-4265170 +
          skoXC1 *
          (-20906760 +
           skoXC1 *
           (-65907220 +
            skoXC1 *
            (-139277376 +
             skoXC1 *
             (-200190144 +
              skoXC1 *
              (-193597440 +
               skoXC1 *
               (-120671232 +
                skoXC1 * (-43794432 + skoXC1 * -7028736)))))))))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) |
     0 <=
     -1/7 +
     skoXC1 *
     (-283/10 +
      skoXC1 *
      (-2136/5 +
       skoXC1 *
       (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716))))) |
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) *
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     skoXC1 *
     (-7 +
      skoXC1 *
      (-952 +
       skoXC1 *
       (-33404 +
        skoXC1 *
        (-518952 +
         skoXC1 *
         (-4265170 +
          skoXC1 *
          (-20906760 +
           skoXC1 *
           (-65907220 +
            skoXC1 *
            (-139277376 +
             skoXC1 *
             (-200190144 +
              skoXC1 *
              (-193597440 +
               skoXC1 *
               (-120671232 +
                skoXC1 * (-43794432 + skoXC1 * -7028736))))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_2, plain,
    ((1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) *
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) |
     1 +
     skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_3, plain,
    (1/7 +
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) <
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) *
     sqrt(skoXC1 * (1 + skoXC1)) | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_4, plain,
    ((1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) <
     1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) <=
     1/7 +
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_5, plain,
    (1 / sqrt(skoXC1 ^ 2 + skoXC1) <= ln(1 + 1 / skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_6, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) <= ln(1 + 1 / skoXC1)),
    inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (ln(X_000104) < X_000105 | ~ lgen(1, ln(X_000104), X_000105)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_8, plain,
    (~ lgen(1,
            1/70 *
            (353 + 1921 * X_000104 ^ 5 + 8389 * X_000104 +
             38524 * X_000104 ^ 2 + 20149 * X_000104 ^ 4 +
             50774 * X_000104 ^ 3 + 10 * X_000104 ^ 6) * (X_000104 - 1) /
            (1 + 126 * X_000104 ^ 5 + 42 * X_000104 + 315 * X_000104 ^ 2 +
             525 * X_000104 ^ 4 + 700 * X_000104 ^ 3 + 7 * X_000104 ^ 6),
            X_000105) | X_000104 <= 0 | lgen(1, ln(X_000104), X_000105)),
    inference(subst, [], [ln_upper_bound_cf13])).

cnf(refute_0_9, plain,
    (ln(X_000104) < X_000105 |
     ~ lgen(1,
            1/70 *
            (353 + 1921 * X_000104 ^ 5 + 8389 * X_000104 +
             38524 * X_000104 ^ 2 + 20149 * X_000104 ^ 4 +
             50774 * X_000104 ^ 3 + 10 * X_000104 ^ 6) * (X_000104 - 1) /
            (1 + 126 * X_000104 ^ 5 + 42 * X_000104 + 315 * X_000104 ^ 2 +
             525 * X_000104 ^ 4 + 700 * X_000104 ^ 3 + 7 * X_000104 ^ 6),
            X_000105) | X_000104 <= 0),
    inference(resolve, [], [refute_0_8, refute_0_7])).

cnf(refute_0_10, plain,
    (ln(X_000104) < X_000105 | X_000104 <= 0 |
     X_000105 <=
     (-353/70 +
      X_000104 *
      (-574/5 +
       X_000104 *
       (-861/2 +
        X_000104 *
        (-175 +
         X_000104 *
         (875/2 +
          X_000104 * (1302/5 + X_000104 * (273/10 + X_000104 * 1/7))))))) /
     (1 +
      X_000104 *
      (42 +
       X_000104 *
       (315 +
        X_000104 *
        (700 + X_000104 * (525 + X_000104 * (126 + X_000104 * 7))))))),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (ln(1 + 1 / skoXC1) < 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-353/70 +
      (1 + 1 / skoXC1) *
      (-574/5 +
       (1 + 1 / skoXC1) *
       (-861/2 +
        (1 + 1 / skoXC1) *
        (-175 +
         (1 + 1 / skoXC1) *
         (875/2 +
          (1 + 1 / skoXC1) *
          (1302/5 +
           (1 + 1 / skoXC1) * (273/10 + (1 + 1 / skoXC1) * 1/7))))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (42 +
       (1 + 1 / skoXC1) *
       (315 +
        (1 + 1 / skoXC1) *
        (700 +
         (1 + 1 / skoXC1) *
         (525 + (1 + 1 / skoXC1) * (126 + (1 + 1 / skoXC1) * 7))))))),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-353/70 +
      (1 + 1 / skoXC1) *
      (-574/5 +
       (1 + 1 / skoXC1) *
       (-861/2 +
        (1 + 1 / skoXC1) *
        (-175 +
         (1 + 1 / skoXC1) *
         (875/2 +
          (1 + 1 / skoXC1) *
          (1302/5 +
           (1 + 1 / skoXC1) * (273/10 + (1 + 1 / skoXC1) * 1/7))))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (42 +
       (1 + 1 / skoXC1) *
       (315 +
        (1 + 1 / skoXC1) *
        (700 +
         (1 + 1 / skoXC1) *
         (525 + (1 + 1 / skoXC1) * (126 + (1 + 1 / skoXC1) * 7))))))),
    inference(resolve, [], [refute_0_6, refute_0_11])).

cnf(refute_0_13, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) <
     (-1/40 +
      skoXC1 *
      (-673/240 +
       skoXC1 *
       (-2261/30 +
        skoXC1 *
        (-9443/15 +
         skoXC1 *
         (-20037/8 +
          skoXC1 *
          (-1387103/240 +
           skoXC1 *
           (-507391/60 +
            skoXC1 *
            (-489241/60 +
             skoXC1 *
             (-637523/120 +
              skoXC1 *
              (-113585/48 +
               skoXC1 *
               (-658 + skoXC1 * (-721/8 + skoXC1 * -63/4)))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     -1/40 +
     skoXC1 *
     (-673/240 +
      skoXC1 *
      (-2261/30 +
       skoXC1 *
       (-9443/15 +
        skoXC1 *
        (-20037/8 +
         skoXC1 *
         (-1387103/240 +
          skoXC1 *
          (-507391/60 +
           skoXC1 *
           (-489241/60 +
            skoXC1 *
            (-637523/120 +
             skoXC1 *
             (-113585/48 +
              skoXC1 *
              (-658 + skoXC1 * (-721/8 + skoXC1 * -63/4))))))))))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_15, plain,
    ((1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) |
     0 <=
     -1/5 +
     skoXC1 *
     (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126))) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_16, plain,
    (sqrt(X_000062) < X_000063 | ~ lgen(1, sqrt(X_000062), X_000063)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain,
    (X_000062 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1) /
            ((X_000062 + 1) * (X_000062 ^ 2 + 6 * X_000062 + 1)),
            X_000063) | lgen(1, sqrt(X_000062), X_000063)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_18, plain,
    (sqrt(X_000062) < X_000063 | X_000062 < 0 |
     ~ lgen(1,
            1/8 *
            (X_000062 ^ 4 + 28 * X_000062 ^ 3 + 70 * X_000062 ^ 2 +
             28 * X_000062 + 1) /
            ((X_000062 + 1) * (X_000062 ^ 2 + 6 * X_000062 + 1)),
            X_000063)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (sqrt(X_000062) < X_000063 | X_000062 < 0 |
     X_000063 <=
     (1/8 +
      X_000062 *
      (7/2 + X_000062 * (35/4 + X_000062 * (7/2 + X_000062 * 1/8)))) /
     (1 + X_000062 * (7 + X_000062 * (7 + X_000062)))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1/5 +
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126))) <
     skoXC1 *
     (5 + skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 *
     (5 +
      skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) *
     sqrt(skoXC1 * (1 + skoXC1)) | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_22, plain,
    ((1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) <
     1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) <=
     1/5 +
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126))) |
     skoXC1 *
     (5 +
      skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_23, plain,
    (ln(X_000083) < X_000084 | ~ lgen(1, ln(X_000083), X_000084)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_24, plain,
    (~ lgen(1,
            1/30 *
            (6 * X_000083 ^ 4 + 481 * X_000083 ^ 3 + 1881 * X_000083 ^ 2 +
             1281 * X_000083 + 131) * (X_000083 - 1) /
            (5 * X_000083 ^ 4 + 40 * X_000083 ^ 3 + 60 * X_000083 ^ 2 +
             20 * X_000083 + 1), X_000084) | X_000083 <= 0 |
     lgen(1, ln(X_000083), X_000084)),
    inference(subst, [], [ln_upper_bound_cf9])).

cnf(refute_0_25, plain,
    (ln(X_000083) < X_000084 |
     ~ lgen(1,
            1/30 *
            (6 * X_000083 ^ 4 + 481 * X_000083 ^ 3 + 1881 * X_000083 ^ 2 +
             1281 * X_000083 + 131) * (X_000083 - 1) /
            (5 * X_000083 ^ 4 + 40 * X_000083 ^ 3 + 60 * X_000083 ^ 2 +
             20 * X_000083 + 1), X_000084) | X_000083 <= 0),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (ln(X_000083) < X_000084 | X_000083 <= 0 |
     X_000084 <=
     (-131/30 +
      X_000083 *
      (-115/3 +
       X_000083 *
       (-20 + X_000083 * (140/3 + X_000083 * (95/6 + X_000083 * 1/5))))) /
     (1 +
      X_000083 * (20 + X_000083 * (60 + X_000083 * (40 + X_000083 * 5))))),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (ln(1 + 1 / skoXC1) < 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-131/30 +
      (1 + 1 / skoXC1) *
      (-115/3 +
       (1 + 1 / skoXC1) *
       (-20 +
        (1 + 1 / skoXC1) *
        (140/3 + (1 + 1 / skoXC1) * (95/6 + (1 + 1 / skoXC1) * 1/5))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (20 +
       (1 + 1 / skoXC1) *
       (60 + (1 + 1 / skoXC1) * (40 + (1 + 1 / skoXC1) * 5))))),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-131/30 +
      (1 + 1 / skoXC1) *
      (-115/3 +
       (1 + 1 / skoXC1) *
       (-20 +
        (1 + 1 / skoXC1) *
        (140/3 + (1 + 1 / skoXC1) * (95/6 + (1 + 1 / skoXC1) * 1/5))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (20 +
       (1 + 1 / skoXC1) *
       (60 + (1 + 1 / skoXC1) * (40 + (1 + 1 / skoXC1) * 5))))),
    inference(resolve, [], [refute_0_6, refute_0_27])).

cnf(refute_0_29, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) <
     (-1/32 +
      skoXC1 *
      (-7/3 +
       skoXC1 *
       (-1181/24 +
        skoXC1 *
        (-7231/24 +
         skoXC1 *
         (-28575/32 +
          skoXC1 *
          (-37201/24 +
           skoXC1 *
           (-10127/6 +
            skoXC1 *
            (-3586/3 +
             skoXC1 *
             (-54863/96 +
              skoXC1 *
              (-4115/24 + skoXC1 * (-365/16 + skoXC1 * -35/8))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     -1/32 +
     skoXC1 *
     (-7/3 +
      skoXC1 *
      (-1181/24 +
       skoXC1 *
       (-7231/24 +
        skoXC1 *
        (-28575/32 +
         skoXC1 *
         (-37201/24 +
          skoXC1 *
          (-10127/6 +
           skoXC1 *
           (-3586/3 +
            skoXC1 *
            (-54863/96 +
             skoXC1 *
             (-4115/24 + skoXC1 * (-365/16 + skoXC1 * -35/8)))))))))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_31, plain,
    ((1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) |
     0 <= -1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35)) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_32, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_33, plain,
    (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) <
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) *
     sqrt(skoXC1 * (1 + skoXC1)) | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_34, plain,
    ((1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) /
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) <
     1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) <=
     1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) |
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_35, plain,
    (ln(X_000076) < X_000077 | ~ lgen(1, ln(X_000076), X_000077)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_36, plain,
    (~ lgen(1,
            1/12 *
            (3 * X_000076 ^ 3 + 131 * X_000076 ^ 2 + 239 * X_000076 + 47) *
            (X_000076 - 1) /
            (4 * X_000076 ^ 3 + 18 * X_000076 ^ 2 + 12 * X_000076 + 1),
            X_000077) | X_000076 <= 0 | lgen(1, ln(X_000076), X_000077)),
    inference(subst, [], [ln_upper_bound_cf7])).

cnf(refute_0_37, plain,
    (ln(X_000076) < X_000077 |
     ~ lgen(1,
            1/12 *
            (3 * X_000076 ^ 3 + 131 * X_000076 ^ 2 + 239 * X_000076 + 47) *
            (X_000076 - 1) /
            (4 * X_000076 ^ 3 + 18 * X_000076 ^ 2 + 12 * X_000076 + 1),
            X_000077) | X_000076 <= 0),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (ln(X_000076) < X_000077 | X_000076 <= 0 |
     X_000077 <=
     (-47/12 +
      X_000076 *
      (-16 + X_000076 * (9 + X_000076 * (32/3 + X_000076 * 1/4)))) /
     (1 + X_000076 * (12 + X_000076 * (18 + X_000076 * 4)))),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (ln(1 + 1 / skoXC1) < 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-47/12 +
      (1 + 1 / skoXC1) *
      (-16 +
       (1 + 1 / skoXC1) *
       (9 + (1 + 1 / skoXC1) * (32/3 + (1 + 1 / skoXC1) * 1/4)))) /
     (1 +
      (1 + 1 / skoXC1) *
      (12 + (1 + 1 / skoXC1) * (18 + (1 + 1 / skoXC1) * 4)))),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-47/12 +
      (1 + 1 / skoXC1) *
      (-16 +
       (1 + 1 / skoXC1) *
       (9 + (1 + 1 / skoXC1) * (32/3 + (1 + 1 / skoXC1) * 1/4)))) /
     (1 +
      (1 + 1 / skoXC1) *
      (12 + (1 + 1 / skoXC1) * (18 + (1 + 1 / skoXC1) * 4)))),
    inference(resolve, [], [refute_0_6, refute_0_39])).

cnf(refute_0_41, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) /
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) <
     (-2048/3 +
      skoXC1 *
      (-46592/3 +
       skoXC1 *
       (-140848/3 +
        skoXC1 *
        (-65226 +
         skoXC1 *
         (-4759553/96 +
          skoXC1 *
          (-1494553/96 +
           skoXC1 *
           (-75839/12 +
            skoXC1 *
            (-24989/48 +
             skoXC1 *
             (-13705/96 + skoXC1 * (-47/32 + skoXC1 * -5/16)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     -2048/3 +
     skoXC1 *
     (-46592/3 +
      skoXC1 *
      (-140848/3 +
       skoXC1 *
       (-65226 +
        skoXC1 *
        (-4759553/96 +
         skoXC1 *
         (-1494553/96 +
          skoXC1 *
          (-75839/12 +
           skoXC1 *
           (-24989/48 +
            skoXC1 *
            (-13705/96 + skoXC1 * (-47/32 + skoXC1 * -5/16))))))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_43, plain,
    ((2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_44, plain,
    (1/3 + skoXC1 * (7 + skoXC1 * 10) <
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <=
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) * sqrt(skoXC1 * (1 + skoXC1)) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_45, plain,
    ((1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <
     1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <=
     1/3 + skoXC1 * (7 + skoXC1 * 10) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_46, plain,
    (ln(X_000055) < X_000056 | ~ lgen(1, ln(X_000055), X_000056)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_47, plain,
    (~ lgen(1,
            1/3 * (X_000055 ^ 2 + 19 * X_000055 + 10) * (X_000055 - 1) /
            (3 * X_000055 ^ 2 + 6 * X_000055 + 1), X_000056) |
     X_000055 <= 0 | lgen(1, ln(X_000055), X_000056)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_48, plain,
    (ln(X_000055) < X_000056 |
     ~ lgen(1,
            1/3 * (X_000055 ^ 2 + 19 * X_000055 + 10) * (X_000055 - 1) /
            (3 * X_000055 ^ 2 + 6 * X_000055 + 1), X_000056) |
     X_000055 <= 0), inference(resolve, [], [refute_0_47, refute_0_46])).

cnf(refute_0_49, plain,
    (ln(X_000055) < X_000056 | X_000055 <= 0 |
     X_000056 <=
     (-10/3 + X_000055 * (-3 + X_000055 * (6 + X_000055 * 1/3))) /
     (1 + X_000055 * (6 + X_000055 * 3))),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (ln(1 + 1 / skoXC1) < 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-10/3 +
      (1 + 1 / skoXC1) *
      (-3 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 1/3))) /
     (1 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 3))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-10/3 +
      (1 + 1 / skoXC1) *
      (-3 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 1/3))) /
     (1 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 3))),
    inference(resolve, [], [refute_0_6, refute_0_50])).

cnf(refute_0_52, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) | skoXC1 * skoXC1 = 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     0 <= -1/2 + skoXC1 * -3 |
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) * (-1/2 + skoXC1 * -3) <=
     skoXC1 * (-2 + skoXC1 * -3)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_54, plain,
    (sqrt(X_000041) < X_000042 | ~ lgen(1, sqrt(X_000041), X_000042)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_55, plain,
    (X_000041 < 0 | ~ lgen(1, (X_000041 + 1) / 2, X_000042) |
     lgen(1, sqrt(X_000041), X_000042)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_56, plain,
    (sqrt(X_000041) < X_000042 | X_000041 < 0 |
     ~ lgen(1, (X_000041 + 1) / 2, X_000042)),
    inference(resolve, [], [refute_0_55, refute_0_54])).

cnf(refute_0_57, plain,
    (sqrt(X_000041) < X_000042 | X_000041 < 0 |
     X_000042 <= 1/2 + X_000041 * 1/2),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     1/2 + skoXC1 * (1 + skoXC1) * 1/2),
    inference(subst, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (1/2 + skoXC1 * 3 <
     skoXC1 * (2 + skoXC1 * 3) / sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (2 + skoXC1 * 3) <=
     (1/2 + skoXC1 * 3) * sqrt(skoXC1 * (1 + skoXC1)) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_60, plain,
    ((1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3)) <
     1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 / sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (2 + skoXC1 * 3)) <=
     1/2 + skoXC1 * 3 | skoXC1 * (2 + skoXC1 * 3) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_61, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_62, plain,
    (~ lgen(1, 1/2 * (X_000048 + 5) * (X_000048 - 1) / (2 * X_000048 + 1),
            X_000049) | X_000048 <= 0 | lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_63, plain,
    (ln(X_000048) < X_000049 |
     ~ lgen(1, 1/2 * (X_000048 + 5) * (X_000048 - 1) / (2 * X_000048 + 1),
            X_000049) | X_000048 <= 0),
    inference(resolve, [], [refute_0_62, refute_0_61])).

cnf(refute_0_64, plain,
    (ln(X_000048) < X_000049 | X_000048 <= 0 |
     X_000049 <=
     (-5/2 + X_000048 * (2 + X_000048 * 1/2)) / (1 + X_000048 * 2)),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (ln(1 + 1 / skoXC1) < 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-5/2 + (1 + 1 / skoXC1) * (2 + (1 + 1 / skoXC1) * 1/2)) /
     (1 + (1 + 1 / skoXC1) * 2)), inference(subst, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (1 + 1 / skoXC1 <= 0 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (-5/2 + (1 + 1 / skoXC1) * (2 + (1 + 1 / skoXC1) * 1/2)) /
     (1 + (1 + 1 / skoXC1) * 2)),
    inference(resolve, [], [refute_0_6, refute_0_65])).

cnf(refute_0_67, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain, (1/1000 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_69, plain, (skoXC1 < 8),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_70, plain, (skoXC1 != 0),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3))),
    inference(resolve, [], [refute_0_67, refute_0_70])).

cnf(refute_0_72, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (2 + skoXC1 * 3)) <=
     1/2 + skoXC1 * 3 | skoXC1 * (2 + skoXC1 * 3) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_71, refute_0_60])).

cnf(refute_0_73, plain,
    (skoXC1 * (2 + skoXC1 * 3) <= 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (2 + skoXC1 * 3) / sqrt(skoXC1 * (1 + skoXC1)) <=
     1/2 + skoXC1 * 3), inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain, (0 < skoXC1 * (2 + skoXC1 * 3)),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_75, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (2 + skoXC1 * 3) / sqrt(skoXC1 * (1 + skoXC1)) <=
     1/2 + skoXC1 * 3),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * (2 + skoXC1 * 3) <=
     (1/2 + skoXC1 * 3) * sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_75, refute_0_59])).

cnf(refute_0_77, plain,
    (skoXC1 * (1 + skoXC1) <= 0 |
     sqrt(skoXC1 * (1 + skoXC1)) * (-1/2 + skoXC1 * -3) <=
     skoXC1 * (-2 + skoXC1 * -3) | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain, (0 < skoXC1 * (1 + skoXC1)),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_79, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) * (-1/2 + skoXC1 * -3) <=
     skoXC1 * (-2 + skoXC1 * -3) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain,
    (0 <= -1/2 + skoXC1 * -3 | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     sqrt(skoXC1 * (1 + skoXC1))), inference(split, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     sqrt(skoXC1 * (1 + skoXC1)) | skoXC1 <= -1/6),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain, (-1/6 < skoXC1),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_83, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     1/2 + skoXC1 * (1 + skoXC1) * 1/2),
    inference(resolve, [], [refute_0_83, refute_0_58])).

cnf(refute_0_85, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain, (skoXC1 * (-1 + skoXC1 * -1) <= 0),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_87, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_86, refute_0_85])).

cnf(refute_0_88, plain,
    (0 <= -1/2 + skoXC1 * -3 |
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) * (-1/2 + skoXC1 * -3) <=
     skoXC1 * (-2 + skoXC1 * -3) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_87, refute_0_53])).

cnf(refute_0_89, plain,
    (skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 |
     1 / skoXC1 <= -1 | skoXC1 <= -1/6),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (-1/6 < skoXC1 |
     skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_91, plain,
    (skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_91, refute_0_0])).

cnf(refute_0_93, plain,
    (skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (-1 < skoXC1 | skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 |
     skoXC1 <= 0), inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_95, plain,
    (skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (0 < skoXC1 | skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4),
    inference(decision, [], [refute_0_68, refute_0_69])).

cnf(refute_0_97, plain,
    (skoXC1 * (1/4 + skoXC1 * (5/4 + skoXC1 * -3/2)) <= 1/4),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain, (skoXC1 * skoXC1 != 0 | skoXC1 = 0),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_99, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) | skoXC1 = 0),
    inference(resolve, [], [refute_0_52, refute_0_98])).

cnf(refute_0_100, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)))),
    inference(resolve, [], [refute_0_99, refute_0_70])).

cnf(refute_0_101, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <=
     1/3 + skoXC1 * (7 + skoXC1 * 10) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_100, refute_0_45])).

cnf(refute_0_102, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     sqrt(skoXC1 * (1 + skoXC1)) <= 1/3 + skoXC1 * (7 + skoXC1 * 10)),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain, (0 < skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_104, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     sqrt(skoXC1 * (1 + skoXC1)) <= 1/3 + skoXC1 * (7 + skoXC1 * 10)),
    inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <=
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) * sqrt(skoXC1 * (1 + skoXC1)) |
     1 / skoXC1 <= -1 | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_104, refute_0_44])).

cnf(refute_0_106, plain,
    (skoXC1 * (1 + skoXC1) <= 0 |
     sqrt(skoXC1 * (1 + skoXC1)) * (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) * (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_106, refute_0_78])).

cnf(refute_0_108, plain,
    (0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <= sqrt(skoXC1 * (1 + skoXC1))),
    inference(split, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 * (7 + skoXC1 * 10) <= -1/3 | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <= sqrt(skoXC1 * (1 + skoXC1))),
    inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain, (-1/3 < skoXC1 * (7 + skoXC1 * 10)),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_111, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <= sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    (sqrt(X_000069) < X_000070 | ~ lgen(1, sqrt(X_000069), X_000070)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_113, plain,
    (X_000069 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000069 + 17920 * X_000069 ^ 2 +
             448 * X_000069 ^ 3 + X_000069 ^ 4) /
            ((X_000069 + 16) * (256 + 96 * X_000069 + X_000069 ^ 2)),
            X_000070) | lgen(1, sqrt(X_000069), X_000070)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_114, plain,
    (sqrt(X_000069) < X_000070 | X_000069 < 1 |
     ~ lgen(1,
            1/32 *
            (65536 + 114688 * X_000069 + 17920 * X_000069 ^ 2 +
             448 * X_000069 ^ 3 + X_000069 ^ 4) /
            ((X_000069 + 16) * (256 + 96 * X_000069 + X_000069 ^ 2)),
            X_000070)),
    inference(resolve, [], [refute_0_113, refute_0_112])).

cnf(refute_0_115, plain,
    (sqrt(X_000069) < X_000070 | X_000069 < 1 |
     X_000070 <=
     (2048 +
      X_000069 *
      (3584 + X_000069 * (560 + X_000069 * (14 + X_000069 * 1/32)))) /
     (4096 + X_000069 * (1792 + X_000069 * (112 + X_000069)))),
    inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_115])).

cnf(refute_0_117, plain,
    (skoXC1 * (1 + skoXC1) < 1 | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_111, refute_0_116])).

cnf(refute_0_118, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_118, refute_0_43])).

cnf(refute_0_120, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (7 + skoXC1 * 10) <= -1/3 | 1 / skoXC1 <= -1 |
     (-2048/3 +
      skoXC1 *
      (-46592/3 +
       skoXC1 *
       (-140848/3 +
        skoXC1 *
        (-65226 +
         skoXC1 *
         (-4759553/96 +
          skoXC1 *
          (-1494553/96 +
           skoXC1 *
           (-75839/12 +
            skoXC1 *
            (-24989/48 +
             skoXC1 *
             (-13705/96 + skoXC1 * (-47/32 + skoXC1 * -5/16)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) <
     (-1/24 +
      skoXC1 *
      (-49/24 +
       skoXC1 *
       (-179/6 +
        skoXC1 *
        (-511/4 +
         skoXC1 *
         (-6623/24 +
          skoXC1 *
          (-8383/24 +
           skoXC1 *
           (-1633/6 +
            skoXC1 *
            (-1679/12 +
             skoXC1 *
             (-1105/24 + skoXC1 * (-47/8 + skoXC1 * -5/4)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     -1/24 +
     skoXC1 *
     (-49/24 +
      skoXC1 *
      (-179/6 +
       skoXC1 *
       (-511/4 +
        skoXC1 *
        (-6623/24 +
         skoXC1 *
         (-8383/24 +
          skoXC1 *
          (-1633/6 +
           skoXC1 *
           (-1679/12 +
            skoXC1 *
            (-1105/24 + skoXC1 * (-47/8 + skoXC1 * -5/4))))))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_122, plain,
    ((1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_123, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_124, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_111, refute_0_123])).

cnf(refute_0_125, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(resolve, [], [refute_0_86, refute_0_125])).

cnf(refute_0_127, plain,
    (0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_126, refute_0_122])).

cnf(refute_0_128, plain,
    (skoXC1 * (7 + skoXC1 * 10) <= -1/3 | 1 / skoXC1 <= -1 |
     (-1/24 +
      skoXC1 *
      (-49/24 +
       skoXC1 *
       (-179/6 +
        skoXC1 *
        (-511/4 +
         skoXC1 *
         (-6623/24 +
          skoXC1 *
          (-8383/24 +
           skoXC1 *
           (-1633/6 +
            skoXC1 *
            (-1679/12 +
             skoXC1 *
             (-1105/24 + skoXC1 * (-47/8 + skoXC1 * -5/4)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (1 / skoXC1 <= -1 |
     (-1/24 +
      skoXC1 *
      (-49/24 +
       skoXC1 *
       (-179/6 +
        skoXC1 *
        (-511/4 +
         skoXC1 *
         (-6623/24 +
          skoXC1 *
          (-8383/24 +
           skoXC1 *
           (-1633/6 +
            skoXC1 *
            (-1679/12 +
             skoXC1 *
             (-1105/24 + skoXC1 * (-47/8 + skoXC1 * -5/4)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(resolve, [], [refute_0_128, refute_0_110])).

cnf(refute_0_130, plain,
    (-1/24 +
     skoXC1 *
     (-49/24 +
      skoXC1 *
      (-179/6 +
       skoXC1 *
       (-511/4 +
        skoXC1 *
        (-6623/24 +
         skoXC1 *
         (-8383/24 +
          skoXC1 *
          (-1633/6 +
           skoXC1 *
           (-1679/12 +
            skoXC1 *
            (-1105/24 + skoXC1 * (-47/8 + skoXC1 * -5/4))))))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_129, refute_0_121])).

cnf(refute_0_131, plain,
    (skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 * (-2 + skoXC1 * -3) <
     (-1/16 +
      skoXC1 *
      (-17/8 +
       skoXC1 *
       (-133/8 +
        skoXC1 *
        (-189/4 +
         skoXC1 *
         (-1163/16 +
          skoXC1 *
          (-509/8 +
           skoXC1 *
           (-281/8 +
            skoXC1 * (-13 + skoXC1 * (-25/16 + skoXC1 * -3/8))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     -1/16 +
     skoXC1 *
     (-17/8 +
      skoXC1 *
      (-133/8 +
       skoXC1 *
       (-189/4 +
        skoXC1 *
        (-1163/16 +
         skoXC1 *
         (-509/8 +
          skoXC1 *
          (-281/8 +
           skoXC1 * (-13 + skoXC1 * (-25/16 + skoXC1 * -3/8)))))))) <=
     skoXC1 * (-2 + skoXC1 * -3) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_133, plain,
    ((1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     0 <= -1/2 + skoXC1 * -3 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/2 + skoXC1 * -3) <= skoXC1 * (-2 + skoXC1 * -3)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_134, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_135, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_83, refute_0_134])).

cnf(refute_0_136, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(resolve, [], [refute_0_86, refute_0_136])).

cnf(refute_0_138, plain,
    (0 <= -1/2 + skoXC1 * -3 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/2 + skoXC1 * -3) <= skoXC1 * (-2 + skoXC1 * -3) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_137, refute_0_133])).

cnf(refute_0_139, plain,
    (1 / skoXC1 <= -1 |
     (-1/16 +
      skoXC1 *
      (-17/8 +
       skoXC1 *
       (-133/8 +
        skoXC1 *
        (-189/4 +
         skoXC1 *
         (-1163/16 +
          skoXC1 *
          (-509/8 +
           skoXC1 *
           (-281/8 +
            skoXC1 * (-13 + skoXC1 * (-25/16 + skoXC1 * -3/8))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-2 + skoXC1 * -3) | skoXC1 <= -1/6),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (1 / skoXC1 <= -1 |
     (-1/16 +
      skoXC1 *
      (-17/8 +
       skoXC1 *
       (-133/8 +
        skoXC1 *
        (-189/4 +
         skoXC1 *
         (-1163/16 +
          skoXC1 *
          (-509/8 +
           skoXC1 *
           (-281/8 +
            skoXC1 * (-13 + skoXC1 * (-25/16 + skoXC1 * -3/8))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-2 + skoXC1 * -3)),
    inference(resolve, [], [refute_0_139, refute_0_82])).

cnf(refute_0_141, plain,
    (-1/16 +
     skoXC1 *
     (-17/8 +
      skoXC1 *
      (-133/8 +
       skoXC1 *
       (-189/4 +
        skoXC1 *
        (-1163/16 +
         skoXC1 *
         (-509/8 +
          skoXC1 *
          (-281/8 +
           skoXC1 * (-13 + skoXC1 * (-25/16 + skoXC1 * -3/8)))))))) <=
     skoXC1 * (-2 + skoXC1 * -3) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_140, refute_0_132])).

cnf(refute_0_142, plain,
    (skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_144, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     1/2 + skoXC1 * (1 + skoXC1) * 1/2),
    inference(subst, [], [refute_0_57])).

cnf(refute_0_145, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     1/2 + skoXC1 * (1 + skoXC1) * 1/2),
    inference(resolve, [], [refute_0_111, refute_0_144])).

cnf(refute_0_146, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)),
    inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) /
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_86, refute_0_146])).

cnf(refute_0_148, plain,
    (0 <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (-1/3 + skoXC1 * (-7 + skoXC1 * -10)) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_147, refute_0_143])).

cnf(refute_0_149, plain,
    (skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | skoXC1 * (7 + skoXC1 * 10) <= -1/3 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (-1/3 < skoXC1 * (7 + skoXC1 * 10) |
     skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_151, plain,
    (skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_151, refute_0_0])).

cnf(refute_0_153, plain,
    (skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (-1 < skoXC1 |
     skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | skoXC1 <= 0),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_155, plain,
    (skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (0 < skoXC1 |
     skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6),
    inference(decision, [], [refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_157, plain,
    (skoXC1 * (-2/3 + skoXC1 * (10/3 + skoXC1 * (3/2 + skoXC1 * -5))) <=
     1/6), inference(resolve, [], [refute_0_155, refute_0_156])).

cnf(refute_0_158, plain,
    (-1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) |
     skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16),
    inference(decision, [],
              [refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_159, plain,
    (skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_142, refute_0_158])).

cnf(refute_0_160, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_159, refute_0_0])).

cnf(refute_0_161, plain,
    (skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_163, plain,
    (skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16),
    inference(decision, [],
              [refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_165, plain,
    (skoXC1 *
     (-1/8 +
      skoXC1 *
      (3/8 +
       skoXC1 *
       (7/4 +
        skoXC1 *
        (-11/16 +
         skoXC1 *
         (11/8 +
          skoXC1 *
          (7/8 + skoXC1 * (-2 + skoXC1 * (23/16 + skoXC1 * -3/8)))))))) <=
     1/16), inference(resolve, [], [refute_0_163, refute_0_164])).

cnf(refute_0_166, plain,
    (-1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) |
     skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24),
    inference(decision, [],
              [refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_167, plain,
    (skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_131, refute_0_166])).

cnf(refute_0_168, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_167, refute_0_0])).

cnf(refute_0_169, plain,
    (skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_171, plain,
    (skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (0 < skoXC1 |
     skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24),
    inference(decision, [],
              [refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_173, plain,
    (skoXC1 *
     (23/24 +
      skoXC1 *
      (19/6 +
       skoXC1 *
       (33/4 +
        skoXC1 *
        (169/24 +
         skoXC1 *
         (17/24 +
          skoXC1 *
          (41/6 +
           skoXC1 *
           (-11/12 +
            skoXC1 * (-97/24 + skoXC1 * (33/8 + skoXC1 * -5/4))))))))) <=
     1/24), inference(resolve, [], [refute_0_171, refute_0_172])).

cnf(refute_0_174, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | -1/3 < skoXC1 * (7 + skoXC1 * 10)),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_175, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     (-2048/3 +
      skoXC1 *
      (-46592/3 +
       skoXC1 *
       (-140848/3 +
        skoXC1 *
        (-65226 +
         skoXC1 *
         (-4759553/96 +
          skoXC1 *
          (-1494553/96 +
           skoXC1 *
           (-75839/12 +
            skoXC1 *
            (-24989/48 +
             skoXC1 *
             (-13705/96 + skoXC1 * (-47/32 + skoXC1 * -5/16)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10))),
    inference(resolve, [], [refute_0_120, refute_0_174])).

cnf(refute_0_176, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -2048/3 +
     skoXC1 *
     (-46592/3 +
      skoXC1 *
      (-140848/3 +
       skoXC1 *
       (-65226 +
        skoXC1 *
        (-4759553/96 +
         skoXC1 *
         (-1494553/96 +
          skoXC1 *
          (-75839/12 +
           skoXC1 *
           (-24989/48 +
            skoXC1 *
            (-13705/96 + skoXC1 * (-47/32 + skoXC1 * -5/16))))))))) <=
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_175, refute_0_42])).

cnf(refute_0_177, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (-9728/3 +
      skoXC1 *
      (22736/3 +
       skoXC1 *
       (2950 +
        skoXC1 *
        (-781025/96 +
         skoXC1 *
         (625607/96 +
          skoXC1 *
          (-32171/12 +
           skoXC1 *
           (32083/48 +
            skoXC1 *
            (-9673/96 + skoXC1 * (273/32 + skoXC1 * -5/16))))))))) <=
     2048/3 |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 * (-2 + skoXC1 * -3) <
     (-1024 +
      skoXC1 *
      (-7936 +
       skoXC1 *
       (-12824 +
        skoXC1 *
        (-12999 +
         skoXC1 *
         (-236993/64 +
          skoXC1 *
          (-58469/32 +
           skoXC1 *
           (-4271/32 +
            skoXC1 * (-341/8 + skoXC1 * (-25/64 + skoXC1 * -3/32))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     -1024 +
     skoXC1 *
     (-7936 +
      skoXC1 *
      (-12824 +
       skoXC1 *
       (-12999 +
        skoXC1 *
        (-236993/64 +
         skoXC1 *
         (-58469/32 +
          skoXC1 *
          (-4271/32 +
           skoXC1 * (-341/8 + skoXC1 * (-25/64 + skoXC1 * -3/32)))))))) <=
     skoXC1 * (-2 + skoXC1 * -3) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_179, plain,
    ((2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     0 <= -1/2 + skoXC1 * -3 |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (-1/2 + skoXC1 * -3) <= skoXC1 * (-2 + skoXC1 * -3)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_180, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1))))),
    inference(subst, [], [refute_0_115])).

cnf(refute_0_181, plain,
    (skoXC1 * (1 + skoXC1) < 1 | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_83, refute_0_180])).

cnf(refute_0_182, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-2 + skoXC1 * -3) / (-1/2 + skoXC1 * -3) <=
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 0 <= -1/2 + skoXC1 * -3 |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (-1/2 + skoXC1 * -3) <= skoXC1 * (-2 + skoXC1 * -3) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_182, refute_0_179])).

cnf(refute_0_184, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     (-1024 +
      skoXC1 *
      (-7936 +
       skoXC1 *
       (-12824 +
        skoXC1 *
        (-12999 +
         skoXC1 *
         (-236993/64 +
          skoXC1 *
          (-58469/32 +
           skoXC1 *
           (-4271/32 +
            skoXC1 * (-341/8 + skoXC1 * (-25/64 + skoXC1 * -3/32))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-2 + skoXC1 * -3) | skoXC1 <= -1/6),
    inference(arithmetic, [], [refute_0_183])).

cnf(refute_0_185, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | -1/6 < skoXC1),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_186, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     (-1024 +
      skoXC1 *
      (-7936 +
       skoXC1 *
       (-12824 +
        skoXC1 *
        (-12999 +
         skoXC1 *
         (-236993/64 +
          skoXC1 *
          (-58469/32 +
           skoXC1 *
           (-4271/32 +
            skoXC1 * (-341/8 + skoXC1 * (-25/64 + skoXC1 * -3/32))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-2 + skoXC1 * -3)),
    inference(resolve, [], [refute_0_184, refute_0_185])).

cnf(refute_0_187, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -1024 +
     skoXC1 *
     (-7936 +
      skoXC1 *
      (-12824 +
       skoXC1 *
       (-12999 +
        skoXC1 *
        (-236993/64 +
         skoXC1 *
         (-58469/32 +
          skoXC1 *
          (-4271/32 +
           skoXC1 * (-341/8 + skoXC1 * (-25/64 + skoXC1 * -3/32)))))))) <=
     skoXC1 * (-2 + skoXC1 * -3) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_186, refute_0_178])).

cnf(refute_0_188, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -4096 <
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_190, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_188, refute_0_189])).

cnf(refute_0_191, plain,
    (skoXC1 * (-1 + skoXC1 * -1) <= -1 |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_192, plain,
    (skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_191, refute_0_190])).

cnf(refute_0_193, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_192, refute_0_0])).

cnf(refute_0_194, plain,
    (skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_196, plain,
    (skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (0 < skoXC1 |
     skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024),
    inference(decision, [],
              [refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_198, plain,
    (skoXC1 *
     (256 +
      skoXC1 *
      (3048 +
       skoXC1 *
       (-3815 +
        skoXC1 *
        (157375/64 +
         skoXC1 *
         (-29509/32 +
          skoXC1 *
          (6961/32 +
           skoXC1 * (-253/8 + skoXC1 * (167/64 + skoXC1 * -3/32)))))))) <=
     1024), inference(resolve, [], [refute_0_196, refute_0_197])).

cnf(refute_0_199, plain,
    (2048/3 <
     skoXC1 *
     (-9728/3 +
      skoXC1 *
      (22736/3 +
       skoXC1 *
       (2950 +
        skoXC1 *
        (-781025/96 +
         skoXC1 *
         (625607/96 +
          skoXC1 *
          (-32171/12 +
           skoXC1 *
           (32083/48 +
            skoXC1 *
            (-9673/96 + skoXC1 * (273/32 + skoXC1 * -5/16))))))))) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096),
    inference(decision, [],
              [refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_200, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_177, refute_0_199])).

cnf(refute_0_201, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -4096 <
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_202, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_200, refute_0_201])).

cnf(refute_0_203, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_202, refute_0_0])).

cnf(refute_0_204, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_203])).

cnf(refute_0_205, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_206, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_204, refute_0_205])).

cnf(refute_0_207, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_208, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (skoXC1 * (skoXC1 * skoXC1) != 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_210, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) /
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) |
     skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_41, refute_0_209])).

cnf(refute_0_211, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) /
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) |
     skoXC1 = 0), inference(resolve, [], [refute_0_210, refute_0_98])).

cnf(refute_0_212, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) /
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))))),
    inference(resolve, [], [refute_0_211, refute_0_70])).

cnf(refute_0_213, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))) <=
     1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) |
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_212, refute_0_34])).

cnf(refute_0_214, plain,
    (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) <= 0 |
     1 / skoXC1 <= -1 |
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))),
    inference(arithmetic, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (0 < skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35)))),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_216, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (skoXC1 * (4 + skoXC1 * (30 + skoXC1 * (60 + skoXC1 * 35))) <=
     (1/4 + skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))) *
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_216, refute_0_33])).

cnf(refute_0_218, plain,
    (skoXC1 * (1 + skoXC1) <= 0 |
     sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_218, refute_0_78])).

cnf(refute_0_220, plain,
    (0 <= -1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35)) |
     1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     sqrt(skoXC1 * (1 + skoXC1))), inference(split, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) <= -1/4 |
     1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (-1/4 < skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35))),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_223, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_221, refute_0_222])).

cnf(refute_0_224, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_223, refute_0_32])).

cnf(refute_0_225, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(resolve, [], [refute_0_86, refute_0_225])).

cnf(refute_0_227, plain,
    (0 <= -1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35)) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_226, refute_0_31])).

cnf(refute_0_228, plain,
    (skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) <= -1/4 |
     1 / skoXC1 <= -1 |
     (-1/32 +
      skoXC1 *
      (-7/3 +
       skoXC1 *
       (-1181/24 +
        skoXC1 *
        (-7231/24 +
         skoXC1 *
         (-28575/32 +
          skoXC1 *
          (-37201/24 +
           skoXC1 *
           (-10127/6 +
            skoXC1 *
            (-3586/3 +
             skoXC1 *
             (-54863/96 +
              skoXC1 *
              (-4115/24 + skoXC1 * (-365/16 + skoXC1 * -35/8))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35)))),
    inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (1 / skoXC1 <= -1 |
     (-1/32 +
      skoXC1 *
      (-7/3 +
       skoXC1 *
       (-1181/24 +
        skoXC1 *
        (-7231/24 +
         skoXC1 *
         (-28575/32 +
          skoXC1 *
          (-37201/24 +
           skoXC1 *
           (-10127/6 +
            skoXC1 *
            (-3586/3 +
             skoXC1 *
             (-54863/96 +
              skoXC1 *
              (-4115/24 + skoXC1 * (-365/16 + skoXC1 * -35/8))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35)))),
    inference(resolve, [], [refute_0_228, refute_0_222])).

cnf(refute_0_230, plain,
    (-1/32 +
     skoXC1 *
     (-7/3 +
      skoXC1 *
      (-1181/24 +
       skoXC1 *
       (-7231/24 +
        skoXC1 *
        (-28575/32 +
         skoXC1 *
         (-37201/24 +
          skoXC1 *
          (-10127/6 +
           skoXC1 *
           (-3586/3 +
            skoXC1 *
            (-54863/96 +
             skoXC1 *
             (-4115/24 + skoXC1 * (-365/16 + skoXC1 * -35/8)))))))))) <=
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_229, refute_0_30])).

cnf(refute_0_231, plain,
    (skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_230])).

cnf(refute_0_232, plain,
    (-1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) |
     skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_233, plain,
    (skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_231, refute_0_232])).

cnf(refute_0_234, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_233, refute_0_0])).

cnf(refute_0_235, plain,
    (skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_234])).

cnf(refute_0_236, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_237, plain,
    (skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_235, refute_0_236])).

cnf(refute_0_238, plain,
    (0 < skoXC1 |
     skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32),
    inference(decision, [],
              [refute_0_208, refute_0_198, refute_0_173, refute_0_165,
               refute_0_157, refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_239, plain,
    (skoXC1 *
     (5/3 +
      skoXC1 *
      (211/24 +
       skoXC1 *
       (593/24 +
        skoXC1 *
        (1345/32 +
         skoXC1 *
         (599/24 +
          skoXC1 *
          (85/6 +
           skoXC1 *
           (71/3 +
            skoXC1 *
            (-1103/96 +
             skoXC1 *
             (-155/24 + skoXC1 * (195/16 + skoXC1 * -35/8)))))))))) <=
     1/32), inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) != 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_241, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_29, refute_0_240])).

cnf(refute_0_242, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) |
     skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_241, refute_0_209])).

cnf(refute_0_243, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) |
     skoXC1 = 0), inference(resolve, [], [refute_0_242, refute_0_98])).

cnf(refute_0_244, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) /
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))))),
    inference(resolve, [], [refute_0_243, refute_0_70])).

cnf(refute_0_245, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (5 +
       skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))) <=
     1/5 +
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126))) |
     skoXC1 *
     (5 +
      skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_244, refute_0_22])).

cnf(refute_0_246, plain,
    (skoXC1 *
     (5 +
      skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) <=
     0 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (5 + skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/5 +
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (0 <
     skoXC1 *
     (5 + skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126))))),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_248, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (5 + skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/5 +
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))),
    inference(resolve, [], [refute_0_246, refute_0_247])).

cnf(refute_0_249, plain,
    (skoXC1 *
     (5 +
      skoXC1 * (60 + skoXC1 * (210 + skoXC1 * (280 + skoXC1 * 126)))) <=
     (1/5 +
      skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))) *
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_248, refute_0_21])).

cnf(refute_0_250, plain,
    (skoXC1 * (1 + skoXC1) <= 0 |
     sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_250, refute_0_78])).

cnf(refute_0_252, plain,
    (0 <=
     -1/5 +
     skoXC1 *
     (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126))) |
     1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     sqrt(skoXC1 * (1 + skoXC1))), inference(split, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126))) <=
     -1/5 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(arithmetic, [], [refute_0_252])).

cnf(refute_0_254, plain,
    (-1/5 <
     skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126)))),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_255, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_253, refute_0_254])).

cnf(refute_0_256, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1))))),
    inference(resolve, [], [refute_0_255, refute_0_20])).

cnf(refute_0_257, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) /
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))))),
    inference(resolve, [], [refute_0_86, refute_0_257])).

cnf(refute_0_259, plain,
    (0 <=
     -1/5 +
     skoXC1 *
     (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126))) |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (-1/5 +
      skoXC1 *
      (-101/6 + skoXC1 * (-112 + skoXC1 * (-217 + skoXC1 * -126)))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_258, refute_0_15])).

cnf(refute_0_260, plain,
    (skoXC1 * (101/6 + skoXC1 * (112 + skoXC1 * (217 + skoXC1 * 126))) <=
     -1/5 | 1 / skoXC1 <= -1 |
     (-1/40 +
      skoXC1 *
      (-673/240 +
       skoXC1 *
       (-2261/30 +
        skoXC1 *
        (-9443/15 +
         skoXC1 *
         (-20037/8 +
          skoXC1 *
          (-1387103/240 +
           skoXC1 *
           (-507391/60 +
            skoXC1 *
            (-489241/60 +
             skoXC1 *
             (-637523/120 +
              skoXC1 *
              (-113585/48 +
               skoXC1 *
               (-658 + skoXC1 * (-721/8 + skoXC1 * -63/4)))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126))))),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (1 / skoXC1 <= -1 |
     (-1/40 +
      skoXC1 *
      (-673/240 +
       skoXC1 *
       (-2261/30 +
        skoXC1 *
        (-9443/15 +
         skoXC1 *
         (-20037/8 +
          skoXC1 *
          (-1387103/240 +
           skoXC1 *
           (-507391/60 +
            skoXC1 *
            (-489241/60 +
             skoXC1 *
             (-637523/120 +
              skoXC1 *
              (-113585/48 +
               skoXC1 *
               (-658 + skoXC1 * (-721/8 + skoXC1 * -63/4)))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126))))),
    inference(resolve, [], [refute_0_260, refute_0_254])).

cnf(refute_0_262, plain,
    (-1/40 +
     skoXC1 *
     (-673/240 +
      skoXC1 *
      (-2261/30 +
       skoXC1 *
       (-9443/15 +
        skoXC1 *
        (-20037/8 +
         skoXC1 *
         (-1387103/240 +
          skoXC1 *
          (-507391/60 +
           skoXC1 *
           (-489241/60 +
            skoXC1 *
            (-637523/120 +
             skoXC1 *
             (-113585/48 +
              skoXC1 *
              (-658 + skoXC1 * (-721/8 + skoXC1 * -63/4))))))))))) <=
     skoXC1 *
     (-5 +
      skoXC1 * (-60 + skoXC1 * (-210 + skoXC1 * (-280 + skoXC1 * -126)))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_261, refute_0_14])).

cnf(refute_0_263, plain,
    (skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (-1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) |
     skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_265, plain,
    (skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_263, refute_0_264])).

cnf(refute_0_266, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_265, refute_0_0])).

cnf(refute_0_267, plain,
    (skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_266])).

cnf(refute_0_268, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_269, plain,
    (skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_267, refute_0_268])).

cnf(refute_0_270, plain,
    (0 < skoXC1 |
     skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40),
    inference(decision, [],
              [refute_0_239, refute_0_208, refute_0_198, refute_0_173,
               refute_0_165, refute_0_157, refute_0_97, refute_0_68,
               refute_0_69])).

cnf(refute_0_271, plain,
    (skoXC1 *
     (527/240 +
      skoXC1 *
      (589/30 +
       skoXC1 *
       (1057/15 +
        skoXC1 *
        (1283/8 +
         skoXC1 *
         (47137/240 +
          skoXC1 *
          (6629/60 +
           skoXC1 *
           (5699/60 +
            skoXC1 *
            (8077/120 +
             skoXC1 *
             (-2705/48 +
              skoXC1 * (skoXC1 * (287/8 + skoXC1 * -63/4))))))))))) <=
     1/40), inference(resolve, [], [refute_0_269, refute_0_270])).

cnf(refute_0_272, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_273, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_13, refute_0_272])).

cnf(refute_0_274, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_275, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_273, refute_0_274])).

cnf(refute_0_276, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_275, refute_0_240])).

cnf(refute_0_277, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_276, refute_0_209])).

cnf(refute_0_278, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) |
     skoXC1 = 0), inference(resolve, [], [refute_0_277, refute_0_98])).

cnf(refute_0_279, plain,
    (1 / skoXC1 <= -1 |
     1 / sqrt(skoXC1 * (1 + skoXC1)) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))))),
    inference(resolve, [], [refute_0_278, refute_0_70])).

cnf(refute_0_280, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (168 +
        skoXC1 *
        (1260 +
         skoXC1 *
         (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))) <=
     1/7 +
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_279, refute_0_4])).

cnf(refute_0_281, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) <=
     0 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/7 +
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))),
    inference(arithmetic, [], [refute_0_280])).

cnf(refute_0_282, plain,
    (0 <
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716))))))),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_283, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     1/7 +
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))),
    inference(resolve, [], [refute_0_281, refute_0_282])).

cnf(refute_0_284, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (168 +
       skoXC1 *
       (1260 +
        skoXC1 *
        (4200 + skoXC1 * (6930 + skoXC1 * (5544 + skoXC1 * 1716)))))) <=
     (1/7 +
      skoXC1 *
      (283/10 +
       skoXC1 *
       (2136/5 +
        skoXC1 *
        (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))) *
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_283, refute_0_3])).

cnf(refute_0_285, plain,
    (skoXC1 * (1 + skoXC1) <= 0 |
     sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) *
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_285, refute_0_78])).

cnf(refute_0_287, plain,
    (0 <=
     -1/7 +
     skoXC1 *
     (-283/10 +
      skoXC1 *
      (-2136/5 +
       skoXC1 *
       (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716))))) |
     1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     sqrt(skoXC1 * (1 + skoXC1))), inference(split, [], [refute_0_286])).

cnf(refute_0_288, plain,
    (skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) <=
     -1/7 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(arithmetic, [], [refute_0_287])).

cnf(refute_0_289, plain,
    (-1/7 <
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716)))))),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_290, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_288, refute_0_289])).

cnf(refute_0_291, plain,
    (sqrt(X_000142) < X_000143 | ~ lgen(1, sqrt(X_000142), X_000143)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_292, plain,
    (X_000142 < 0 |
     ~ lgen(1,
            1/32 *
            (65536 * X_000142 ^ 4 + 114688 * X_000142 ^ 3 +
             17920 * X_000142 ^ 2 + 448 * X_000142 + 1) /
            ((16 * X_000142 + 1) *
             (256 * X_000142 ^ 2 + 96 * X_000142 + 1)), X_000143) |
     1 <= X_000142 | lgen(1, sqrt(X_000142), X_000143)),
    inference(subst, [], [sqrt_upper_bound_2_small])).

cnf(refute_0_293, plain,
    (sqrt(X_000142) < X_000143 | X_000142 < 0 |
     ~ lgen(1,
            1/32 *
            (65536 * X_000142 ^ 4 + 114688 * X_000142 ^ 3 +
             17920 * X_000142 ^ 2 + 448 * X_000142 + 1) /
            ((16 * X_000142 + 1) *
             (256 * X_000142 ^ 2 + 96 * X_000142 + 1)), X_000143) |
     1 <= X_000142), inference(resolve, [], [refute_0_292, refute_0_291])).

cnf(refute_0_294, plain,
    (sqrt(X_000142) < X_000143 | X_000142 < 0 |
     X_000143 <=
     (1/32 +
      X_000142 *
      (14 + X_000142 * (560 + X_000142 * (3584 + X_000142 * 2048)))) /
     (1 + X_000142 * (112 + X_000142 * (1792 + X_000142 * 4096))) |
     1 <= X_000142), inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) |
     1 <= skoXC1 * (1 + skoXC1) |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     (1/32 +
      skoXC1 * (1 + skoXC1) *
      (14 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (3584 + skoXC1 * (1 + skoXC1) * 2048)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (112 +
       skoXC1 * (1 + skoXC1) * (1792 + skoXC1 * (1 + skoXC1) * 4096)))),
    inference(subst, [], [refute_0_294])).

cnf(refute_0_296, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 <= skoXC1 * (1 + skoXC1) |
     1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     (1/32 +
      skoXC1 * (1 + skoXC1) *
      (14 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (3584 + skoXC1 * (1 + skoXC1) * 2048)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (112 +
       skoXC1 * (1 + skoXC1) * (1792 + skoXC1 * (1 + skoXC1) * 4096)))),
    inference(resolve, [], [refute_0_290, refute_0_295])).

cnf(refute_0_297, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 * (-1 + skoXC1 * -1) <= -1 |
     1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(arithmetic, [], [refute_0_296])).

cnf(refute_0_298, plain,
    (skoXC1 * (-1 + skoXC1 * -1) <= -1 | skoXC1 * (-1 + skoXC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_299, plain,
    (skoXC1 * (-1 + skoXC1 * -1) <= -1 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(resolve, [], [refute_0_298, refute_0_297])).

cnf(refute_0_300, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(resolve, [], [refute_0_299, refute_0_208])).

cnf(refute_0_301, plain,
    (skoXC1 *
     (-7 +
      skoXC1 *
      (-168 +
       skoXC1 *
       (-1260 +
        skoXC1 *
        (-4200 + skoXC1 * (-6930 + skoXC1 * (-5544 + skoXC1 * -1716)))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) *
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) |
     1 +
     skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_300, refute_0_2])).

cnf(refute_0_302, plain,
    (skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     -1 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-952 +
       skoXC1 *
       (-33404 +
        skoXC1 *
        (-518952 +
         skoXC1 *
         (-4265170 +
          skoXC1 *
          (-20906760 +
           skoXC1 *
           (-65907220 +
            skoXC1 *
            (-139277376 +
             skoXC1 *
             (-200190144 +
              skoXC1 *
              (-193597440 +
               skoXC1 *
               (-120671232 +
                skoXC1 * (-43794432 + skoXC1 * -7028736)))))))))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))),
    inference(arithmetic, [], [refute_0_301])).

cnf(refute_0_303, plain,
    (-1 <
     skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_304, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-952 +
       skoXC1 *
       (-33404 +
        skoXC1 *
        (-518952 +
         skoXC1 *
         (-4265170 +
          skoXC1 *
          (-20906760 +
           skoXC1 *
           (-65907220 +
            skoXC1 *
            (-139277376 +
             skoXC1 *
             (-200190144 +
              skoXC1 *
              (-193597440 +
               skoXC1 *
               (-120671232 +
                skoXC1 * (-43794432 + skoXC1 * -7028736)))))))))))) /
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))),
    inference(resolve, [], [refute_0_302, refute_0_303])).

cnf(refute_0_305, plain,
    (0 <=
     -1/7 +
     skoXC1 *
     (-283/10 +
      skoXC1 *
      (-2136/5 +
       skoXC1 *
       (-2154 + skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716))))) |
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) *
     (-1/7 +
      skoXC1 *
      (-283/10 +
       skoXC1 *
       (-2136/5 +
        skoXC1 *
        (-2154 +
         skoXC1 * (-4730 + skoXC1 * (-4686 + skoXC1 * -1716)))))) <=
     skoXC1 *
     (-7 +
      skoXC1 *
      (-952 +
       skoXC1 *
       (-33404 +
        skoXC1 *
        (-518952 +
         skoXC1 *
         (-4265170 +
          skoXC1 *
          (-20906760 +
           skoXC1 *
           (-65907220 +
            skoXC1 *
            (-139277376 +
             skoXC1 *
             (-200190144 +
              skoXC1 *
              (-193597440 +
               skoXC1 *
               (-120671232 +
                skoXC1 * (-43794432 + skoXC1 * -7028736)))))))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_304, refute_0_1])).

cnf(refute_0_306, plain,
    (skoXC1 *
     (1317/320 +
      skoXC1 *
      (9209/20 +
       skoXC1 *
       (167035/16 +
        skoXC1 *
        (12141205/112 +
         skoXC1 *
         (320356107/560 +
          skoXC1 *
          (489409689/280 +
           skoXC1 *
           (126799656/35 +
            skoXC1 *
            (36670392/7 +
             skoXC1 *
             (10810656/5 +
              skoXC1 *
              (-58474688/5 +
               skoXC1 *
               (-29372416 +
                skoXC1 *
                (-31516672 +
                 skoXC1 * (-16625664 + skoXC1 * -3514368))))))))))))) <=
     1/224 |
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) <=
     -1/7 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) <
     skoXC1 *
     (-4 +
      skoXC1 *
      (-478 +
       skoXC1 *
       (-11036 +
        skoXC1 *
        (-94595 +
         skoXC1 *
         (-404880 +
          skoXC1 *
          (-998992 +
           skoXC1 *
           (-1498624 +
            skoXC1 *
            (-1352960 + skoXC1 * (-675840 + skoXC1 * -143360))))))))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) |
     0 <= -1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35)) |
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (-478 +
       skoXC1 *
       (-11036 +
        skoXC1 *
        (-94595 +
         skoXC1 *
         (-404880 +
          skoXC1 *
          (-998992 +
           skoXC1 *
           (-1498624 +
            skoXC1 *
            (-1352960 + skoXC1 * (-675840 + skoXC1 * -143360)))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_308, plain,
    ((1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) *
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) |
     1 +
     skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_309, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) |
     1 <= skoXC1 * (1 + skoXC1) |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/32 +
      skoXC1 * (1 + skoXC1) *
      (14 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (3584 + skoXC1 * (1 + skoXC1) * 2048)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (112 +
       skoXC1 * (1 + skoXC1) * (1792 + skoXC1 * (1 + skoXC1) * 4096)))),
    inference(subst, [], [refute_0_294])).

cnf(refute_0_310, plain,
    (skoXC1 * (1 + skoXC1) < 0 | 1 <= skoXC1 * (1 + skoXC1) |
     1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/32 +
      skoXC1 * (1 + skoXC1) *
      (14 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (3584 + skoXC1 * (1 + skoXC1) * 2048)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (112 +
       skoXC1 * (1 + skoXC1) * (1792 + skoXC1 * (1 + skoXC1) * 4096)))),
    inference(resolve, [], [refute_0_223, refute_0_309])).

cnf(refute_0_311, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 * (-1 + skoXC1 * -1) <= -1 |
     1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(arithmetic, [], [refute_0_310])).

cnf(refute_0_312, plain,
    (skoXC1 * (-1 + skoXC1 * -1) <= -1 | 1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(resolve, [], [refute_0_298, refute_0_311])).

cnf(refute_0_313, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) /
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))))),
    inference(resolve, [], [refute_0_312, refute_0_208])).

cnf(refute_0_314, plain,
    (skoXC1 * (-4 + skoXC1 * (-30 + skoXC1 * (-60 + skoXC1 * -35))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) *
     (1 +
      skoXC1 *
      (112 +
       skoXC1 *
       (1904 +
        skoXC1 *
        (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096)))))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048))))))) |
     1 +
     skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_313, refute_0_308])).

cnf(refute_0_315, plain,
    (skoXC1 *
     (112 +
      skoXC1 *
      (1904 +
       skoXC1 *
       (7680 + skoXC1 * (14080 + skoXC1 * (12288 + skoXC1 * 4096))))) <=
     -1 | 1 / skoXC1 <= -1 |
     skoXC1 *
     (-4 +
      skoXC1 *
      (-478 +
       skoXC1 *
       (-11036 +
        skoXC1 *
        (-94595 +
         skoXC1 *
         (-404880 +
          skoXC1 *
          (-998992 +
           skoXC1 *
           (-1498624 +
            skoXC1 *
            (-1352960 + skoXC1 * (-675840 + skoXC1 * -143360))))))))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))),
    inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (1 / skoXC1 <= -1 |
     skoXC1 *
     (-4 +
      skoXC1 *
      (-478 +
       skoXC1 *
       (-11036 +
        skoXC1 *
        (-94595 +
         skoXC1 *
         (-404880 +
          skoXC1 *
          (-998992 +
           skoXC1 *
           (-1498624 +
            skoXC1 *
            (-1352960 + skoXC1 * (-675840 + skoXC1 * -143360))))))))) /
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     1/32 +
     skoXC1 *
     (14 +
      skoXC1 *
      (574 +
       skoXC1 *
       (4704 +
        skoXC1 *
        (13360 +
         skoXC1 *
         (18944 + skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))),
    inference(resolve, [], [refute_0_315, refute_0_303])).

cnf(refute_0_317, plain,
    (0 <= -1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35)) |
     (1/32 +
      skoXC1 *
      (14 +
       skoXC1 *
       (574 +
        skoXC1 *
        (4704 +
         skoXC1 *
         (13360 +
          skoXC1 *
          (18944 +
           skoXC1 * (15872 + skoXC1 * (8192 + skoXC1 * 2048)))))))) *
     (-1/4 + skoXC1 * (-35/3 + skoXC1 * (-85/2 + skoXC1 * -35))) <=
     skoXC1 *
     (-4 +
      skoXC1 *
      (-478 +
       skoXC1 *
       (-11036 +
        skoXC1 *
        (-94595 +
         skoXC1 *
         (-404880 +
          skoXC1 *
          (-998992 +
           skoXC1 *
           (-1498624 +
            skoXC1 *
            (-1352960 + skoXC1 * (-675840 + skoXC1 * -143360))))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_316, refute_0_307])).

cnf(refute_0_318, plain,
    (skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) <= -1/4 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_317])).

cnf(refute_0_319, plain,
    (-1/4 < skoXC1 * (35/3 + skoXC1 * (85/2 + skoXC1 * 35)) |
     skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_320, plain,
    (skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_318, refute_0_319])).

cnf(refute_0_321, plain,
    (1 <= -1 * skoXC1 |
     skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_320, refute_0_0])).

cnf(refute_0_322, plain,
    (skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_321])).

cnf(refute_0_323, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_324, plain,
    (skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_322, refute_0_323])).

cnf(refute_0_325, plain,
    (0 < skoXC1 |
     skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128),
    inference(decision, [],
              [refute_0_271, refute_0_239, refute_0_208, refute_0_198,
               refute_0_173, refute_0_165, refute_0_157, refute_0_97,
               refute_0_68, refute_0_69])).

cnf(refute_0_326, plain,
    (skoXC1 *
     (13/96 +
      skoXC1 *
      (32609/192 +
       skoXC1 *
       (246455/96 +
        skoXC1 *
        (11490 +
         skoXC1 *
         (72802/3 +
          skoXC1 *
          (124712/3 +
           skoXC1 *
           (116048/3 +
            skoXC1 *
            (-242176/3 +
             skoXC1 *
             (-755200/3 + skoXC1 * (-230400 + skoXC1 * -71680)))))))))) <=
     1/128), inference(resolve, [], [refute_0_324, refute_0_325])).

cnf(refute_0_327, plain,
    (1/224 <
     skoXC1 *
     (1317/320 +
      skoXC1 *
      (9209/20 +
       skoXC1 *
       (167035/16 +
        skoXC1 *
        (12141205/112 +
         skoXC1 *
         (320356107/560 +
          skoXC1 *
          (489409689/280 +
           skoXC1 *
           (126799656/35 +
            skoXC1 *
            (36670392/7 +
             skoXC1 *
             (10810656/5 +
              skoXC1 *
              (-58474688/5 +
               skoXC1 *
               (-29372416 +
                skoXC1 *
                (-31516672 +
                 skoXC1 * (-16625664 + skoXC1 * -3514368))))))))))))) |
     skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) <=
     -1/7),
    inference(decision, [],
              [refute_0_326, refute_0_271, refute_0_239, refute_0_208,
               refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_328, plain,
    (skoXC1 *
     (283/10 +
      skoXC1 *
      (2136/5 +
       skoXC1 *
       (2154 + skoXC1 * (4730 + skoXC1 * (4686 + skoXC1 * 1716))))) <=
     -1/7 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_306, refute_0_327])).

cnf(refute_0_329, plain, (1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_328, refute_0_289])).

cnf(refute_0_330, plain, (1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_329, refute_0_0])).

cnf(refute_0_331, plain, (skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_330])).

cnf(refute_0_332, plain, (-1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_326, refute_0_271, refute_0_239, refute_0_208,
               refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_333, plain, (skoXC1 <= 0),
    inference(resolve, [], [refute_0_331, refute_0_332])).

cnf(refute_0_334, plain, (0 < skoXC1),
    inference(decision, [],
              [refute_0_326, refute_0_271, refute_0_239, refute_0_208,
               refute_0_198, refute_0_173, refute_0_165, refute_0_157,
               refute_0_97, refute_0_68, refute_0_69])).

cnf(refute_0_335, plain, ($false),
    inference(resolve, [], [refute_0_333, refute_0_334])).
% SZS output end CNFRefutation for log-upper-21-sqrt-weak2.tptp
