% SZS output start CNFRefutation for log-error-analysis-1.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/12 * (3 * X ^ 3 + 131 * X ^ 2 + 239 * X + 47) * (X - 1) /
            (4 * X ^ 3 + 18 * X ^ 2 + 12 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/3 * (10 * X ^ 2 + 19 * X + 1) * (X - 1) /
            (X * (X ^ 2 + 6 * X + 3))) | lgen(R, Y, ln(X)))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(log_error_analysis_1, conjecture,
    (! [X] :
       ((0 <= X & X <= 1/128) =>
        abs(ln((1 + X / 2) / (1 - X / 2)) - X -
            208334991/2500000000 * X ^ 3) <= 3/62500000000000))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 1/128) =>
        abs(ln((1 + X / 2) / (1 - X / 2)) - X -
            208334991/2500000000 * X ^ 3) <= 3/62500000000000)),
    inference(strip, [], [log_error_analysis_1])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 1/128) =>
          abs(ln((1 + X / 2) / (1 - X / 2)) - X -
              208334991/2500000000 * X ^ 3) <= 3/62500000000000)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (3/62500000000000 <
        abs(ln((1 + X / 2) / (1 - X / 2)) - X -
            208334991/2500000000 * X ^ 3) & X <= 1/128 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (3/62500000000000 <
     abs(ln((1 + skoXC1 / 2) / (1 - skoXC1 / 2)) - skoXC1 -
         208334991/2500000000 * skoXC1 ^ 3) & 0 <= skoXC1 &
     skoXC1 <= 1/128), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (3/62500000000000 <
     abs(ln((1 + skoXC1 / 2) / (1 - skoXC1 / 2)) - skoXC1 -
         208334991/2500000000 * skoXC1 ^ 3)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 1/128),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 < (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     10 +
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000071, ln(X_000070)) | X_000071 <= ln(X_000070)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | lgen(0, X_000071, ln(X_000070))),
    inference(subst, [], [ln_lower_bound_cf5])).

cnf(refute_0_4, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    ((-1/3 + X_000070 * (-6 + X_000070 * (3 + X_000070 * 10/3))) /
     (X_000070 * (3 + X_000070 * (6 + X_000070))) < X_000071 |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((-1/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-6 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 10/3))) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (3/62500000000000 <
     abs(ln((1 + skoXC1 / 2) / (1 - skoXC1 / 2)) - skoXC1 -
         208334991/2500000000 * skoXC1 ^ 3)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (3/62500000000000 <
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (0 <=
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) =
     -(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
      ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_10, plain,
    (3/62500000000000 <
     -(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
      ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) !=
     -(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
      ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <=
     3/62500000000000), introduced(tautology, [equality])).

cnf(refute_0_11, plain,
    (3/62500000000000 <
     -(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
      ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) |
     0 <=
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <=
     3/62500000000000), inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (3/62500000000000 <
     -(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
      ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) |
     0 <=
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_11, refute_0_8])).

cnf(refute_0_13, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) |
     35 +
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) <=
     0 |
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_15, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) < 0 |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) =
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_16, plain,
    (3/62500000000000 <
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) !=
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <=
     3/62500000000000), introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) < 0 |
     3/62500000000000 <
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     abs(skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
         ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <=
     3/62500000000000),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) < 0 |
     3/62500000000000 <
     skoXC1 * (-1 + skoXC1 * (skoXC1 * -208334991/2500000000)) +
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_17, refute_0_8])).

cnf(refute_0_19, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    ((-1/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-6 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 10/3))) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_13])).

cnf(refute_0_21, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) /
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) |
     3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000057, ln(X_000056)) | X_000057 <= ln(X_000056)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_24, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     lgen(0, X_000057, ln(X_000056))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_25, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     X_000057 <= ln(X_000056)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    ((-1/2 + X_000056 * (-2 + X_000056 * 5/2)) /
     (X_000056 * (2 + X_000056)) < X_000057 | X_000056 <= 0 |
     X_000057 <= ln(X_000056)), inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    ((-1/2 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 5/2)) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    ((-1/2 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 5/2)) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_27, refute_0_13])).

cnf(refute_0_29, plain,
    (skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) /
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (skoXC1 <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (1 + skoXC1 * 1/2) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 / (1 + skoXC1 * 1/2) | 1 + skoXC1 * 1/2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_31, plain,
    (~ lgen(0, X_000043, ln(X_000042)) | X_000043 <= ln(X_000042)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_32, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     lgen(0, X_000043, ln(X_000042))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_33, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     X_000043 <= ln(X_000042)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    ((-1 + X_000042) / X_000042 < X_000043 | X_000042 <= 0 |
     X_000043 <= ln(X_000042)), inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    ((-1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    ((-1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_35, refute_0_13])).

cnf(refute_0_37, plain,
    (skoXC1 / (1 + skoXC1 * 1/2) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_39, plain, (skoXC1 <= 1/128),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_40, plain, (2 != skoXC1),
    inference(decision, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain,
    (skoXC1 / (1 + skoXC1 * 1/2) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_37, refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (1 + skoXC1 * 1/2) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | 1 + skoXC1 * 1/2 <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_30, refute_0_41])).

cnf(refute_0_43, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (skoXC1 <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (1 + skoXC1 * 1/2) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 / (1 + skoXC1 * 1/2) | 1 + skoXC1 * 1/2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_45, plain,
    ((-1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_46, plain,
    (skoXC1 / (1 + skoXC1 * 1/2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (skoXC1 / (1 + skoXC1 * 1/2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_46, refute_0_40])).

cnf(refute_0_48, plain,
    (skoXC1 <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (1 + skoXC1 * 1/2) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | 1 + skoXC1 * 1/2 <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_44, refute_0_47])).

cnf(refute_0_49, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     -2 < skoXC1), inference(decision, [], [refute_0_38, refute_0_39])).

cnf(refute_0_51, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_51, refute_0_0])).

cnf(refute_0_53, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     2 <= skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < 2 |
     0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 <= -2), inference(decision, [], [refute_0_38, refute_0_39])).

cnf(refute_0_55, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_55, refute_0_50])).

cnf(refute_0_57, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     -2 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [], [refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_58, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_43, refute_0_57])).

cnf(refute_0_59, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_58, refute_0_0])).

cnf(refute_0_60, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     2 <= skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 < 2 |
     3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(decision, [], [refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_62, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain,
    (3/62500000000000 <
     skoXC1 *
     (-3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_62, refute_0_57])).

cnf(refute_0_64, plain,
    (skoXC1 * (-1 + skoXC1 * 1/4) != -1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | 2 = skoXC1),
    inference(decision, [],
              [refute_0_63, refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_65, plain,
    (skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) /
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_29, refute_0_64])).

cnf(refute_0_66, plain,
    (2 != skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_63, refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_67, plain,
    (skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) /
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * (3 + skoXC1 * (-1/2 + skoXC1 * -1/2)) <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     3 + skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8)) <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_67])).

cnf(refute_0_69, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8)) <= -3 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (-3 < skoXC1 * (-1/2 + skoXC1 * (-3/4 + skoXC1 * 1/8)) |
     9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_63, refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_71, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_71, refute_0_0])).

cnf(refute_0_73, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     2 <= skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 < 2 |
     9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_63, refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_75, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     -2 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_63, refute_0_56, refute_0_38, refute_0_39])).

cnf(refute_0_77, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (9/250000000000000 +
       skoXC1 *
       (994599997/500000000000000 +
        skoXC1 *
        (416665009/5000000000 +
         skoXC1 *
         (-625004973/10000000000 + skoXC1 * 208334991/20000000000))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) != -1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(decision, [],
              [refute_0_77, refute_0_63, refute_0_56, refute_0_38,
               refute_0_39])).

cnf(refute_0_79, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_21, refute_0_78])).

cnf(refute_0_80, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_79, refute_0_64])).

cnf(refute_0_81, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_80, refute_0_66])).

cnf(refute_0_82, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     10 +
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     0 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_1, refute_0_81])).

cnf(refute_0_83, plain,
    (3/6250000000000 <
     skoXC1 *
     (21/62500000000000 +
      skoXC1 *
      (3/31250000000000 +
       skoXC1 *
       (621624991/93750000000000 +
        skoXC1 *
        (-6962199991/1500000000000000 +
         skoXC1 *
         (-375003978399991/3000000000000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     10 +
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_85, plain,
    ((-1/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-6 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 10/3))) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_86, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_86, refute_0_78])).

cnf(refute_0_88, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_87, refute_0_64])).

cnf(refute_0_89, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_66])).

cnf(refute_0_90, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     10 +
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     0 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_84, refute_0_89])).

cnf(refute_0_91, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     -10 <
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_77, refute_0_63, refute_0_56, refute_0_38,
               refute_0_39])).

cnf(refute_0_93, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_93, refute_0_0])).

cnf(refute_0_95, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) | 2 <= skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (skoXC1 < 2 |
     0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_77, refute_0_63, refute_0_56, refute_0_38,
               refute_0_39])).

cnf(refute_0_97, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) | -2 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_77, refute_0_63, refute_0_56, refute_0_38,
               refute_0_39])).

cnf(refute_0_99, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-34811/7500000000 +
         skoXC1 *
         (-468754973/3750000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 |
     skoXC1 *
     (21/62500000000000 +
      skoXC1 *
      (3/31250000000000 +
       skoXC1 *
       (621624991/93750000000000 +
        skoXC1 *
        (-6962199991/1500000000000000 +
         skoXC1 *
         (-375003978399991/3000000000000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) <= 3/6250000000000 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_99, refute_0_77, refute_0_63, refute_0_56,
               refute_0_38, refute_0_39])).

cnf(refute_0_101, plain,
    (skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_100, refute_0_83])).

cnf(refute_0_102, plain,
    (-10 <
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_99, refute_0_77, refute_0_63, refute_0_56,
               refute_0_38, refute_0_39])).

cnf(refute_0_103, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_103, refute_0_0])).

cnf(refute_0_105, plain,
    (2 <= skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (skoXC1 < 2 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_99, refute_0_77, refute_0_63, refute_0_56,
               refute_0_38, refute_0_39])).

cnf(refute_0_107, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) | skoXC1 <= -2),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (-2 < skoXC1 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(decision, [],
              [refute_0_99, refute_0_77, refute_0_63, refute_0_56,
               refute_0_38, refute_0_39])).

cnf(refute_0_109, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2))),
    inference(resolve, [], [refute_0_109, refute_0_19])).

cnf(refute_0_111, plain,
    (~ lgen(0, ln(X_000077), X_000078) | ln(X_000077) <= X_000078),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_112, plain,
    (~ lgen(0,
            1/12 *
            (3 * X_000077 ^ 3 + 131 * X_000077 ^ 2 + 239 * X_000077 + 47) *
            (X_000077 - 1) /
            (4 * X_000077 ^ 3 + 18 * X_000077 ^ 2 + 12 * X_000077 + 1),
            X_000078) | X_000077 <= 0 | lgen(0, ln(X_000077), X_000078)),
    inference(subst, [], [ln_upper_bound_cf7])).

cnf(refute_0_113, plain,
    (~ lgen(0,
            1/12 *
            (3 * X_000077 ^ 3 + 131 * X_000077 ^ 2 + 239 * X_000077 + 47) *
            (X_000077 - 1) /
            (4 * X_000077 ^ 3 + 18 * X_000077 ^ 2 + 12 * X_000077 + 1),
            X_000078) | X_000077 <= 0 | ln(X_000077) <= X_000078),
    inference(resolve, [], [refute_0_112, refute_0_111])).

cnf(refute_0_114, plain,
    (X_000078 <
     (-47/12 +
      X_000077 *
      (-16 + X_000077 * (9 + X_000077 * (32/3 + X_000077 * 1/4)))) /
     (1 + X_000077 * (12 + X_000077 * (18 + X_000077 * 4))) |
     X_000077 <= 0 | ln(X_000077) <= X_000078),
    inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-47/12 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-16 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (9 +
        (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
        (32/3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/4)))) /
     (1 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (12 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (18 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 4))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-47/12 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-16 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (9 +
        (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
        (32/3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/4)))) /
     (1 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (12 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (18 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 4))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_115, refute_0_110])).

cnf(refute_0_117, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-2 + skoXC1 * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/16))) =
     -1 | skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    ((3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) |
     10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= 0 |
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_119, plain,
    (~ lgen(0, ln(X_000063), X_000064) | ln(X_000063) <= X_000064),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_120, plain,
    (~ lgen(0,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0 | lgen(0, ln(X_000063), X_000064)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_121, plain,
    (~ lgen(0,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0 | ln(X_000063) <= X_000064),
    inference(resolve, [], [refute_0_120, refute_0_119])).

cnf(refute_0_122, plain,
    (X_000064 <
     (-10/3 + X_000063 * (-3 + X_000063 * (6 + X_000063 * 1/3))) /
     (1 + X_000063 * (6 + X_000063 * 3)) | X_000063 <= 0 |
     ln(X_000063) <= X_000064), inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-10/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/3))) /
     (1 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 3)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-10/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/3))) /
     (1 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 3)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_110])).

cnf(refute_0_125, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) != -1 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(decision, [], [refute_0_39, refute_0_38])).

cnf(refute_0_127, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain, (skoXC1 * (-1 + skoXC1 * 1/4) != -1 | 2 = skoXC1),
    inference(decision, [], [refute_0_39, refute_0_38])).

cnf(refute_0_129, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_40])).

cnf(refute_0_131, plain,
    ((3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) |
     10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_118, refute_0_130])).

cnf(refute_0_132, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= -10 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) |
     10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_134, plain,
    (ln(X_000063) < X_000064 | ~ lgen(1, ln(X_000063), X_000064)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_135, plain,
    (~ lgen(1,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0 | lgen(1, ln(X_000063), X_000064)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_136, plain,
    (ln(X_000063) < X_000064 |
     ~ lgen(1,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0), inference(resolve, [], [refute_0_135, refute_0_134])).

cnf(refute_0_137, plain,
    (ln(X_000063) < X_000064 | X_000063 <= 0 |
     X_000064 <=
     (-10/3 + X_000063 * (-3 + X_000063 * (6 + X_000063 * 1/3))) /
     (1 + X_000063 * (6 + X_000063 * 3))),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     (-10/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/3))) /
     (1 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 3)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_137])).

cnf(refute_0_139, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_139, refute_0_126])).

cnf(refute_0_141, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_140, refute_0_128])).

cnf(refute_0_142, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) /
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_141, refute_0_40])).

cnf(refute_0_143, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) *
     (10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8)))) <=
     skoXC1 * (10 + skoXC1 * (-8 + skoXC1 * (5/6 + skoXC1 * 1/3))) |
     10 + skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_142, refute_0_133])).

cnf(refute_0_144, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= -10 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    ((3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (3 + skoXC1 * (-1 + skoXC1 * -1/4)) < skoXC1 * (3 + skoXC1 * -1) |
     3 + skoXC1 * (-1 + skoXC1 * -1/4) <= 0 |
     skoXC1 * (3 + skoXC1 * -1) / (3 + skoXC1 * (-1 + skoXC1 * -1/4)) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_146, plain,
    (~ lgen(0, ln(X_000049), X_000050) | ln(X_000049) <= X_000050),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_147, plain,
    (~ lgen(0, 1/2 * (X_000049 + 5) * (X_000049 - 1) / (2 * X_000049 + 1),
            X_000050) | X_000049 <= 0 | lgen(0, ln(X_000049), X_000050)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_148, plain,
    (~ lgen(0, 1/2 * (X_000049 + 5) * (X_000049 - 1) / (2 * X_000049 + 1),
            X_000050) | X_000049 <= 0 | ln(X_000049) <= X_000050),
    inference(resolve, [], [refute_0_147, refute_0_146])).

cnf(refute_0_149, plain,
    (X_000050 <
     (-5/2 + X_000049 * (2 + X_000049 * 1/2)) / (1 + X_000049 * 2) |
     X_000049 <= 0 | ln(X_000049) <= X_000050),
    inference(arithmetic, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-5/2 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/2)) /
     (1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 2) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (-5/2 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (2 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 1/2)) /
     (1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 2) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_150, refute_0_110])).

cnf(refute_0_152, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (3 + skoXC1 * -1) / (3 + skoXC1 * (-1 + skoXC1 * -1/4)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (3 + skoXC1 * -1) / (3 + skoXC1 * (-1 + skoXC1 * -1/4)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_152, refute_0_128])).

cnf(refute_0_154, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 * (3 + skoXC1 * -1) / (3 + skoXC1 * (-1 + skoXC1 * -1/4)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_153, refute_0_40])).

cnf(refute_0_155, plain,
    ((3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (3 + skoXC1 * (-1 + skoXC1 * -1/4)) < skoXC1 * (3 + skoXC1 * -1) |
     3 + skoXC1 * (-1 + skoXC1 * -1/4) <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_145, refute_0_154])).

cnf(refute_0_156, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     skoXC1 * (-1 + skoXC1 * -1/4) <= -3 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    ((62500000000003/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (1 + skoXC1 * -1/2) < 1 + skoXC1 * 1/2 | 1 + skoXC1 * -1/2 <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <=
     62500000000003/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_158, plain,
    (~ lgen(0, ln(X_000035), X_000036) | ln(X_000035) <= X_000036),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_159, plain,
    (~ lgen(0, X_000035 - 1, X_000036) | X_000035 <= 0 |
     lgen(0, ln(X_000035), X_000036)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_160, plain,
    (~ lgen(0, X_000035 - 1, X_000036) | X_000035 <= 0 |
     ln(X_000035) <= X_000036),
    inference(resolve, [], [refute_0_159, refute_0_158])).

cnf(refute_0_161, plain,
    (X_000036 < -1 + X_000035 | X_000035 <= 0 | ln(X_000035) <= X_000036),
    inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     -1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <=
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(subst, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     -1 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_162, refute_0_110])).

cnf(refute_0_164, plain,
    (62500000000003/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    ((62500000000003/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (1 + skoXC1 * -1/2) < 1 + skoXC1 * 1/2 | 1 + skoXC1 * -1/2 <= 0 |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_157, refute_0_164])).

cnf(refute_0_166, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     2 <= skoXC1 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (skoXC1 < 2 |
     3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000)))),
    inference(decision, [], [refute_0_39, refute_0_38])).

cnf(refute_0_168, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_166, refute_0_167])).

cnf(refute_0_169, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_168, refute_0_0])).

cnf(refute_0_170, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (skoXC1 < 2 |
     3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 <= -2), inference(decision, [], [refute_0_39, refute_0_38])).

cnf(refute_0_172, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     skoXC1 <= -2), inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000))) |
     -2 < skoXC1), inference(decision, [], [refute_0_39, refute_0_38])).

cnf(refute_0_174, plain,
    (3/62500000000000 <
     skoXC1 *
     (3/125000000000000 +
      skoXC1 *
      (1/2 +
       skoXC1 * (-208334991/2500000000 + skoXC1 * 208334991/5000000000)))),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (-3 < skoXC1 * (-1 + skoXC1 * -1/4) |
     9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000))))),
    inference(decision, [], [refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_176, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_156, refute_0_175])).

cnf(refute_0_177, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_176, refute_0_0])).

cnf(refute_0_178, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_177])).

cnf(refute_0_179, plain,
    (skoXC1 < 2 |
     9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     skoXC1 <= -2),
    inference(decision, [], [refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_180, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     skoXC1 <= -2), inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000)))) |
     -2 < skoXC1),
    inference(decision, [], [refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_182, plain,
    (9/62500000000000 <
     skoXC1 *
     (3/62500000000000 +
      skoXC1 *
      (3/250000000000000 +
       skoXC1 *
       (-4973/2500000000 +
        skoXC1 *
        (208334991/2500000000 + skoXC1 * 208334991/10000000000))))),
    inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (-10 < skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0),
    inference(decision, [],
              [refute_0_182, refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_184, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_144, refute_0_183])).

cnf(refute_0_185, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | 1 + skoXC1 * -1/2 <= 0 |
     1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_184, refute_0_0])).

cnf(refute_0_186, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     2 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | skoXC1 <= -2), inference(arithmetic, [], [refute_0_185])).

cnf(refute_0_187, plain,
    (skoXC1 < 2 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_182, refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_188, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0 | skoXC1 <= -2),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (-2 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0),
    inference(decision, [],
              [refute_0_182, refute_0_174, refute_0_39, refute_0_38])).

cnf(refute_0_190, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (4973/750000000 +
        skoXC1 *
        (-4973/937500000 +
         skoXC1 *
         (-1/8 +
          skoXC1 *
          (208334991/2500000000 + skoXC1 * -208334991/20000000000)))))) <=
     0), inference(resolve, [], [refute_0_188, refute_0_189])).

cnf(refute_0_191, plain,
    (-10 < skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) |
     3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(decision, [],
              [refute_0_190, refute_0_182, refute_0_174, refute_0_39,
               refute_0_38])).

cnf(refute_0_192, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_132, refute_0_191])).

cnf(refute_0_193, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_192, refute_0_0])).

cnf(refute_0_194, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 < 2 |
     3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 <= -2),
    inference(decision, [],
              [refute_0_190, refute_0_182, refute_0_174, refute_0_39,
               refute_0_38])).

cnf(refute_0_196, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 <= -2), inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     -2 < skoXC1),
    inference(decision, [],
              [refute_0_190, refute_0_182, refute_0_174, refute_0_39,
               refute_0_38])).

cnf(refute_0_198, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(resolve, [], [refute_0_196, refute_0_197])).

cnf(refute_0_199, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (-2 + skoXC1 * (3/2 + skoXC1 * (-1/2 + skoXC1 * 1/16))) !=
     -1 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_200, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_117, refute_0_199])).

cnf(refute_0_201, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) != -1 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_202, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_200, refute_0_201])).

cnf(refute_0_203, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 * (-1 + skoXC1 * 1/4) != -1 | 2 = skoXC1),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_204, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     2 != skoXC1),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_206, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) /
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_204, refute_0_205])).

cnf(refute_0_207, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (45/4 +
        skoXC1 *
        (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64)))))) <
     skoXC1 *
     (35 +
      skoXC1 *
      (-45 +
       skoXC1 *
       (85/6 + skoXC1 * (11/4 + skoXC1 * (-29/16 + skoXC1 * 1/6))))) |
     35 +
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) <=
     0 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_14, refute_0_206])).

cnf(refute_0_208, plain,
    (21/12500000000000 <
     skoXC1 *
     (27/12500000000000 +
      skoXC1 *
      (-27/50000000000000 +
       skoXC1 *
       (-8702750117/375000000000000 +
        skoXC1 *
        (29838000153/1000000000000000 +
         skoXC1 *
         (437492540499991/1000000000000000 +
          skoXC1 *
          (-6750051719200027/12000000000000000 +
           skoXC1 *
           (8750084541/40000000000 +
            skoXC1 *
            (-625004973/40000000000 +
             skoXC1 * -625004973/160000000000)))))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) <=
     -35 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) <=
     -35 |
     skoXC1 *
     (27/12500000000000 +
      skoXC1 *
      (-27/50000000000000 +
       skoXC1 *
       (-8702750117/375000000000000 +
        skoXC1 *
        (29838000153/1000000000000000 +
         skoXC1 *
         (437492540499991/1000000000000000 +
          skoXC1 *
          (-6750051719200027/12000000000000000 +
           skoXC1 *
           (8750084541/40000000000 +
            skoXC1 *
            (-625004973/40000000000 +
             skoXC1 * -625004973/160000000000)))))))) <=
     21/12500000000000),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_210, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) <=
     -35 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_209, refute_0_208])).

cnf(refute_0_211, plain,
    (-35 <
     skoXC1 *
     (-45 +
      skoXC1 *
      (45/4 +
       skoXC1 *
       (13/2 + skoXC1 * (-51/16 + skoXC1 * (3/16 + skoXC1 * 3/64))))) |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_212, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_212, refute_0_0])).

cnf(refute_0_214, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (skoXC1 < 2 |
     ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 <= -2),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_216, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     skoXC1 <= -2), inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     -2 < skoXC1),
    inference(decision, [],
              [refute_0_198, refute_0_190, refute_0_182, refute_0_174,
               refute_0_39, refute_0_38])).

cnf(refute_0_218, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(resolve, [], [refute_0_216, refute_0_217])).

cnf(refute_0_219, plain,
    (ln((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))),
    inference(resolve, [], [refute_0_13, refute_0_218])).

cnf(refute_0_220, plain,
    ((-1/3 +
      (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (-6 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (3 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) * 10/3))) /
     ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
      (3 +
       (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) *
       (6 + (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2)))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_219])).

cnf(refute_0_221, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-3/2 + skoXC1 * (3/4 + skoXC1 * -1/8)) = -1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1 |
     skoXC1 * (-1 + skoXC1 * 1/4) = -1),
    inference(resolve, [], [refute_0_221, refute_0_126])).

cnf(refute_0_223, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0 | 2 = skoXC1),
    inference(resolve, [], [refute_0_222, refute_0_128])).

cnf(refute_0_224, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) /
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) <
     -3/62500000000000 +
     skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000)) |
     (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_223, refute_0_40])).

cnf(refute_0_225, plain,
    (skoXC1 *
     (10 +
      skoXC1 * (-7 + skoXC1 * (-7/6 + skoXC1 * (17/12 + skoXC1 * -1/6)))) <
     (-3/62500000000000 +
      skoXC1 * (1 + skoXC1 * (skoXC1 * 208334991/2500000000))) *
     (10 +
      skoXC1 *
      (-7 +
       skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))) |
     10 +
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     0 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_1, refute_0_224])).

cnf(refute_0_226, plain,
    (3/6250000000000 <
     skoXC1 *
     (21/62500000000000 +
      skoXC1 *
      (3/31250000000000 +
       skoXC1 *
       (621624991/93750000000000 +
        skoXC1 *
        (-6962199991/1500000000000000 +
         skoXC1 *
         (-375003978399991/3000000000000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) |
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_225])).

cnf(refute_0_227, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= -10 |
     1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_132, refute_0_0])).

cnf(refute_0_228, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     2 <= skoXC1 |
     skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) <= -10 |
     skoXC1 <= -2), inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (-10 < skoXC1 * (-8 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/8))) |
     3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_38, refute_0_39, refute_0_174, refute_0_182])).

cnf(refute_0_230, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     2 <= skoXC1 | skoXC1 <= -2),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    (skoXC1 < 2 |
     3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     skoXC1 <= -2),
    inference(decision, [],
              [refute_0_38, refute_0_39, refute_0_174, refute_0_182])).

cnf(refute_0_232, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     skoXC1 <= -2), inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000)))))) |
     -2 < skoXC1),
    inference(decision, [],
              [refute_0_38, refute_0_39, refute_0_174, refute_0_182])).

cnf(refute_0_234, plain,
    (3/6250000000000 <
     skoXC1 *
     (3/7812500000000 +
      skoXC1 *
      (skoXC1 *
       (-1243250009/187500000000000 +
        skoXC1 *
        (7956800009/1500000000000000 +
         skoXC1 *
         (1/8 +
          skoXC1 *
          (-208334991/2500000000 + skoXC1 * 208334991/20000000000))))))),
    inference(resolve, [], [refute_0_232, refute_0_233])).

cnf(refute_0_235, plain,
    (skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 |
     skoXC1 *
     (21/62500000000000 +
      skoXC1 *
      (3/31250000000000 +
       skoXC1 *
       (621624991/93750000000000 +
        skoXC1 *
        (-6962199991/1500000000000000 +
         skoXC1 *
         (-375003978399991/3000000000000000 +
          skoXC1 *
          (130209991/1250000000 +
           skoXC1 *
           (-208334991/20000000000 +
            skoXC1 * -208334991/40000000000))))))) <= 3/6250000000000),
    inference(decision, [],
              [refute_0_234, refute_0_38, refute_0_39, refute_0_174,
               refute_0_182])).

cnf(refute_0_236, plain,
    (skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16)))) <=
     -10 | (1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_235, refute_0_226])).

cnf(refute_0_237, plain,
    (-10 <
     skoXC1 *
     (-7 +
      skoXC1 * (-2 + skoXC1 * (2 + skoXC1 * (-1/8 + skoXC1 * -1/16))))),
    inference(decision, [],
              [refute_0_234, refute_0_38, refute_0_39, refute_0_174,
               refute_0_182])).

cnf(refute_0_238, plain, ((1 + skoXC1 * 1/2) / (1 + skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_236, refute_0_237])).

cnf(refute_0_239, plain,
    (1 + skoXC1 * -1/2 <= 0 | 1 + skoXC1 * 1/2 <= 0 * (1 + skoXC1 * -1/2)),
    inference(resolve, [], [refute_0_238, refute_0_0])).

cnf(refute_0_240, plain, (2 <= skoXC1 | skoXC1 <= -2),
    inference(arithmetic, [], [refute_0_239])).

cnf(refute_0_241, plain, (skoXC1 < 2 | skoXC1 <= -2),
    inference(decision, [],
              [refute_0_234, refute_0_38, refute_0_39, refute_0_174,
               refute_0_182])).

cnf(refute_0_242, plain, (skoXC1 <= -2),
    inference(resolve, [], [refute_0_240, refute_0_241])).

cnf(refute_0_243, plain, (-2 < skoXC1),
    inference(decision, [],
              [refute_0_234, refute_0_38, refute_0_39, refute_0_174,
               refute_0_182])).

cnf(refute_0_244, plain, ($false),
    inference(resolve, [], [refute_0_242, refute_0_243])).
% SZS output end CNFRefutation for log-error-analysis-1.tptp
