% SZS output start CNFRefutation for log-nest-exp-upper2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_1, axiom,
    (X < -1 | ~ lgen(R, Y, 1 + X) | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

fof(log_nest_exp_upper2, conjecture,
    (! [X] : ((0 < X & X < 3) => ln(exp(X) - 1) < X + 3))).

fof(subgoal_0, plain,
    (! [X] : ((0 < X & X < 3) => ln(exp(X) - 1) < X + 3)),
    inference(strip, [], [log_nest_exp_upper2])).

fof(negate_0_0, plain,
    (~ ! [X] : ((0 < X & X < 3) => ln(exp(X) - 1) < X + 3)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (0 < X & X < 3 & X + 3 <= ln(exp(X) - 1))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 3 & 0 < skoXC1 & skoXC1 + 3 <= ln(exp(skoXC1) - 1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (skoXC1 + 3 <= ln(exp(skoXC1) - 1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 3),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     7247757312 + skoXC1 * 7247757312 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0 |
     (7247757312 + skoXC1 * 7247757312) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     4 + skoXC1 * (7 + skoXC1 * 3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (1/3 + skoXC1 * 1/3) | 1/3 + skoXC1 * 1/3 <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     (4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    (~ lgen(0, exp(X_000075), X_000076) | exp(X_000075) <= X_000076),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (3 < X_000075 |
     ~ lgen(0,
            21743271936 /
            (-(X_000075 ^ 3) + 12 * X_000075 ^ 2 - 96 * X_000075 + 384) ^
            4, X_000076) | lgen(0, exp(X_000075), X_000076)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_4, plain,
    (3 < X_000075 |
     ~ lgen(0,
            21743271936 /
            (-(X_000075 ^ 3) + 12 * X_000075 ^ 2 - 96 * X_000075 + 384) ^
            4, X_000076) | exp(X_000075) <= X_000076),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (3 < X_000075 |
     X_000076 <
     21743271936 /
     (21743271936 +
      X_000075 *
      (-21743271936 +
       X_000075 *
       (10871635968 +
        X_000075 *
        (-3623878656 +
         X_000075 *
         (891813888 +
          X_000075 *
          (-169869312 +
           X_000075 *
           (25657344 +
            X_000075 *
            (-3096576 +
             X_000075 *
             (297216 +
              X_000075 *
              (-22272 +
               X_000075 * (1248 + X_000075 * (-48 + X_000075)))))))))))) |
     exp(X_000075) <= X_000076), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 |
     exp(skoXC1) <=
     (4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    ((34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <
     1 + skoXC1 |
     (1 + skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2 |
     14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_8, plain,
    (X_000048 < exp(X_000047) | ~ lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (X_000047 < -1 | ~ lgen(1, X_000048, 1 + X_000047) |
     lgen(1, X_000048, exp(X_000047))),
    inference(subst, [], [exp_lower_taylor_1])).

cnf(refute_0_10, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) |
     ~ lgen(1, X_000048, 1 + X_000047)),
    inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (X_000047 < -1 | X_000048 < exp(X_000047) | 1 + X_000047 <= X_000048),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (skoXC1 < -1 |
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <
     exp(skoXC1) |
     1 + skoXC1 <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((16/3 + exp(skoXC1) * (-14 + exp(skoXC1) * (5 + exp(skoXC1) * 1/3))) /
     (-2 + exp(skoXC1) * (exp(skoXC1) * 3)) < 3 + skoXC1 |
     (3 + skoXC1) * (-2 + exp(skoXC1) * (exp(skoXC1) * 3)) <=
     16/3 + exp(skoXC1) * (-14 + exp(skoXC1) * (5 + exp(skoXC1) * 1/3)) |
     -2 + exp(skoXC1) * (exp(skoXC1) * 3) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_14, plain, (skoXC1 + 3 <= ln(exp(skoXC1) - 1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_15, plain, (3 + skoXC1 <= ln(-1 + exp(skoXC1))),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (ln(X_000096) < X_000097 | ~ lgen(1, ln(X_000096), X_000097)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain,
    (~ lgen(1,
            1/3 * (X_000096 ^ 2 + 19 * X_000096 + 10) * (X_000096 - 1) /
            (3 * X_000096 ^ 2 + 6 * X_000096 + 1), X_000097) |
     X_000096 <= 0 | lgen(1, ln(X_000096), X_000097)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_18, plain,
    (ln(X_000096) < X_000097 |
     ~ lgen(1,
            1/3 * (X_000096 ^ 2 + 19 * X_000096 + 10) * (X_000096 - 1) /
            (3 * X_000096 ^ 2 + 6 * X_000096 + 1), X_000097) |
     X_000096 <= 0), inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (ln(X_000096) < X_000097 | X_000096 <= 0 |
     X_000097 <=
     (-10/3 + X_000096 * (-3 + X_000096 * (6 + X_000096 * 1/3))) /
     (1 + X_000096 * (6 + X_000096 * 3))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (ln(-1 + exp(skoXC1)) < 3 + skoXC1 | -1 + exp(skoXC1) <= 0 |
     3 + skoXC1 <=
     (-10/3 +
      (-1 + exp(skoXC1)) *
      (-3 + (-1 + exp(skoXC1)) * (6 + (-1 + exp(skoXC1)) * 1/3))) /
     (1 + (-1 + exp(skoXC1)) * (6 + (-1 + exp(skoXC1)) * 3))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (-1 + exp(skoXC1) <= 0 |
     3 + skoXC1 <=
     (-10/3 +
      (-1 + exp(skoXC1)) *
      (-3 + (-1 + exp(skoXC1)) * (6 + (-1 + exp(skoXC1)) * 1/3))) /
     (1 + (-1 + exp(skoXC1)) * (6 + (-1 + exp(skoXC1)) * 3))),
    inference(resolve, [], [refute_0_15, refute_0_20])).

cnf(refute_0_22, plain,
    (3 + skoXC1 <=
     (16/3 + exp(skoXC1) * (-14 + exp(skoXC1) * (5 + exp(skoXC1) * 1/3))) /
     (-2 + exp(skoXC1) * (exp(skoXC1) * 3)) | exp(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 < -1 | 1 < exp(skoXC1) | 1 + skoXC1 <= 1),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_24, plain, (skoXC1 < -1 | 1 < exp(skoXC1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_26, plain, (skoXC1 < 3),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_27, plain, (-1 <= skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_25, refute_0_26])).

cnf(refute_0_28, plain, (1 < exp(skoXC1) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_27, refute_0_24])).

cnf(refute_0_29, plain, (1 < exp(skoXC1)),
    inference(resolve, [], [refute_0_28, refute_0_25])).

cnf(refute_0_30, plain,
    (3 + skoXC1 <=
     (16/3 + exp(skoXC1) * (-14 + exp(skoXC1) * (5 + exp(skoXC1) * 1/3))) /
     (-2 + exp(skoXC1) * (exp(skoXC1) * 3))),
    inference(resolve, [], [refute_0_22, refute_0_29])).

cnf(refute_0_31, plain,
    ((3 + skoXC1) * (-2 + exp(skoXC1) * (exp(skoXC1) * 3)) <=
     16/3 + exp(skoXC1) * (-14 + exp(skoXC1) * (5 + exp(skoXC1) * 1/3)) |
     -2 + exp(skoXC1) * (exp(skoXC1) * 3) <= 0),
    inference(resolve, [], [refute_0_30, refute_0_13])).

cnf(refute_0_32, plain,
    (exp(skoXC1) * (exp(skoXC1) * 3) <= 2 |
     exp(skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2), inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (2 / (3 + skoXC1 * 3) < 1 + skoXC1 |
     (1 + skoXC1) * (3 + skoXC1 * 3) <= 2 | 3 + skoXC1 * 3 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_34, plain,
    (skoXC1 < -1 | 2 / (3 + skoXC1 * 3) < exp(skoXC1) |
     1 + skoXC1 <= 2 / (3 + skoXC1 * 3)),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_35, plain,
    (2 / (exp(skoXC1) * 3) < 1 + skoXC1 |
     (1 + skoXC1) * (exp(skoXC1) * 3) <= 2 | exp(skoXC1) * 3 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_36, plain,
    (skoXC1 < -1 | 2 / (exp(skoXC1) * 3) < exp(skoXC1) |
     1 + skoXC1 <= 2 / (exp(skoXC1) * 3)),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_37, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) * (exp(skoXC1) * 3) <= 2),
    introduced(tautology, [assume])).

cnf(refute_0_38, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | exp(skoXC1) * 3 <= 0 |
     exp(skoXC1) <= 2 / (exp(skoXC1) * 3)),
    inference(split, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) <= 2 / (exp(skoXC1) * 3)),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (skoXC1 < -1 | 2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (exp(skoXC1) * 3)),
    inference(resolve, [], [refute_0_39, refute_0_36])).

cnf(refute_0_41, plain,
    ((10 + skoXC1 * 4) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     10 + skoXC1 * 4), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_42, plain,
    (10 + skoXC1 * 4 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | exp(skoXC1) <= 10 + skoXC1 * 4),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_43, plain,
    ((-4 + exp(skoXC1) * (1 + exp(skoXC1) * 1/2)) /
     (-1 + exp(skoXC1) * 2) < 3 + skoXC1 |
     (3 + skoXC1) * (-1 + exp(skoXC1) * 2) <=
     -4 + exp(skoXC1) * (1 + exp(skoXC1) * 1/2) |
     -1 + exp(skoXC1) * 2 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_44, plain,
    (ln(X_000061) < X_000062 | ~ lgen(1, ln(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_45, plain,
    (~ lgen(1, 1/2 * (X_000061 + 5) * (X_000061 - 1) / (2 * X_000061 + 1),
            X_000062) | X_000061 <= 0 | lgen(1, ln(X_000061), X_000062)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_46, plain,
    (ln(X_000061) < X_000062 |
     ~ lgen(1, 1/2 * (X_000061 + 5) * (X_000061 - 1) / (2 * X_000061 + 1),
            X_000062) | X_000061 <= 0),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain,
    (ln(X_000061) < X_000062 | X_000061 <= 0 |
     X_000062 <=
     (-5/2 + X_000061 * (2 + X_000061 * 1/2)) / (1 + X_000061 * 2)),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (ln(-1 + exp(skoXC1)) < 3 + skoXC1 | -1 + exp(skoXC1) <= 0 |
     3 + skoXC1 <=
     (-5/2 + (-1 + exp(skoXC1)) * (2 + (-1 + exp(skoXC1)) * 1/2)) /
     (1 + (-1 + exp(skoXC1)) * 2)), inference(subst, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (-1 + exp(skoXC1) <= 0 |
     3 + skoXC1 <=
     (-5/2 + (-1 + exp(skoXC1)) * (2 + (-1 + exp(skoXC1)) * 1/2)) /
     (1 + (-1 + exp(skoXC1)) * 2)),
    inference(resolve, [], [refute_0_15, refute_0_48])).

cnf(refute_0_50, plain,
    (3 + skoXC1 <=
     (-4 + exp(skoXC1) * (1 + exp(skoXC1) * 1/2)) /
     (-1 + exp(skoXC1) * 2) | exp(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (3 + skoXC1 <=
     (-4 + exp(skoXC1) * (1 + exp(skoXC1) * 1/2)) /
     (-1 + exp(skoXC1) * 2)),
    inference(resolve, [], [refute_0_50, refute_0_29])).

cnf(refute_0_52, plain,
    ((3 + skoXC1) * (-1 + exp(skoXC1) * 2) <=
     -4 + exp(skoXC1) * (1 + exp(skoXC1) * 1/2) |
     -1 + exp(skoXC1) * 2 <= 0),
    inference(resolve, [], [refute_0_51, refute_0_43])).

cnf(refute_0_53, plain,
    (exp(skoXC1) * (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) <= -1 + skoXC1 |
     exp(skoXC1) <= 1/2), inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < -1 | 1/2 < exp(skoXC1) | 1 + skoXC1 <= 1/2),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_55, plain, (skoXC1 < -1 | 1/2 < exp(skoXC1) | skoXC1 <= -1/2),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain, (-1 <= skoXC1 | skoXC1 <= -1/2),
    inference(decision, [], [refute_0_25, refute_0_26])).

cnf(refute_0_57, plain, (1/2 < exp(skoXC1) | skoXC1 <= -1/2),
    inference(resolve, [], [refute_0_56, refute_0_55])).

cnf(refute_0_58, plain, (-1/2 < skoXC1),
    inference(decision, [], [refute_0_25, refute_0_26])).

cnf(refute_0_59, plain, (1/2 < exp(skoXC1)),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (exp(skoXC1) * (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) <= -1 + skoXC1),
    inference(resolve, [], [refute_0_53, refute_0_59])).

cnf(refute_0_61, plain,
    (5 + skoXC1 * 2 + exp(skoXC1) * -1/2 < 0 |
     0 < 5 + skoXC1 * 2 + exp(skoXC1) * -1/2 | 0 <= -1 + skoXC1),
    inference(split, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (exp(skoXC1) < 10 + skoXC1 * 4 | 10 + skoXC1 * 4 < exp(skoXC1) |
     1 <= skoXC1), inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (0 * (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) < -1 + skoXC1 |
     5 + skoXC1 * 2 + exp(skoXC1) * -1/2 <= 0 |
     (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_64, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_65, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_66, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_65, refute_0_64])).

cnf(refute_0_67, plain,
    (0 < (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) |
     (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) < exp(skoXC1)),
    inference(subst, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (5 + skoXC1 * 2 + exp(skoXC1) * -1/2 <= 0 |
     exp(skoXC1) <= (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2)),
    inference(split, [], [refute_0_60])).

cnf(refute_0_69, plain,
    (10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <= (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2)),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (0 < (-1 + skoXC1) / (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_69, refute_0_67])).

cnf(refute_0_71, plain,
    (0 * (5 + skoXC1 * 2 + exp(skoXC1) * -1/2) < -1 + skoXC1 |
     10 + skoXC1 * 4 <= exp(skoXC1) |
     5 + skoXC1 * 2 + exp(skoXC1) * -1/2 <= 0),
    inference(resolve, [], [refute_0_63, refute_0_70])).

cnf(refute_0_72, plain, (1 < skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain, (-1 <= skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_26, refute_0_25, refute_0_72])).

cnf(refute_0_74, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (exp(skoXC1) * 3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_73, refute_0_40])).

cnf(refute_0_75, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     (1 + skoXC1) * (exp(skoXC1) * 3) <= 2 | exp(skoXC1) * 3 <= 0 |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_74, refute_0_35])).

cnf(refute_0_76, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) * (3 + skoXC1 * 3) <= 2 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 3 + skoXC1 * 3 <= 0 |
     10 + skoXC1 * 4 <= exp(skoXC1) | exp(skoXC1) <= 2 / (3 + skoXC1 * 3)),
    inference(split, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <= 2 / (3 + skoXC1 * 3) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain, (-1 < skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_26, refute_0_25, refute_0_72])).

cnf(refute_0_80, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <= 2 / (3 + skoXC1 * 3)),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain,
    (skoXC1 < -1 | 2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (3 + skoXC1 * 3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_80, refute_0_34])).

cnf(refute_0_82, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (3 + skoXC1 * 3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_73, refute_0_81])).

cnf(refute_0_83, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     (1 + skoXC1) * (3 + skoXC1 * 3) <= 2 | 3 + skoXC1 * 3 <= 0 |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_82, refute_0_33])).

cnf(refute_0_84, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     skoXC1 * (6 + skoXC1 * 3) <= -1 | 10 + skoXC1 * 4 <= exp(skoXC1) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (-1 < skoXC1 * (6 + skoXC1 * 3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     skoXC1 <= -1),
    inference(decision, [], [refute_0_26, refute_0_25, refute_0_72])).

cnf(refute_0_86, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     skoXC1 <= -1), inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_86, refute_0_79])).

cnf(refute_0_88, plain,
    (exp(skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_32, refute_0_87])).

cnf(refute_0_89, plain,
    (10 + skoXC1 * 4 <= exp(skoXC1) |
     14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0 |
     exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(split, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= -14 |
     10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    ((12 + skoXC1 * 9) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <=
     12 + skoXC1 * 9), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_92, plain,
    (12 + skoXC1 * 9 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | exp(skoXC1) <= 12 + skoXC1 * 9),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_93, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= -14),
    introduced(tautology, [assume])).

cnf(refute_0_94, plain,
    (4 + skoXC1 * 3 + exp(skoXC1) * -1/3 < 0 |
     -14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     0 < 4 + skoXC1 * 3 + exp(skoXC1) * -1/3 | 0 <= -14),
    inference(split, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     exp(skoXC1) < 12 + skoXC1 * 9 | 12 + skoXC1 * 9 < exp(skoXC1)),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) < -14 |
     4 + skoXC1 * 3 + exp(skoXC1) * -1/3 <= 0 |
     -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_97, plain,
    (0 < -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) < exp(skoXC1)),
    inference(subst, [], [refute_0_66])).

cnf(refute_0_98, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     4 + skoXC1 * 3 + exp(skoXC1) * -1/3 <= 0 |
     exp(skoXC1) <= -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)),
    inference(split, [], [refute_0_93])).

cnf(refute_0_99, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <= exp(skoXC1) |
     exp(skoXC1) <= -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     0 < -14 / (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_99, refute_0_97])).

cnf(refute_0_101, plain,
    (0 * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) < -14 |
     -14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <= exp(skoXC1) |
     4 + skoXC1 * 3 + exp(skoXC1) * -1/3 <= 0),
    inference(resolve, [], [refute_0_96, refute_0_100])).

cnf(refute_0_102, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 < exp(skoXC1)),
    inference(resolve, [], [refute_0_102, refute_0_95])).

cnf(refute_0_104, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1), inference(resolve, [], [refute_0_92, refute_0_103])).

cnf(refute_0_105, plain, (10 + skoXC1 * 4 <= exp(skoXC1) | skoXC1 <= 3),
    inference(decision, [], [refute_0_72, refute_0_25, refute_0_26])).

cnf(refute_0_106, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    ((12 + skoXC1 * 9) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 |
     -14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     10 + skoXC1 * 4 <= exp(skoXC1) |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_91, refute_0_106])).

cnf(refute_0_108, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     239175991296 <
     skoXC1 *
     (65229815808 +
      skoXC1 *
      (65229815808 +
       skoXC1 *
       (-54358179840 +
        skoXC1 *
        (21913141248 +
         skoXC1 *
         (-5987893248 +
          skoXC1 *
          (1220935680 +
           skoXC1 *
           (-193757184 +
            skoXC1 *
            (24302592 +
             skoXC1 *
             (-2407680 +
              skoXC1 *
              (185472 +
               skoXC1 * (-10656 + skoXC1 * (420 + skoXC1 * -9)))))))))))) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 |
     skoXC1 *
     (65229815808 +
      skoXC1 *
      (65229815808 +
       skoXC1 *
       (-54358179840 +
        skoXC1 *
        (21913141248 +
         skoXC1 *
         (-5987893248 +
          skoXC1 *
          (1220935680 +
           skoXC1 *
           (-193757184 +
            skoXC1 *
            (24302592 +
             skoXC1 *
             (-2407680 +
              skoXC1 *
              (185472 +
               skoXC1 *
               (-10656 + skoXC1 * (420 + skoXC1 * -9)))))))))))) <=
     239175991296 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_72, refute_0_25, refute_0_26])).

cnf(refute_0_110, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_109, refute_0_108])).

cnf(refute_0_111, plain,
    (-21743271936 <
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 + skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_72, refute_0_25, refute_0_26])).

cnf(refute_0_112, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(resolve, [], [refute_0_90, refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_113, refute_0_12])).

cnf(refute_0_115, plain,
    (1 + skoXC1 <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_73, refute_0_114])).

cnf(refute_0_116, plain,
    ((1 + skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2 | 10 + skoXC1 * 4 <= exp(skoXC1) |
     14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0),
    inference(resolve, [], [refute_0_115, refute_0_7])).

cnf(refute_0_117, plain,
    (exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= -14 |
     exp(skoXC1) *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) <= -8/3 + skoXC1 * -12 |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (exp(skoXC1) *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) <= -8/3 + skoXC1 * -12 |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_117, refute_0_112])).

cnf(refute_0_119, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) + exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <
     0 |
     0 <
     4 + skoXC1 * (7 + skoXC1 * 3) + exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     0 <= -8/3 + skoXC1 * -12 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(split, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (-4 + skoXC1 * (-7 + skoXC1 * -3) <
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3) |
     10 + skoXC1 * 4 <= exp(skoXC1) | skoXC1 <= -2/9),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain, (-2/9 < skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_72, refute_0_25, refute_0_26])).

cnf(refute_0_122, plain,
    (-4 + skoXC1 * (-7 + skoXC1 * -3) <
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (0 *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) < -8/3 + skoXC1 * -12 |
     4 + skoXC1 * (7 + skoXC1 * 3) +
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <= 0 |
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_124, plain,
    (0 <
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) |
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) < exp(skoXC1)),
    inference(subst, [], [refute_0_66])).

cnf(refute_0_125, plain,
    (10 + skoXC1 * 4 <= exp(skoXC1) |
     4 + skoXC1 * (7 + skoXC1 * 3) +
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <= 0 |
     exp(skoXC1) <=
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3))),
    inference(split, [], [refute_0_118])).

cnf(refute_0_126, plain,
    (exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     exp(skoXC1) <=
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3))),
    inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (0 <
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_126, refute_0_124])).

cnf(refute_0_128, plain,
    (0 *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) < -8/3 + skoXC1 * -12 |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) | 10 + skoXC1 * 4 <= exp(skoXC1) |
     4 + skoXC1 * (7 + skoXC1 * 3) +
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 < -2/9 |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain, (-2/9 <= skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_26, refute_0_25, refute_0_72])).

cnf(refute_0_131, plain,
    (exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_130, refute_0_129])).

cnf(refute_0_132, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_131, refute_0_122])).

cnf(refute_0_133, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1) |
     1/3 + skoXC1 * 1/3 <= 0 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(split, [], [refute_0_132])).

cnf(refute_0_134, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1) |
     10 + skoXC1 * 4 <= exp(skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_134, refute_0_79])).

cnf(refute_0_136, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_6, refute_0_135])).

cnf(refute_0_137, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_105, refute_0_136])).

cnf(refute_0_138, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (1/3 + skoXC1 * 1/3) | 1/3 + skoXC1 * 1/3 <= 0 |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_1, refute_0_137])).

cnf(refute_0_139, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     (7247757312 + skoXC1 * 7247757312) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     10 + skoXC1 * 4 <= exp(skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     (7247757312 + skoXC1 * 7247757312) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_139, refute_0_79])).

cnf(refute_0_141, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     7247757312 + skoXC1 * 7247757312 | 10 + skoXC1 * 4 <= exp(skoXC1) |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_140])).

cnf(refute_0_142, plain,
    (79725330432 <
     skoXC1 *
     (-57982058496 +
      skoXC1 *
      (43486543872 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-10815012864 +
         skoXC1 *
         (5308416000 +
          skoXC1 *
          (-1588985856 +
           skoXC1 *
           (342392832 +
            skoXC1 *
            (-56484864 +
             skoXC1 *
             (7298304 +
              skoXC1 *
              (-740736 +
               skoXC1 *
               (58272 +
                skoXC1 *
                (-3412 + skoXC1 * (137 + skoXC1 * -3))))))))))))) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (skoXC1 *
     (-57982058496 +
      skoXC1 *
      (43486543872 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-10815012864 +
         skoXC1 *
         (5308416000 +
          skoXC1 *
          (-1588985856 +
           skoXC1 *
           (342392832 +
            skoXC1 *
            (-56484864 +
             skoXC1 *
             (7298304 +
              skoXC1 *
              (-740736 +
               skoXC1 *
               (58272 +
                skoXC1 *
                (-3412 + skoXC1 * (137 + skoXC1 * -3))))))))))))) <=
     79725330432 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(decision, [], [refute_0_72, refute_0_25, refute_0_26])).

cnf(refute_0_144, plain,
    (skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 | 10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_143, refute_0_142])).

cnf(refute_0_145, plain, (10 + skoXC1 * 4 <= exp(skoXC1)),
    inference(resolve, [], [refute_0_144, refute_0_111])).

cnf(refute_0_146, plain, (10 + skoXC1 * 4 < exp(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_145, refute_0_62])).

cnf(refute_0_147, plain,
    (10 + skoXC1 * 4 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1 | 1 <= skoXC1),
    inference(resolve, [], [refute_0_42, refute_0_146])).

cnf(refute_0_148, plain, (1 <= skoXC1 | skoXC1 <= 3),
    inference(decision, [], [refute_0_26, refute_0_25])).

cnf(refute_0_149, plain,
    (10 + skoXC1 * 4 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_148, refute_0_147])).

cnf(refute_0_150, plain,
    ((10 + skoXC1 * 4) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 | 1 <= skoXC1 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_41, refute_0_149])).

cnf(refute_0_151, plain,
    (195689447424 <
     skoXC1 *
     (130459631616 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (-7247757312 +
        skoXC1 *
        (5577375744 +
         skoXC1 *
         (-1868562432 +
          skoXC1 *
          (422903808 +
           skoXC1 *
           (-71663616 +
            skoXC1 *
            (9414144 +
             skoXC1 *
             (-966144 +
              skoXC1 *
              (76608 +
               skoXC1 * (-4512 + skoXC1 * (182 + skoXC1 * -4)))))))))))) |
     1 <= skoXC1 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (1 <= skoXC1 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 |
     skoXC1 *
     (130459631616 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (-7247757312 +
        skoXC1 *
        (5577375744 +
         skoXC1 *
         (-1868562432 +
          skoXC1 *
          (422903808 +
           skoXC1 *
           (-71663616 +
            skoXC1 *
            (9414144 +
             skoXC1 *
             (-966144 +
              skoXC1 *
              (76608 +
               skoXC1 * (-4512 + skoXC1 * (182 + skoXC1 * -4)))))))))))) <=
     195689447424), inference(decision, [], [refute_0_26, refute_0_25])).

cnf(refute_0_153, plain,
    (1 <= skoXC1 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_152, refute_0_151])).

cnf(refute_0_154, plain,
    (-21743271936 <
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 + skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) |
     1 <= skoXC1), inference(decision, [], [refute_0_26, refute_0_25])).

cnf(refute_0_155, plain, (1 <= skoXC1),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain, (-1 <= skoXC1),
    inference(decision, [], [refute_0_25, refute_0_26, refute_0_155])).

cnf(refute_0_157, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (exp(skoXC1) * 3)),
    inference(resolve, [], [refute_0_156, refute_0_40])).

cnf(refute_0_158, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     (1 + skoXC1) * (exp(skoXC1) * 3) <= 2 | exp(skoXC1) * 3 <= 0),
    inference(resolve, [], [refute_0_157, refute_0_35])).

cnf(refute_0_159, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) * (3 + skoXC1 * 3) <= 2),
    inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | 3 + skoXC1 * 3 <= 0 |
     exp(skoXC1) <= 2 / (3 + skoXC1 * 3)),
    inference(split, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) <= 2 / (3 + skoXC1 * 3) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain, (-1 < skoXC1),
    inference(decision, [], [refute_0_25, refute_0_26, refute_0_155])).

cnf(refute_0_163, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     exp(skoXC1) <= 2 / (3 + skoXC1 * 3)),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (skoXC1 < -1 | 2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (3 + skoXC1 * 3)),
    inference(resolve, [], [refute_0_163, refute_0_34])).

cnf(refute_0_165, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     1 + skoXC1 <= 2 / (3 + skoXC1 * 3)),
    inference(resolve, [], [refute_0_156, refute_0_164])).

cnf(refute_0_166, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     (1 + skoXC1) * (3 + skoXC1 * 3) <= 2 | 3 + skoXC1 * 3 <= 0),
    inference(resolve, [], [refute_0_165, refute_0_33])).

cnf(refute_0_167, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) |
     skoXC1 * (6 + skoXC1 * 3) <= -1 | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain, (-1 < skoXC1 * (6 + skoXC1 * 3) | skoXC1 <= -1),
    inference(decision, [], [refute_0_25, refute_0_26, refute_0_155])).

cnf(refute_0_169, plain,
    (2 < exp(skoXC1) * (exp(skoXC1) * 3) | skoXC1 <= -1),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain, (2 < exp(skoXC1) * (exp(skoXC1) * 3)),
    inference(resolve, [], [refute_0_169, refute_0_162])).

cnf(refute_0_171, plain,
    (exp(skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2),
    inference(resolve, [], [refute_0_32, refute_0_170])).

cnf(refute_0_172, plain,
    (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0 |
     exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(split, [], [refute_0_171])).

cnf(refute_0_173, plain,
    (exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= -14 |
     exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(arithmetic, [], [refute_0_172])).

cnf(refute_0_174, plain, (skoXC1 <= 3),
    inference(decision, [], [refute_0_155, refute_0_26, refute_0_25])).

cnf(refute_0_175, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     12 + skoXC1 * 9 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))))),
    inference(resolve, [], [refute_0_174, refute_0_104])).

cnf(refute_0_176, plain,
    ((12 + skoXC1 * 9) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     21743271936 |
     -14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_91, refute_0_175])).

cnf(refute_0_177, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     239175991296 <
     skoXC1 *
     (65229815808 +
      skoXC1 *
      (65229815808 +
       skoXC1 *
       (-54358179840 +
        skoXC1 *
        (21913141248 +
         skoXC1 *
         (-5987893248 +
          skoXC1 *
          (1220935680 +
           skoXC1 *
           (-193757184 +
            skoXC1 *
            (24302592 +
             skoXC1 *
             (-2407680 +
              skoXC1 *
              (185472 +
               skoXC1 * (-10656 + skoXC1 * (420 + skoXC1 * -9)))))))))))) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936 |
     skoXC1 *
     (65229815808 +
      skoXC1 *
      (65229815808 +
       skoXC1 *
       (-54358179840 +
        skoXC1 *
        (21913141248 +
         skoXC1 *
         (-5987893248 +
          skoXC1 *
          (1220935680 +
           skoXC1 *
           (-193757184 +
            skoXC1 *
            (24302592 +
             skoXC1 *
             (-2407680 +
              skoXC1 *
              (185472 +
               skoXC1 *
               (-10656 + skoXC1 * (420 + skoXC1 * -9)))))))))))) <=
     239175991296),
    inference(decision, [], [refute_0_155, refute_0_26, refute_0_25])).

cnf(refute_0_179, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_178, refute_0_177])).

cnf(refute_0_180, plain,
    (-21743271936 <
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 + skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))),
    inference(decision, [], [refute_0_155, refute_0_26, refute_0_25])).

cnf(refute_0_181, plain,
    (-14 < exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (exp(skoXC1) <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(resolve, [], [refute_0_173, refute_0_181])).

cnf(refute_0_183, plain,
    (skoXC1 < -1 |
     1 + skoXC1 <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(resolve, [], [refute_0_182, refute_0_12])).

cnf(refute_0_184, plain,
    (1 + skoXC1 <=
     (34/3 + skoXC1 * 2) /
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3))),
    inference(resolve, [], [refute_0_156, refute_0_183])).

cnf(refute_0_185, plain,
    ((1 + skoXC1) *
     (14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3)) <=
     34/3 + skoXC1 * 2 |
     14 + exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= 0),
    inference(resolve, [], [refute_0_184, refute_0_7])).

cnf(refute_0_186, plain,
    (exp(skoXC1) * (4 + skoXC1 * 3 + exp(skoXC1) * -1/3) <= -14 |
     exp(skoXC1) *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) <= -8/3 + skoXC1 * -12),
    inference(arithmetic, [], [refute_0_185])).

cnf(refute_0_187, plain,
    (exp(skoXC1) *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) <= -8/3 + skoXC1 * -12),
    inference(resolve, [], [refute_0_186, refute_0_181])).

cnf(refute_0_188, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) + exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <
     0 |
     0 <
     4 + skoXC1 * (7 + skoXC1 * 3) + exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     0 <= -8/3 + skoXC1 * -12), inference(split, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (-4 + skoXC1 * (-7 + skoXC1 * -3) <
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3) |
     skoXC1 <= -2/9), inference(arithmetic, [], [refute_0_188])).

cnf(refute_0_190, plain, (-2/9 < skoXC1),
    inference(decision, [], [refute_0_155, refute_0_26, refute_0_25])).

cnf(refute_0_191, plain,
    (-4 + skoXC1 * (-7 + skoXC1 * -3) <
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) |
     4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3)),
    inference(resolve, [], [refute_0_189, refute_0_190])).

cnf(refute_0_192, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) +
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <= 0 |
     exp(skoXC1) <=
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3))),
    inference(split, [], [refute_0_187])).

cnf(refute_0_193, plain,
    (exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) |
     exp(skoXC1) <=
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3))),
    inference(arithmetic, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (0 <
     (-8/3 + skoXC1 * -12) /
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3)),
    inference(resolve, [], [refute_0_193, refute_0_124])).

cnf(refute_0_195, plain,
    (0 *
     (4 + skoXC1 * (7 + skoXC1 * 3) +
      exp(skoXC1) * (-1/3 + skoXC1 * -1/3)) < -8/3 + skoXC1 * -12 |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3) |
     4 + skoXC1 * (7 + skoXC1 * 3) +
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <= 0),
    inference(resolve, [], [refute_0_123, refute_0_194])).

cnf(refute_0_196, plain,
    (skoXC1 < -2/9 |
     exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3)),
    inference(arithmetic, [], [refute_0_195])).

cnf(refute_0_197, plain, (-2/9 <= skoXC1),
    inference(decision, [], [refute_0_25, refute_0_26, refute_0_155])).

cnf(refute_0_198, plain,
    (exp(skoXC1) * (-1/3 + skoXC1 * -1/3) <=
     -4 + skoXC1 * (-7 + skoXC1 * -3)),
    inference(resolve, [], [refute_0_197, refute_0_196])).

cnf(refute_0_199, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) < exp(skoXC1) * (1/3 + skoXC1 * 1/3)),
    inference(resolve, [], [refute_0_198, refute_0_191])).

cnf(refute_0_200, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1) |
     1/3 + skoXC1 * 1/3 <= 0), inference(split, [], [refute_0_199])).

cnf(refute_0_201, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) < exp(skoXC1)),
    inference(resolve, [], [refute_0_201, refute_0_162])).

cnf(refute_0_203, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     3 < skoXC1), inference(resolve, [], [refute_0_6, refute_0_202])).

cnf(refute_0_204, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) / (1/3 + skoXC1 * 1/3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))))),
    inference(resolve, [], [refute_0_174, refute_0_203])).

cnf(refute_0_205, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) *
     (1/3 + skoXC1 * 1/3) | 1/3 + skoXC1 * 1/3 <= 0),
    inference(resolve, [], [refute_0_1, refute_0_204])).

cnf(refute_0_206, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     (7247757312 + skoXC1 * 7247757312) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    (4 + skoXC1 * (7 + skoXC1 * 3) <
     (7247757312 + skoXC1 * 7247757312) /
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))))),
    inference(resolve, [], [refute_0_206, refute_0_162])).

cnf(refute_0_208, plain,
    ((4 + skoXC1 * (7 + skoXC1 * 3)) *
     (21743271936 +
      skoXC1 *
      (-21743271936 +
       skoXC1 *
       (10871635968 +
        skoXC1 *
        (-3623878656 +
         skoXC1 *
         (891813888 +
          skoXC1 *
          (-169869312 +
           skoXC1 *
           (25657344 +
            skoXC1 *
            (-3096576 +
             skoXC1 *
             (297216 +
              skoXC1 *
              (-22272 +
               skoXC1 * (1248 + skoXC1 * (-48 + skoXC1)))))))))))) <
     7247757312 + skoXC1 * 7247757312 |
     21743271936 +
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_207])).

cnf(refute_0_209, plain,
    (79725330432 <
     skoXC1 *
     (-57982058496 +
      skoXC1 *
      (43486543872 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-10815012864 +
         skoXC1 *
         (5308416000 +
          skoXC1 *
          (-1588985856 +
           skoXC1 *
           (342392832 +
            skoXC1 *
            (-56484864 +
             skoXC1 *
             (7298304 +
              skoXC1 *
              (-740736 +
               skoXC1 *
               (58272 +
                skoXC1 *
                (-3412 + skoXC1 * (137 + skoXC1 * -3))))))))))))) |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (skoXC1 *
     (-57982058496 +
      skoXC1 *
      (43486543872 +
       skoXC1 *
       (3623878656 +
        skoXC1 *
        (-10815012864 +
         skoXC1 *
         (5308416000 +
          skoXC1 *
          (-1588985856 +
           skoXC1 *
           (342392832 +
            skoXC1 *
            (-56484864 +
             skoXC1 *
             (7298304 +
              skoXC1 *
              (-740736 +
               skoXC1 *
               (58272 +
                skoXC1 *
                (-3412 + skoXC1 * (137 + skoXC1 * -3))))))))))))) <=
     79725330432 |
     skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936),
    inference(decision, [], [refute_0_155, refute_0_26, refute_0_25])).

cnf(refute_0_211, plain,
    (skoXC1 *
     (-21743271936 +
      skoXC1 *
      (10871635968 +
       skoXC1 *
       (-3623878656 +
        skoXC1 *
        (891813888 +
         skoXC1 *
         (-169869312 +
          skoXC1 *
          (25657344 +
           skoXC1 *
           (-3096576 +
            skoXC1 *
            (297216 +
             skoXC1 *
             (-22272 +
              skoXC1 * (1248 + skoXC1 * (-48 + skoXC1))))))))))) <=
     -21743271936), inference(resolve, [], [refute_0_210, refute_0_209])).

cnf(refute_0_212, plain, ($false),
    inference(resolve, [], [refute_0_211, refute_0_180])).
% SZS output end CNFRefutation for log-nest-exp-upper2.tptp
