% SZS output start CNFRefutation for log-3615-sqrt.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/3 * (10 * X ^ 2 + 19 * X + 1) * (X - 1) /
            (X * (X ^ 2 + 6 * X + 3))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/12 * (47 * X ^ 3 + 239 * X ^ 2 + 131 * X + 3) * (X - 1) /
            (X * (X ^ 3 + 12 * X ^ 2 + 18 * X + 4))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/12 * (3 * X ^ 3 + 131 * X ^ 2 + 239 * X + 47) * (X - 1) /
            (4 * X ^ 3 + 18 * X ^ 2 + 12 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (6 * X ^ 4 + 481 * X ^ 3 + 1881 * X ^ 2 + 1281 * X + 131) *
            (X - 1) / (5 * X ^ 4 + 40 * X ^ 3 + 60 * X ^ 2 + 20 * X + 1),
            Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (5 * X ^ 5 + 647 * X ^ 4 + 4397 * X ^ 3 + 6397 * X ^ 2 +
             2272 * X + 142) * (X - 1) /
            (6 * X ^ 5 + 75 * X ^ 4 + 200 * X ^ 3 + 150 * X ^ 2 + 30 * X +
             1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/70 *
            (353 + 1921 * X ^ 5 + 8389 * X + 38524 * X ^ 2 +
             20149 * X ^ 4 + 50774 * X ^ 3 + 10 * X ^ 6) * (X - 1) /
            (1 + 126 * X ^ 5 + 42 * X + 315 * X ^ 2 + 525 * X ^ 4 +
             700 * X ^ 3 + 7 * X ^ 6), Y) | lgen(R, ln(X), Y))).

cnf(ln_lower_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (142 * X ^ 5 + 2272 * X ^ 4 + 6397 * X ^ 3 + 4397 * X ^ 2 +
             647 * X + 5) * (X - 1) /
            (X *
             (X ^ 5 + 30 * X ^ 4 + 150 * X ^ 3 + 200 * X ^ 2 + 75 * X +
              6))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf13, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/70 *
            (10 + 1921 * X + 20149 * X ^ 2 + 50774 * X ^ 3 +
             38524 * X ^ 4 + 8389 * X ^ 5 + 353 * X ^ 6) * (X - 1) /
            (X *
             (7 + 126 * X + 525 * X ^ 2 + 700 * X ^ 3 + 315 * X ^ 4 +
              42 * X ^ 5 + X ^ 6))) | lgen(R, Y, ln(X)))).

cnf(sqrt_upper_bound_2_small, axiom,
    (X < 0 | 1 <= X |
     ~ lgen(R,
            1/32 *
            (65536 * X ^ 4 + 114688 * X ^ 3 + 17920 * X ^ 2 + 448 * X +
             1) / ((16 * X + 1) * (256 * X ^ 2 + 96 * X + 1)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

fof(log_3615, conjecture,
    (! [X] :
       ((10 ^ -3 < X & X != 1 & X < 1000) =>
        ln(X) / (X - 1) <= 1 / sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((10 ^ -3 < X & X != 1 & X < 1000) =>
        ln(X) / (X - 1) <= 1 / sqrt(X))),
    inference(strip, [], [log_3615])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((10 ^ -3 < X & X != 1 & X < 1000) =>
          ln(X) / (X - 1) <= 1 / sqrt(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (10 ^ -3 < X & X < 1000 & 1 / sqrt(X) < ln(X) / (X - 1) & X != 1)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 1000 & 1 / sqrt(skoXC1) < ln(skoXC1) / (skoXC1 - 1) &
     10 ^ -3 < skoXC1 & skoXC1 != 1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (1 / sqrt(skoXC1) < ln(skoXC1) / (skoXC1 - 1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (10 ^ -3 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 != 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (skoXC1 < 1000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/224 +
     skoXC1 *
     (913/320 +
      skoXC1 *
      (37627/80 +
       skoXC1 *
       (6226967/320 +
        skoXC1 *
        (280088819/1120 +
         skoXC1 *
         (78830115/64 +
          skoXC1 *
          (159781489/80 +
           skoXC1 *
           (58492319/2240 +
            skoXC1 *
            (-9828353/5 +
             skoXC1 *
             (-6479656/5 +
              skoXC1 * (-253184 + skoXC1 * -361472/35)))))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0 |
     (1/224 +
      skoXC1 *
      (913/320 +
       skoXC1 *
       (37627/80 +
        skoXC1 *
        (6226967/320 +
         skoXC1 *
         (280088819/1120 +
          skoXC1 *
          (78830115/64 +
           skoXC1 *
           (159781489/80 +
            skoXC1 *
            (58492319/2240 +
             skoXC1 *
             (-9828353/5 +
              skoXC1 *
              (-6479656/5 +
               skoXC1 * (-253184 + skoXC1 * -361472/35))))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1)))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) |
     1/7 +
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))) <=
     0 |
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-119 +
       skoXC1 *
       (-399 +
        skoXC1 *
        (-175 +
         skoXC1 * (385 + skoXC1 * (273 + skoXC1 * (41 + skoXC1))))))) |
     skoXC1 *
     (-7 +
      skoXC1 *
      (-119 +
       skoXC1 *
       (-399 +
        skoXC1 *
        (-175 +
         skoXC1 * (385 + skoXC1 * (273 + skoXC1 * (41 + skoXC1))))))) /
     sqrt(skoXC1) <=
     -1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) |
     sqrt(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_3, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0 |
     (-1 + skoXC1) / sqrt(skoXC1) <=
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_4, plain, (1 / sqrt(skoXC1) < ln(skoXC1) / (skoXC1 - 1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_5, plain, (1 / sqrt(skoXC1) < ln(skoXC1) / (-1 + skoXC1)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (ln(skoXC1) < 1 / sqrt(skoXC1) * (-1 + skoXC1) | 0 <= -1 + skoXC1 |
     ln(skoXC1) / (-1 + skoXC1) <= 1 / sqrt(skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_7, plain,
    (ln(skoXC1) < 1 / sqrt(skoXC1) * (-1 + skoXC1) | 0 <= -1 + skoXC1),
    inference(resolve, [], [refute_0_6, refute_0_5])).

cnf(refute_0_8, plain,
    (ln(skoXC1) < (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000147, ln(X_000146)) | X_000147 <= ln(X_000146)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (~ lgen(0, X_000147,
            1/70 *
            (10 + 1921 * X_000146 + 20149 * X_000146 ^ 2 +
             50774 * X_000146 ^ 3 + 38524 * X_000146 ^ 4 +
             8389 * X_000146 ^ 5 + 353 * X_000146 ^ 6) * (X_000146 - 1) /
            (X_000146 *
             (7 + 126 * X_000146 + 525 * X_000146 ^ 2 +
              700 * X_000146 ^ 3 + 315 * X_000146 ^ 4 + 42 * X_000146 ^ 5 +
              X_000146 ^ 6))) | X_000146 <= 0 |
     lgen(0, X_000147, ln(X_000146))),
    inference(subst, [], [ln_lower_bound_cf13])).

cnf(refute_0_11, plain,
    (~ lgen(0, X_000147,
            1/70 *
            (10 + 1921 * X_000146 + 20149 * X_000146 ^ 2 +
             50774 * X_000146 ^ 3 + 38524 * X_000146 ^ 4 +
             8389 * X_000146 ^ 5 + 353 * X_000146 ^ 6) * (X_000146 - 1) /
            (X_000146 *
             (7 + 126 * X_000146 + 525 * X_000146 ^ 2 +
              700 * X_000146 ^ 3 + 315 * X_000146 ^ 4 + 42 * X_000146 ^ 5 +
              X_000146 ^ 6))) | X_000146 <= 0 | X_000147 <= ln(X_000146)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    ((-1/7 +
      X_000146 *
      (-273/10 +
       X_000146 *
       (-1302/5 +
        X_000146 *
        (-875/2 +
         X_000146 *
         (175 +
          X_000146 *
          (861/2 + X_000146 * (574/5 + X_000146 * 353/70))))))) /
     (X_000146 *
      (7 +
       X_000146 *
       (126 +
        X_000146 *
        (525 +
         X_000146 *
         (700 + X_000146 * (315 + X_000146 * (42 + X_000146))))))) <
     X_000147 | X_000146 <= 0 | X_000147 <= ln(X_000146)),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (-1 + skoXC1) / sqrt(skoXC1) |
     (-1 + skoXC1) / sqrt(skoXC1) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) <
     -361472/35 +
     skoXC1 *
     (-253184 +
      skoXC1 *
      (-6479656/5 +
       skoXC1 *
       (-9828353/5 +
        skoXC1 *
        (58492319/2240 +
         skoXC1 *
         (159781489/80 +
          skoXC1 *
          (78830115/64 +
           skoXC1 *
           (280088819/1120 +
            skoXC1 *
            (6226967/320 +
             skoXC1 *
             (37627/80 + skoXC1 * (913/320 + skoXC1 * 1/224)))))))))) |
     1 +
     skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) <=
     0 |
     (-361472/35 +
      skoXC1 *
      (-253184 +
       skoXC1 *
       (-6479656/5 +
        skoXC1 *
        (-9828353/5 +
         skoXC1 *
         (58492319/2240 +
          skoXC1 *
          (159781489/80 +
           skoXC1 *
           (78830115/64 +
            skoXC1 *
            (280088819/1120 +
             skoXC1 *
             (6226967/320 +
              skoXC1 *
              (37627/80 + skoXC1 * (913/320 + skoXC1 * 1/224))))))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) <=
     -4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_16, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 |
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_17, plain,
    ((-2048 +
      skoXC1 *
      (-3584 + skoXC1 * (-560 + skoXC1 * (-14 + skoXC1 * -1/32)))) *
     ln(skoXC1) <
     4096 +
     skoXC1 * (-2304 + skoXC1 * (-1680 + skoXC1 * (-111 + skoXC1 * -1))) |
     (4096 +
      skoXC1 *
      (-2304 + skoXC1 * (-1680 + skoXC1 * (-111 + skoXC1 * -1)))) /
     ln(skoXC1) <=
     -2048 +
     skoXC1 * (-3584 + skoXC1 * (-560 + skoXC1 * (-14 + skoXC1 * -1/32))) |
     ln(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_18, plain,
    ((-1 + skoXC1) / ln(skoXC1) *
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) <
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) |
     4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1)) <= 0 |
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) /
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) <=
     (-1 + skoXC1) / ln(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_19, plain,
    (ln(skoXC1) * -1 * sqrt(skoXC1) < 1 + skoXC1 * -1 |
     (1 + skoXC1 * -1) / sqrt(skoXC1) <= ln(skoXC1) * -1 |
     sqrt(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_20, plain,
    (1 / sqrt(skoXC1) * (-1 + skoXC1) < ln(skoXC1) | -1 + skoXC1 <= 0 |
     ln(skoXC1) / (-1 + skoXC1) <= 1 / sqrt(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_21, plain,
    (1 / sqrt(skoXC1) * (-1 + skoXC1) < ln(skoXC1) | -1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_20, refute_0_5])).

cnf(refute_0_22, plain,
    (ln(skoXC1) * -1 < (1 + skoXC1 * -1) / sqrt(skoXC1) | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (ln(skoXC1) * -1 * sqrt(skoXC1) < 1 + skoXC1 * -1 | skoXC1 <= 1 |
     sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_19, refute_0_22])).

cnf(refute_0_24, plain,
    (-1 + skoXC1 < sqrt(skoXC1) * ln(skoXC1) | skoXC1 <= 0 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain, (10 ^ -3 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_26, plain, (1/1000 < skoXC1),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain, (skoXC1 != 1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_28, plain, (skoXC1 < 1000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_29, plain, (0 < skoXC1 | skoXC1 <= 1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (-1 + skoXC1 < sqrt(skoXC1) * ln(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_24, refute_0_29])).

cnf(refute_0_31, plain,
    ((-1 + skoXC1) / ln(skoXC1) < sqrt(skoXC1) | ln(skoXC1) <= 0 |
     skoXC1 <= 1), inference(split, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((-1 + skoXC1) / ln(skoXC1) < sqrt(skoXC1) | skoXC1 <= 0 |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_31])).

cnf(refute_0_33, plain,
    ((-1 + skoXC1) / ln(skoXC1) < sqrt(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_32, refute_0_29])).

cnf(refute_0_34, plain,
    (~ lgen(0, sqrt(X_000090), X_000091) | sqrt(X_000090) <= X_000091),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_35, plain,
    (X_000090 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000090 + 17920 * X_000090 ^ 2 +
             448 * X_000090 ^ 3 + X_000090 ^ 4) /
            ((X_000090 + 16) * (256 + 96 * X_000090 + X_000090 ^ 2)),
            X_000091) | lgen(0, sqrt(X_000090), X_000091)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_36, plain,
    (X_000090 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000090 + 17920 * X_000090 ^ 2 +
             448 * X_000090 ^ 3 + X_000090 ^ 4) /
            ((X_000090 + 16) * (256 + 96 * X_000090 + X_000090 ^ 2)),
            X_000091) | sqrt(X_000090) <= X_000091),
    inference(resolve, [], [refute_0_35, refute_0_34])).

cnf(refute_0_37, plain,
    (X_000090 < 1 |
     X_000091 <
     (2048 +
      X_000090 *
      (3584 + X_000090 * (560 + X_000090 * (14 + X_000090 * 1/32)))) /
     (4096 + X_000090 * (1792 + X_000090 * (112 + X_000090))) |
     sqrt(X_000090) <= X_000091),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 < 1 |
     (-1 + skoXC1) / ln(skoXC1) <
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) /
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) |
     sqrt(skoXC1) <= (-1 + skoXC1) / ln(skoXC1)),
    inference(subst, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 < 1 |
     (-1 + skoXC1) / ln(skoXC1) <
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) /
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_38, refute_0_33])).

cnf(refute_0_40, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * (6 + skoXC1 * 3)) <
     -5/12 +
     skoXC1 *
     (-289/24 +
      skoXC1 *
      (-467/12 +
       skoXC1 *
       (-135/8 +
        skoXC1 *
        (171/4 + skoXC1 * (565/24 + skoXC1 * (23/12 + skoXC1 * 1/24)))))) |
     1 + skoXC1 * (6 + skoXC1 * 3) <= 0 |
     (-5/12 +
      skoXC1 *
      (-289/24 +
       skoXC1 *
       (-467/12 +
        skoXC1 *
        (-135/8 +
         skoXC1 *
         (171/4 +
          skoXC1 * (565/24 + skoXC1 * (23/12 + skoXC1 * 1/24))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) <=
     -1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 |
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_42, plain,
    (~ lgen(0, ln(X_000069), X_000070) | ln(X_000069) <= X_000070),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_43, plain,
    (~ lgen(0,
            1/3 * (X_000069 ^ 2 + 19 * X_000069 + 10) * (X_000069 - 1) /
            (3 * X_000069 ^ 2 + 6 * X_000069 + 1), X_000070) |
     X_000069 <= 0 | lgen(0, ln(X_000069), X_000070)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_44, plain,
    (~ lgen(0,
            1/3 * (X_000069 ^ 2 + 19 * X_000069 + 10) * (X_000069 - 1) /
            (3 * X_000069 ^ 2 + 6 * X_000069 + 1), X_000070) |
     X_000069 <= 0 | ln(X_000069) <= X_000070),
    inference(resolve, [], [refute_0_43, refute_0_42])).

cnf(refute_0_45, plain,
    (X_000070 <
     (-10/3 + X_000069 * (-3 + X_000069 * (6 + X_000069 * 1/3))) /
     (1 + X_000069 * (6 + X_000069 * 3)) | X_000069 <= 0 |
     ln(X_000069) <= X_000070), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) |
     ln(skoXC1) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    ((-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) *
     ln(skoXC1) <
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     ln(skoXC1) <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     ln(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_48, plain,
    ((-1 + skoXC1) / ln(skoXC1) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0 |
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     (-1 + skoXC1) / ln(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_49, plain,
    (~ lgen(0, sqrt(X_000083), X_000084) | sqrt(X_000083) <= X_000084),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_50, plain,
    (X_000083 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000083 ^ 4 + 28 * X_000083 ^ 3 + 70 * X_000083 ^ 2 +
             28 * X_000083 + 1) /
            ((X_000083 + 1) * (X_000083 ^ 2 + 6 * X_000083 + 1)),
            X_000084) | lgen(0, sqrt(X_000083), X_000084)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_51, plain,
    (X_000083 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000083 ^ 4 + 28 * X_000083 ^ 3 + 70 * X_000083 ^ 2 +
             28 * X_000083 + 1) /
            ((X_000083 + 1) * (X_000083 ^ 2 + 6 * X_000083 + 1)),
            X_000084) | sqrt(X_000083) <= X_000084),
    inference(resolve, [], [refute_0_50, refute_0_49])).

cnf(refute_0_52, plain,
    (X_000083 < 0 |
     X_000084 <
     (1/8 +
      X_000083 *
      (7/2 + X_000083 * (35/4 + X_000083 * (7/2 + X_000083 * 1/8)))) /
     (1 + X_000083 * (7 + X_000083 * (7 + X_000083))) |
     sqrt(X_000083) <= X_000084),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (skoXC1 < 0 |
     (-1 + skoXC1) / ln(skoXC1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) |
     sqrt(skoXC1) <= (-1 + skoXC1) / ln(skoXC1)),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < 0 |
     (-1 + skoXC1) / ln(skoXC1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_53, refute_0_33])).

cnf(refute_0_55, plain, (0 <= skoXC1 | skoXC1 <= 1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_56, plain,
    ((-1 + skoXC1) / ln(skoXC1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_55, refute_0_54])).

cnf(refute_0_57, plain,
    ((-1 + skoXC1) / ln(skoXC1) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_48, refute_0_56])).

cnf(refute_0_58, plain,
    (-1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     ln(skoXC1) | skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= -1 |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (-1 < skoXC1 * (7 + skoXC1 * (7 + skoXC1)) | skoXC1 <= 1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_60, plain,
    (-1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     ln(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    ((-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) *
     ln(skoXC1) <
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) |
     ln(skoXC1) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_47, refute_0_60])).

cnf(refute_0_62, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     ln(skoXC1) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     skoXC1 <= 0 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     ln(skoXC1) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_62, refute_0_29])).

cnf(refute_0_64, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     ln(skoXC1) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 | skoXC1 <= 1), inference(split, [], [refute_0_63])).

cnf(refute_0_65, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     ln(skoXC1) |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/24 +
     skoXC1 *
     (23/12 +
      skoXC1 *
      (565/24 +
       skoXC1 *
       (171/4 +
        skoXC1 *
        (-135/8 +
         skoXC1 * (-467/12 + skoXC1 * (-289/24 + skoXC1 * -5/12)))))) |
     1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0 |
     (1/24 +
      skoXC1 *
      (23/12 +
       skoXC1 *
       (565/24 +
        skoXC1 *
        (171/4 +
         skoXC1 *
         (-135/8 +
          skoXC1 * (-467/12 + skoXC1 * (-289/24 + skoXC1 * -5/12))))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_67, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) *
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) |
     1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) <= 0 |
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_68, plain,
    ((-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) * sqrt(skoXC1) <
     skoXC1 * (-3 + skoXC1 * (-3 + skoXC1 * (5 + skoXC1))) |
     skoXC1 * (-3 + skoXC1 * (-3 + skoXC1 * (5 + skoXC1))) /
     sqrt(skoXC1) <= -1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3)) |
     sqrt(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_69, plain,
    (-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3)) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1))) |
     skoXC1 * (3 + skoXC1 * (6 + skoXC1)) <= 0 |
     (-1 + skoXC1) / sqrt(skoXC1) <=
     (-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) /
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_70, plain,
    (~ lgen(0, X_000077, ln(X_000076)) | X_000077 <= ln(X_000076)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_71, plain,
    (~ lgen(0, X_000077,
            1/3 * (10 * X_000076 ^ 2 + 19 * X_000076 + 1) *
            (X_000076 - 1) /
            (X_000076 * (X_000076 ^ 2 + 6 * X_000076 + 3))) |
     X_000076 <= 0 | lgen(0, X_000077, ln(X_000076))),
    inference(subst, [], [ln_lower_bound_cf5])).

cnf(refute_0_72, plain,
    (~ lgen(0, X_000077,
            1/3 * (10 * X_000076 ^ 2 + 19 * X_000076 + 1) *
            (X_000076 - 1) /
            (X_000076 * (X_000076 ^ 2 + 6 * X_000076 + 3))) |
     X_000076 <= 0 | X_000077 <= ln(X_000076)),
    inference(resolve, [], [refute_0_71, refute_0_70])).

cnf(refute_0_73, plain,
    ((-1/3 + X_000076 * (-6 + X_000076 * (3 + X_000076 * 10/3))) /
     (X_000076 * (3 + X_000076 * (6 + X_000076))) < X_000077 |
     X_000076 <= 0 | X_000077 <= ln(X_000076)),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    ((-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) /
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1))) <
     (-1 + skoXC1) / sqrt(skoXC1) |
     (-1 + skoXC1) / sqrt(skoXC1) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_73])).

cnf(refute_0_75, plain,
    ((-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) /
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_74, refute_0_8])).

cnf(refute_0_76, plain, (0 < skoXC1 | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_77, plain,
    ((-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) /
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3)) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 * (3 + skoXC1 * (6 + skoXC1))) | 1 <= skoXC1 |
     skoXC1 * (3 + skoXC1 * (6 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_69, refute_0_77])).

cnf(refute_0_79, plain,
    (-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3)) <
     skoXC1 * (-3 + skoXC1 * (-3 + skoXC1 * (5 + skoXC1))) / sqrt(skoXC1) |
     1 <= skoXC1 | skoXC1 * (3 + skoXC1 * (6 + skoXC1)) <= 0),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (0 < skoXC1 * (3 + skoXC1 * (6 + skoXC1)) | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_81, plain,
    (-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3)) <
     skoXC1 * (-3 + skoXC1 * (-3 + skoXC1 * (5 + skoXC1))) / sqrt(skoXC1) |
     1 <= skoXC1), inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    ((-1/3 + skoXC1 * (-6 + skoXC1 * (3 + skoXC1 * 10/3))) * sqrt(skoXC1) <
     skoXC1 * (-3 + skoXC1 * (-3 + skoXC1 * (5 + skoXC1))) | 1 <= skoXC1 |
     sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_68, refute_0_81])).

cnf(refute_0_83, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     sqrt(skoXC1) * (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) |
     1 <= skoXC1 | skoXC1 <= 0), inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     sqrt(skoXC1) * (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_83, refute_0_76])).

cnf(refute_0_85, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) < sqrt(skoXC1) |
     1 <= skoXC1 |
     1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) <= 0),
    inference(split, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) < sqrt(skoXC1) |
     1 <= skoXC1 | skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) <= -1/3),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (-1/3 < skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_88, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) < sqrt(skoXC1) |
     1 <= skoXC1), inference(resolve, [], [refute_0_86, refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 < 0 |
     skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) |
     sqrt(skoXC1) <=
     skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)))),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_90, plain,
    (skoXC1 < 0 |
     skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_89, refute_0_88])).

cnf(refute_0_91, plain, (0 <= skoXC1 | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_92, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) /
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_90])).

cnf(refute_0_93, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) *
     (1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3))) | 1 <= skoXC1 |
     1/3 + skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) <= 0),
    inference(resolve, [], [refute_0_67, refute_0_92])).

cnf(refute_0_94, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     (1/24 +
      skoXC1 *
      (23/12 +
       skoXC1 *
       (565/24 +
        skoXC1 *
        (171/4 +
         skoXC1 *
         (-135/8 +
          skoXC1 * (-467/12 + skoXC1 * (-289/24 + skoXC1 * -5/12))))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1 |
     skoXC1 * (6 + skoXC1 * (-3 + skoXC1 * -10/3)) <= -1/3),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) <
     (1/24 +
      skoXC1 *
      (23/12 +
       skoXC1 *
       (565/24 +
        skoXC1 *
        (171/4 +
         skoXC1 *
         (-135/8 +
          skoXC1 * (-467/12 + skoXC1 * (-289/24 + skoXC1 * -5/12))))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_94, refute_0_87])).

cnf(refute_0_96, plain,
    (skoXC1 * (3 + skoXC1 * (3 + skoXC1 * (-5 + skoXC1 * -1))) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/24 +
     skoXC1 *
     (23/12 +
      skoXC1 *
      (565/24 +
       skoXC1 *
       (171/4 +
        skoXC1 *
        (-135/8 +
         skoXC1 * (-467/12 + skoXC1 * (-289/24 + skoXC1 * -5/12)))))) |
     1 <= skoXC1 | 1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_66, refute_0_95])).

cnf(refute_0_97, plain,
    (-1/24 <
     skoXC1 *
     (-13/12 +
      skoXC1 *
      (-11/24 +
       skoXC1 *
       (23/4 +
        skoXC1 *
        (-39/8 + skoXC1 * (1/12 + skoXC1 * (-1/24 + skoXC1 * 7/12)))))) |
     1 <= skoXC1 | skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= -1),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/16 +
     skoXC1 *
     (2 +
      skoXC1 *
      (177/16 +
       skoXC1 *
       (21/2 + skoXC1 * (-237/16 + skoXC1 * (-17/2 + skoXC1 * -5/16))))) |
     1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0 |
     (1/16 +
      skoXC1 *
      (2 +
       skoXC1 *
       (177/16 +
        skoXC1 *
        (21/2 +
         skoXC1 * (-237/16 + skoXC1 * (-17/2 + skoXC1 * -5/16)))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_99, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) *
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) |
     1/2 + skoXC1 * (2 + skoXC1 * -5/2) <= 0 |
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_100, plain,
    ((-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) * sqrt(skoXC1) <
     skoXC1 * (-2 + skoXC1 * (1 + skoXC1)) |
     skoXC1 * (-2 + skoXC1 * (1 + skoXC1)) / sqrt(skoXC1) <=
     -1/2 + skoXC1 * (-2 + skoXC1 * 5/2) | sqrt(skoXC1) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_101, plain,
    (-1/2 + skoXC1 * (-2 + skoXC1 * 5/2) <
     (-1 + skoXC1) / sqrt(skoXC1) * (skoXC1 * (2 + skoXC1)) |
     skoXC1 * (2 + skoXC1) <= 0 |
     (-1 + skoXC1) / sqrt(skoXC1) <=
     (-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) / (skoXC1 * (2 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_102, plain,
    (~ lgen(0, X_000063, ln(X_000062)) | X_000063 <= ln(X_000062)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_103, plain,
    (~ lgen(0, X_000063,
            1/2 * (1 + 5 * X_000062) * (X_000062 - 1) /
            (X_000062 * (2 + X_000062))) | X_000062 <= 0 |
     lgen(0, X_000063, ln(X_000062))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_104, plain,
    (~ lgen(0, X_000063,
            1/2 * (1 + 5 * X_000062) * (X_000062 - 1) /
            (X_000062 * (2 + X_000062))) | X_000062 <= 0 |
     X_000063 <= ln(X_000062)),
    inference(resolve, [], [refute_0_103, refute_0_102])).

cnf(refute_0_105, plain,
    ((-1/2 + X_000062 * (-2 + X_000062 * 5/2)) /
     (X_000062 * (2 + X_000062)) < X_000063 | X_000062 <= 0 |
     X_000063 <= ln(X_000062)), inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    ((-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) / (skoXC1 * (2 + skoXC1)) <
     (-1 + skoXC1) / sqrt(skoXC1) |
     (-1 + skoXC1) / sqrt(skoXC1) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_105])).

cnf(refute_0_107, plain,
    ((-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) / (skoXC1 * (2 + skoXC1)) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_106, refute_0_8])).

cnf(refute_0_108, plain,
    ((-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) / (skoXC1 * (2 + skoXC1)) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_107, refute_0_76])).

cnf(refute_0_109, plain,
    (-1/2 + skoXC1 * (-2 + skoXC1 * 5/2) <
     (-1 + skoXC1) / sqrt(skoXC1) * (skoXC1 * (2 + skoXC1)) | 1 <= skoXC1 |
     skoXC1 * (2 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_101, refute_0_108])).

cnf(refute_0_110, plain,
    (-1/2 + skoXC1 * (-2 + skoXC1 * 5/2) <
     skoXC1 * (-2 + skoXC1 * (1 + skoXC1)) / sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 * (2 + skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain, (0 < skoXC1 * (2 + skoXC1) | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_112, plain,
    (-1/2 + skoXC1 * (-2 + skoXC1 * 5/2) <
     skoXC1 * (-2 + skoXC1 * (1 + skoXC1)) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    ((-1/2 + skoXC1 * (-2 + skoXC1 * 5/2)) * sqrt(skoXC1) <
     skoXC1 * (-2 + skoXC1 * (1 + skoXC1)) | 1 <= skoXC1 |
     sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_100, refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     sqrt(skoXC1) * (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) | 1 <= skoXC1 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     sqrt(skoXC1) * (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_114, refute_0_76])).

cnf(refute_0_116, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) < sqrt(skoXC1) | 1 <= skoXC1 |
     1/2 + skoXC1 * (2 + skoXC1 * -5/2) <= 0),
    inference(split, [], [refute_0_115])).

cnf(refute_0_117, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) < sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 * (2 + skoXC1 * -5/2) <= -1/2),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (-1/2 < skoXC1 * (2 + skoXC1 * -5/2) | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_119, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) < sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_117, refute_0_118])).

cnf(refute_0_120, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) |
     sqrt(skoXC1) <=
     skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2))),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_121, plain,
    (skoXC1 < 0 |
     skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_120, refute_0_119])).

cnf(refute_0_122, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) /
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_121])).

cnf(refute_0_123, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) *
     (1/2 + skoXC1 * (2 + skoXC1 * -5/2)) | 1 <= skoXC1 |
     1/2 + skoXC1 * (2 + skoXC1 * -5/2) <= 0),
    inference(resolve, [], [refute_0_99, refute_0_122])).

cnf(refute_0_124, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     (1/16 +
      skoXC1 *
      (2 +
       skoXC1 *
       (177/16 +
        skoXC1 *
        (21/2 +
         skoXC1 * (-237/16 + skoXC1 * (-17/2 + skoXC1 * -5/16)))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1 |
     skoXC1 * (2 + skoXC1 * -5/2) <= -1/2),
    inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) <
     (1/16 +
      skoXC1 *
      (2 +
       skoXC1 *
       (177/16 +
        skoXC1 *
        (21/2 +
         skoXC1 * (-237/16 + skoXC1 * (-17/2 + skoXC1 * -5/16)))))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_124, refute_0_118])).

cnf(refute_0_126, plain,
    (skoXC1 * (2 + skoXC1 * (-1 + skoXC1 * -1)) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/16 +
     skoXC1 *
     (2 +
      skoXC1 *
      (177/16 +
       skoXC1 *
       (21/2 + skoXC1 * (-237/16 + skoXC1 * (-17/2 + skoXC1 * -5/16))))) |
     1 <= skoXC1 | 1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_98, refute_0_125])).

cnf(refute_0_127, plain,
    (-1/16 <
     skoXC1 *
     (skoXC1 *
      (-31/16 +
       skoXC1 *
       (9/2 + skoXC1 * (-45/16 + skoXC1 * (-1/2 + skoXC1 * 11/16))))) |
     1 <= skoXC1 | skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= -1),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain,
    (-1/16 <
     skoXC1 *
     (skoXC1 *
      (-31/16 +
       skoXC1 *
       (9/2 + skoXC1 * (-45/16 + skoXC1 * (-1/2 + skoXC1 * 11/16))))) |
     -1 < skoXC1 * (7 + skoXC1 * (7 + skoXC1)) | 1 <= skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_129, plain,
    (-1/16 <
     skoXC1 *
     (skoXC1 *
      (-31/16 +
       skoXC1 *
       (9/2 + skoXC1 * (-45/16 + skoXC1 * (-1/2 + skoXC1 * 11/16))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (skoXC1 < 1 |
     -1/16 <
     skoXC1 *
     (skoXC1 *
      (-31/16 +
       skoXC1 *
       (9/2 + skoXC1 * (-45/16 + skoXC1 * (-1/2 + skoXC1 * 11/16)))))),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_131, plain,
    (-1/16 <
     skoXC1 *
     (skoXC1 *
      (-31/16 +
       skoXC1 *
       (9/2 + skoXC1 * (-45/16 + skoXC1 * (-1/2 + skoXC1 * 11/16)))))),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (-1/24 <
     skoXC1 *
     (-13/12 +
      skoXC1 *
      (-11/24 +
       skoXC1 *
       (23/4 +
        skoXC1 *
        (-39/8 + skoXC1 * (1/12 + skoXC1 * (-1/24 + skoXC1 * 7/12)))))) |
     -1 < skoXC1 * (7 + skoXC1 * (7 + skoXC1)) | 1 <= skoXC1),
    inference(decision, [],
              [refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_133, plain,
    (-1/24 <
     skoXC1 *
     (-13/12 +
      skoXC1 *
      (-11/24 +
       skoXC1 *
       (23/4 +
        skoXC1 *
        (-39/8 + skoXC1 * (1/12 + skoXC1 * (-1/24 + skoXC1 * 7/12)))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_97, refute_0_132])).

cnf(refute_0_134, plain,
    (skoXC1 < 1 |
     -1/24 <
     skoXC1 *
     (-13/12 +
      skoXC1 *
      (-11/24 +
       skoXC1 *
       (23/4 +
        skoXC1 *
        (-39/8 + skoXC1 * (1/12 + skoXC1 * (-1/24 + skoXC1 * 7/12))))))),
    inference(decision, [],
              [refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_135, plain,
    (-1/24 <
     skoXC1 *
     (-13/12 +
      skoXC1 *
      (-11/24 +
       skoXC1 *
       (23/4 +
        skoXC1 *
        (-39/8 + skoXC1 * (1/12 + skoXC1 * (-1/24 + skoXC1 * 7/12))))))),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain,
    (-1/8 <
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_137, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     ln(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_65, refute_0_136])).

cnf(refute_0_138, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_46, refute_0_137])).

cnf(refute_0_139, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_138, refute_0_29])).

cnf(refute_0_140, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_41, refute_0_139])).

cnf(refute_0_141, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/12 +
      skoXC1 *
      (-289/24 +
       skoXC1 *
       (-467/12 +
        skoXC1 *
        (-135/8 +
         skoXC1 *
         (171/4 +
          skoXC1 * (565/24 + skoXC1 * (23/12 + skoXC1 * 1/24))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/12 +
      skoXC1 *
      (-289/24 +
       skoXC1 *
       (-467/12 +
        skoXC1 *
        (-135/8 +
         skoXC1 *
         (171/4 +
          skoXC1 * (565/24 + skoXC1 * (23/12 + skoXC1 * 1/24))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_141, refute_0_136])).

cnf(refute_0_143, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * (6 + skoXC1 * 3)) <
     -5/12 +
     skoXC1 *
     (-289/24 +
      skoXC1 *
      (-467/12 +
       skoXC1 *
       (-135/8 +
        skoXC1 *
        (171/4 + skoXC1 * (565/24 + skoXC1 * (23/12 + skoXC1 * 1/24)))))) |
     1 + skoXC1 * (6 + skoXC1 * 3) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_40, refute_0_142])).

cnf(refute_0_144, plain,
    (-7/12 <
     skoXC1 *
     (-1/24 +
      skoXC1 *
      (1/12 +
       skoXC1 *
       (-39/8 +
        skoXC1 *
        (23/4 + skoXC1 * (-11/24 + skoXC1 * (-13/12 + skoXC1 * 1/24)))))) |
     skoXC1 * (6 + skoXC1 * 3) <= -1 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * 2) <
     -5/16 +
     skoXC1 *
     (-17/2 +
      skoXC1 *
      (-237/16 +
       skoXC1 *
       (21/2 + skoXC1 * (177/16 + skoXC1 * (2 + skoXC1 * 1/16))))) |
     1 + skoXC1 * 2 <= 0 |
     (-5/16 +
      skoXC1 *
      (-17/2 +
       skoXC1 *
       (-237/16 +
        skoXC1 *
        (21/2 + skoXC1 * (177/16 + skoXC1 * (2 + skoXC1 * 1/16)))))) /
     (1 + skoXC1 * 2) <=
     -1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_146, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 |
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_147, plain,
    (~ lgen(0, ln(X_000055), X_000056) | ln(X_000055) <= X_000056),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_148, plain,
    (~ lgen(0, 1/2 * (X_000055 + 5) * (X_000055 - 1) / (2 * X_000055 + 1),
            X_000056) | X_000055 <= 0 | lgen(0, ln(X_000055), X_000056)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_149, plain,
    (~ lgen(0, 1/2 * (X_000055 + 5) * (X_000055 - 1) / (2 * X_000055 + 1),
            X_000056) | X_000055 <= 0 | ln(X_000055) <= X_000056),
    inference(resolve, [], [refute_0_148, refute_0_147])).

cnf(refute_0_150, plain,
    (X_000056 <
     (-5/2 + X_000055 * (2 + X_000055 * 1/2)) / (1 + X_000055 * 2) |
     X_000055 <= 0 | ln(X_000055) <= X_000056),
    inference(arithmetic, [], [refute_0_149])).

cnf(refute_0_151, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) |
     ln(skoXC1) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_150])).

cnf(refute_0_152, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) |
     skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_151, refute_0_137])).

cnf(refute_0_153, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_152, refute_0_29])).

cnf(refute_0_154, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/2 + skoXC1 * (2 + skoXC1 * 1/2)) / (1 + skoXC1 * 2) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_146, refute_0_153])).

cnf(refute_0_155, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/16 +
      skoXC1 *
      (-17/2 +
       skoXC1 *
       (-237/16 +
        skoXC1 *
        (21/2 + skoXC1 * (177/16 + skoXC1 * (2 + skoXC1 * 1/16)))))) /
     (1 + skoXC1 * 2) |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-5/16 +
      skoXC1 *
      (-17/2 +
       skoXC1 *
       (-237/16 +
        skoXC1 *
        (21/2 + skoXC1 * (177/16 + skoXC1 * (2 + skoXC1 * 1/16)))))) /
     (1 + skoXC1 * 2) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_155, refute_0_136])).

cnf(refute_0_157, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * 2) <
     -5/16 +
     skoXC1 *
     (-17/2 +
      skoXC1 *
      (-237/16 +
       skoXC1 *
       (21/2 + skoXC1 * (177/16 + skoXC1 * (2 + skoXC1 * 1/16))))) |
     1 + skoXC1 * 2 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_145, refute_0_156])).

cnf(refute_0_158, plain,
    (-11/16 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (-45/16 +
       skoXC1 * (9/2 + skoXC1 * (-31/16 + skoXC1 * (skoXC1 * 1/16))))) |
     skoXC1 <= -1/2 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    (-11/16 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (-45/16 +
       skoXC1 * (9/2 + skoXC1 * (-31/16 + skoXC1 * (skoXC1 * 1/16))))) |
     -1/2 < skoXC1 | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_160, plain,
    (-11/16 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (-45/16 +
       skoXC1 * (9/2 + skoXC1 * (-31/16 + skoXC1 * (skoXC1 * 1/16))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (-11/16 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (-45/16 +
       skoXC1 * (9/2 + skoXC1 * (-31/16 + skoXC1 * (skoXC1 * 1/16))))) |
     1 < skoXC1),
    inference(decision, [],
              [refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_162, plain,
    (-11/16 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (-45/16 +
       skoXC1 * (9/2 + skoXC1 * (-31/16 + skoXC1 * (skoXC1 * 1/16)))))),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (-7/12 <
     skoXC1 *
     (-1/24 +
      skoXC1 *
      (1/12 +
       skoXC1 *
       (-39/8 +
        skoXC1 *
        (23/4 + skoXC1 * (-11/24 + skoXC1 * (-13/12 + skoXC1 * 1/24)))))) |
     -1 < skoXC1 * (6 + skoXC1 * 3) | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_164, plain,
    (-7/12 <
     skoXC1 *
     (-1/24 +
      skoXC1 *
      (1/12 +
       skoXC1 *
       (-39/8 +
        skoXC1 *
        (23/4 + skoXC1 * (-11/24 + skoXC1 * (-13/12 + skoXC1 * 1/24)))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_144, refute_0_163])).

cnf(refute_0_165, plain,
    (-7/12 <
     skoXC1 *
     (-1/24 +
      skoXC1 *
      (1/12 +
       skoXC1 *
       (-39/8 +
        skoXC1 *
        (23/4 + skoXC1 * (-11/24 + skoXC1 * (-13/12 + skoXC1 * 1/24)))))) |
     1 < skoXC1),
    inference(decision, [],
              [refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_166, plain,
    (-7/12 <
     skoXC1 *
     (-1/24 +
      skoXC1 *
      (1/12 +
       skoXC1 *
       (-39/8 +
        skoXC1 *
        (23/4 + skoXC1 * (-11/24 + skoXC1 * (-13/12 + skoXC1 * 1/24))))))),
    inference(resolve, [], [refute_0_164, refute_0_165])).

cnf(refute_0_167, plain, (1 <= skoXC1 | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_168, plain,
    ((-1 + skoXC1) / ln(skoXC1) <
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) /
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_167, refute_0_39])).

cnf(refute_0_169, plain,
    ((-1 + skoXC1) / ln(skoXC1) *
     (4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1))) <
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) |
     4096 + skoXC1 * (1792 + skoXC1 * (112 + skoXC1)) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_18, refute_0_168])).

cnf(refute_0_170, plain,
    (-2048 +
     skoXC1 * (-3584 + skoXC1 * (-560 + skoXC1 * (-14 + skoXC1 * -1/32))) <
     (4096 +
      skoXC1 *
      (-2304 + skoXC1 * (-1680 + skoXC1 * (-111 + skoXC1 * -1)))) /
     ln(skoXC1) | skoXC1 * (1792 + skoXC1 * (112 + skoXC1)) <= -4096 |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (-4096 < skoXC1 * (1792 + skoXC1 * (112 + skoXC1)) | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_172, plain,
    (-2048 +
     skoXC1 * (-3584 + skoXC1 * (-560 + skoXC1 * (-14 + skoXC1 * -1/32))) <
     (4096 +
      skoXC1 *
      (-2304 + skoXC1 * (-1680 + skoXC1 * (-111 + skoXC1 * -1)))) /
     ln(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    ((-2048 +
      skoXC1 *
      (-3584 + skoXC1 * (-560 + skoXC1 * (-14 + skoXC1 * -1/32)))) *
     ln(skoXC1) <
     4096 +
     skoXC1 * (-2304 + skoXC1 * (-1680 + skoXC1 * (-111 + skoXC1 * -1))) |
     ln(skoXC1) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_17, refute_0_172])).

cnf(refute_0_174, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     ln(skoXC1) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     ln(skoXC1) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_174, refute_0_29])).

cnf(refute_0_176, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     ln(skoXC1) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1), inference(split, [], [refute_0_175])).

cnf(refute_0_177, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     ln(skoXC1) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (-2048 <
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_179, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     ln(skoXC1) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (~ lgen(0, ln(X_000139), X_000140) | ln(X_000139) <= X_000140),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_181, plain,
    (~ lgen(0,
            1/70 *
            (353 + 1921 * X_000139 ^ 5 + 8389 * X_000139 +
             38524 * X_000139 ^ 2 + 20149 * X_000139 ^ 4 +
             50774 * X_000139 ^ 3 + 10 * X_000139 ^ 6) * (X_000139 - 1) /
            (1 + 126 * X_000139 ^ 5 + 42 * X_000139 + 315 * X_000139 ^ 2 +
             525 * X_000139 ^ 4 + 700 * X_000139 ^ 3 + 7 * X_000139 ^ 6),
            X_000140) | X_000139 <= 0 | lgen(0, ln(X_000139), X_000140)),
    inference(subst, [], [ln_upper_bound_cf13])).

cnf(refute_0_182, plain,
    (~ lgen(0,
            1/70 *
            (353 + 1921 * X_000139 ^ 5 + 8389 * X_000139 +
             38524 * X_000139 ^ 2 + 20149 * X_000139 ^ 4 +
             50774 * X_000139 ^ 3 + 10 * X_000139 ^ 6) * (X_000139 - 1) /
            (1 + 126 * X_000139 ^ 5 + 42 * X_000139 + 315 * X_000139 ^ 2 +
             525 * X_000139 ^ 4 + 700 * X_000139 ^ 3 + 7 * X_000139 ^ 6),
            X_000140) | X_000139 <= 0 | ln(X_000139) <= X_000140),
    inference(resolve, [], [refute_0_181, refute_0_180])).

cnf(refute_0_183, plain,
    (X_000140 <
     (-353/70 +
      X_000139 *
      (-574/5 +
       X_000139 *
       (-861/2 +
        X_000139 *
        (-175 +
         X_000139 *
         (875/2 +
          X_000139 * (1302/5 + X_000139 * (273/10 + X_000139 * 1/7))))))) /
     (1 +
      X_000139 *
      (42 +
       X_000139 *
       (315 +
        X_000139 *
        (700 + X_000139 * (525 + X_000139 * (126 + X_000139 * 7)))))) |
     X_000139 <= 0 | ln(X_000139) <= X_000140),
    inference(arithmetic, [], [refute_0_182])).

cnf(refute_0_184, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) |
     ln(skoXC1) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_183])).

cnf(refute_0_185, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) |
     skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_184, refute_0_179])).

cnf(refute_0_186, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_185, refute_0_29])).

cnf(refute_0_187, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-353/70 +
      skoXC1 *
      (-574/5 +
       skoXC1 *
       (-861/2 +
        skoXC1 *
        (-175 +
         skoXC1 *
         (875/2 +
          skoXC1 * (1302/5 + skoXC1 * (273/10 + skoXC1 * 1/7))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_16, refute_0_186])).

cnf(refute_0_188, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-361472/35 +
      skoXC1 *
      (-253184 +
       skoXC1 *
       (-6479656/5 +
        skoXC1 *
        (-9828353/5 +
         skoXC1 *
         (58492319/2240 +
          skoXC1 *
          (159781489/80 +
           skoXC1 *
           (78830115/64 +
            skoXC1 *
            (280088819/1120 +
             skoXC1 *
             (6226967/320 +
              skoXC1 *
              (37627/80 + skoXC1 * (913/320 + skoXC1 * 1/224))))))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-361472/35 +
      skoXC1 *
      (-253184 +
       skoXC1 *
       (-6479656/5 +
        skoXC1 *
        (-9828353/5 +
         skoXC1 *
         (58492319/2240 +
          skoXC1 *
          (159781489/80 +
           skoXC1 *
           (78830115/64 +
            skoXC1 *
            (280088819/1120 +
             skoXC1 *
             (6226967/320 +
              skoXC1 *
              (37627/80 + skoXC1 * (913/320 + skoXC1 * 1/224))))))))))) /
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_188, refute_0_178])).

cnf(refute_0_190, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 +
      skoXC1 *
      (42 +
       skoXC1 *
       (315 +
        skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7)))))) <
     -361472/35 +
     skoXC1 *
     (-253184 +
      skoXC1 *
      (-6479656/5 +
       skoXC1 *
       (-9828353/5 +
        skoXC1 *
        (58492319/2240 +
         skoXC1 *
         (159781489/80 +
          skoXC1 *
          (78830115/64 +
           skoXC1 *
           (280088819/1120 +
            skoXC1 *
            (6226967/320 +
             skoXC1 *
             (37627/80 + skoXC1 * (913/320 + skoXC1 * 1/224)))))))))) |
     1 +
     skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_15, refute_0_189])).

cnf(refute_0_191, plain,
    (218112/35 <
     skoXC1 *
     (-83456 +
      skoXC1 *
      (-520696/5 +
       skoXC1 *
       (525492/5 +
        skoXC1 *
        (66863199/2240 +
         skoXC1 *
         (7420609/80 +
          skoXC1 *
          (644707/64 +
           skoXC1 *
           (-41108141/1120 +
            skoXC1 *
            (-2179753/320 +
             skoXC1 *
             (-34613/80 + skoXC1 * (-1327/320 + skoXC1 * 1/224)))))))))) |
     skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) <=
     -1 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_190])).

cnf(refute_0_192, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     1/48 +
     skoXC1 *
     (391/120 +
      skoXC1 *
      (5519/60 +
       skoXC1 *
       (1901/3 +
        skoXC1 *
        (83089/60 +
         skoXC1 *
         (32003/60 +
          skoXC1 *
          (-144391/120 +
           skoXC1 *
           (-16661/15 +
            skoXC1 *
            (-14741/48 + skoXC1 * (-3053/120 + skoXC1 * -71/120))))))))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0 |
     (1/48 +
      skoXC1 *
      (391/120 +
       skoXC1 *
       (5519/60 +
        skoXC1 *
        (1901/3 +
         skoXC1 *
         (83089/60 +
          skoXC1 *
          (32003/60 +
           skoXC1 *
           (-144391/120 +
            skoXC1 *
            (-16661/15 +
             skoXC1 *
             (-14741/48 +
              skoXC1 * (-3053/120 + skoXC1 * -71/120)))))))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <=
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_193, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_194, plain,
    ((-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) *
     (ln(skoXC1) * -1) <
     -1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) |
     ln(skoXC1) * -1 <= 0 |
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (ln(skoXC1) * -1) <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_195, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0 |
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <=
     (1 + skoXC1 * -1) / (ln(skoXC1) * -1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_196, plain,
    (1 < ln(skoXC1) / (-1 + skoXC1) * sqrt(skoXC1) |
     ln(skoXC1) / (-1 + skoXC1) <= 1 / sqrt(skoXC1) | sqrt(skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_197, plain,
    (1 < ln(skoXC1) / (-1 + skoXC1) * sqrt(skoXC1) | sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_196, refute_0_5])).

cnf(refute_0_198, plain,
    (1 < sqrt(skoXC1) * ln(skoXC1) / (-1 + skoXC1) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain, (0 < skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_200, plain, (1 < sqrt(skoXC1) * ln(skoXC1) / (-1 + skoXC1)),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (sqrt(skoXC1) * ln(skoXC1) < 1 * (-1 + skoXC1) | 0 <= -1 + skoXC1 |
     sqrt(skoXC1) * ln(skoXC1) / (-1 + skoXC1) <= 1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_202, plain,
    (sqrt(skoXC1) * ln(skoXC1) < 1 * (-1 + skoXC1) | 0 <= -1 + skoXC1),
    inference(resolve, [], [refute_0_201, refute_0_200])).

cnf(refute_0_203, plain,
    (1 + skoXC1 * -1 < sqrt(skoXC1) * (ln(skoXC1) * -1) | 1 <= skoXC1),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) < sqrt(skoXC1) | 1 <= skoXC1 |
     ln(skoXC1) * -1 <= 0), inference(split, [], [refute_0_203])).

cnf(refute_0_205, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) < sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) < sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_205, refute_0_76])).

cnf(refute_0_207, plain,
    (skoXC1 < 0 |
     (1 + skoXC1 * -1) / (ln(skoXC1) * -1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) |
     sqrt(skoXC1) <= (1 + skoXC1 * -1) / (ln(skoXC1) * -1)),
    inference(subst, [], [refute_0_52])).

cnf(refute_0_208, plain,
    (skoXC1 < 0 |
     (1 + skoXC1 * -1) / (ln(skoXC1) * -1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_207, refute_0_206])).

cnf(refute_0_209, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) <
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) /
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_208])).

cnf(refute_0_210, plain,
    ((1 + skoXC1 * -1) / (ln(skoXC1) * -1) *
     (1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1))) <
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     1 <= skoXC1 | 1 + skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_195, refute_0_209])).

cnf(refute_0_211, plain,
    (-1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) <
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (ln(skoXC1) * -1) | 1 <= skoXC1 |
     skoXC1 * (7 + skoXC1 * (7 + skoXC1)) <= -1),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (-1 < skoXC1 * (7 + skoXC1 * (7 + skoXC1)) | 1 <= skoXC1),
    inference(decision, [],
              [refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_213, plain,
    (-1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) <
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (ln(skoXC1) * -1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_211, refute_0_212])).

cnf(refute_0_214, plain,
    ((-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) *
     (ln(skoXC1) * -1) <
     -1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) | 1 <= skoXC1 |
     ln(skoXC1) * -1 <= 0),
    inference(resolve, [], [refute_0_194, refute_0_213])).

cnf(refute_0_215, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     ln(skoXC1) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     ln(skoXC1) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_215, refute_0_76])).

cnf(refute_0_217, plain,
    (ln(skoXC1) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     1 <= skoXC1), inference(split, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (ln(skoXC1) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8), inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (-1/8 <
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_220, plain,
    (ln(skoXC1) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain,
    (~ lgen(0, X_000133, ln(X_000132)) | X_000133 <= ln(X_000132)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_222, plain,
    (~ lgen(0, X_000133,
            1/30 *
            (142 * X_000132 ^ 5 + 2272 * X_000132 ^ 4 +
             6397 * X_000132 ^ 3 + 4397 * X_000132 ^ 2 + 647 * X_000132 +
             5) * (X_000132 - 1) /
            (X_000132 *
             (X_000132 ^ 5 + 30 * X_000132 ^ 4 + 150 * X_000132 ^ 3 +
              200 * X_000132 ^ 2 + 75 * X_000132 + 6))) | X_000132 <= 0 |
     lgen(0, X_000133, ln(X_000132))),
    inference(subst, [], [ln_lower_bound_cf11])).

cnf(refute_0_223, plain,
    (~ lgen(0, X_000133,
            1/30 *
            (142 * X_000132 ^ 5 + 2272 * X_000132 ^ 4 +
             6397 * X_000132 ^ 3 + 4397 * X_000132 ^ 2 + 647 * X_000132 +
             5) * (X_000132 - 1) /
            (X_000132 *
             (X_000132 ^ 5 + 30 * X_000132 ^ 4 + 150 * X_000132 ^ 3 +
              200 * X_000132 ^ 2 + 75 * X_000132 + 6))) | X_000132 <= 0 |
     X_000133 <= ln(X_000132)),
    inference(resolve, [], [refute_0_222, refute_0_221])).

cnf(refute_0_224, plain,
    ((-1/6 +
      X_000132 *
      (-107/5 +
       X_000132 *
       (-125 +
        X_000132 *
        (-200/3 +
         X_000132 * (275/2 + X_000132 * (71 + X_000132 * 71/15)))))) /
     (X_000132 *
      (6 +
       X_000132 *
       (75 +
        X_000132 *
        (200 + X_000132 * (150 + X_000132 * (30 + X_000132)))))) <
     X_000133 | X_000132 <= 0 | X_000133 <= ln(X_000132)),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     ln(skoXC1) | skoXC1 <= 0), inference(subst, [], [refute_0_224])).

cnf(refute_0_226, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_225, refute_0_220])).

cnf(refute_0_227, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_226, refute_0_76])).

cnf(refute_0_228, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_193, refute_0_227])).

cnf(refute_0_229, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/48 +
      skoXC1 *
      (391/120 +
       skoXC1 *
       (5519/60 +
        skoXC1 *
        (1901/3 +
         skoXC1 *
         (83089/60 +
          skoXC1 *
          (32003/60 +
           skoXC1 *
           (-144391/120 +
            skoXC1 *
            (-16661/15 +
             skoXC1 *
             (-14741/48 +
              skoXC1 * (-3053/120 + skoXC1 * -71/120)))))))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     1 <= skoXC1 |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8), inference(arithmetic, [], [refute_0_228])).

cnf(refute_0_230, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/48 +
      skoXC1 *
      (391/120 +
       skoXC1 *
       (5519/60 +
        skoXC1 *
        (1901/3 +
         skoXC1 *
         (83089/60 +
          skoXC1 *
          (32003/60 +
           skoXC1 *
           (-144391/120 +
            skoXC1 *
            (-16661/15 +
             skoXC1 *
             (-14741/48 +
              skoXC1 * (-3053/120 + skoXC1 * -71/120)))))))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_229, refute_0_219])).

cnf(refute_0_231, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     1/48 +
     skoXC1 *
     (391/120 +
      skoXC1 *
      (5519/60 +
       skoXC1 *
       (1901/3 +
        skoXC1 *
        (83089/60 +
         skoXC1 *
         (32003/60 +
          skoXC1 *
          (-144391/120 +
           skoXC1 *
           (-16661/15 +
            skoXC1 *
            (-14741/48 + skoXC1 * (-3053/120 + skoXC1 * -71/120))))))))) |
     1 <= skoXC1 |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_192, refute_0_230])).

cnf(refute_0_232, plain,
    (-1/48 <
     skoXC1 *
     (-329/120 +
      skoXC1 *
      (-1141/60 +
       skoXC1 *
       (-49/3 +
        skoXC1 *
        (4249/60 +
         skoXC1 *
         (3563/60 +
          skoXC1 *
          (-13111/120 +
           skoXC1 *
           (-251/15 +
            skoXC1 *
            (1099/48 + skoXC1 * (1267/120 + skoXC1 * 49/120))))))))) |
     1 <= skoXC1 |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0), inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) <
     -145408/15 +
     skoXC1 *
     (-2435584/15 +
      skoXC1 *
      (-1616144/3 +
       skoXC1 *
       (-5941394/15 +
        skoXC1 *
        (200130809/480 +
         skoXC1 *
         (253071991/480 +
          skoXC1 *
          (142049251/960 +
           skoXC1 *
           (172001/12 +
            skoXC1 *
            (190483/480 + skoXC1 * (1441/480 + skoXC1 * 1/192))))))))) |
     1 +
     skoXC1 *
     (30 + skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6)))) <=
     0 |
     (-145408/15 +
      skoXC1 *
      (-2435584/15 +
       skoXC1 *
       (-1616144/3 +
        skoXC1 *
        (-5941394/15 +
         skoXC1 *
         (200130809/480 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (142049251/960 +
            skoXC1 *
            (172001/12 +
             skoXC1 *
             (190483/480 + skoXC1 * (1441/480 + skoXC1 * 1/192)))))))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) <=
     -4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_234, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 |
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_235, plain,
    (~ lgen(0, ln(X_000125), X_000126) | ln(X_000125) <= X_000126),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_236, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000125 ^ 5 + 647 * X_000125 ^ 4 + 4397 * X_000125 ^ 3 +
             6397 * X_000125 ^ 2 + 2272 * X_000125 + 142) *
            (X_000125 - 1) /
            (6 * X_000125 ^ 5 + 75 * X_000125 ^ 4 + 200 * X_000125 ^ 3 +
             150 * X_000125 ^ 2 + 30 * X_000125 + 1), X_000126) |
     X_000125 <= 0 | lgen(0, ln(X_000125), X_000126)),
    inference(subst, [], [ln_upper_bound_cf11])).

cnf(refute_0_237, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000125 ^ 5 + 647 * X_000125 ^ 4 + 4397 * X_000125 ^ 3 +
             6397 * X_000125 ^ 2 + 2272 * X_000125 + 142) *
            (X_000125 - 1) /
            (6 * X_000125 ^ 5 + 75 * X_000125 ^ 4 + 200 * X_000125 ^ 3 +
             150 * X_000125 ^ 2 + 30 * X_000125 + 1), X_000126) |
     X_000125 <= 0 | ln(X_000125) <= X_000126),
    inference(resolve, [], [refute_0_236, refute_0_235])).

cnf(refute_0_238, plain,
    (X_000126 <
     (-71/15 +
      X_000125 *
      (-71 +
       X_000125 *
       (-275/2 +
        X_000125 *
        (200/3 +
         X_000125 * (125 + X_000125 * (107/5 + X_000125 * 1/6)))))) /
     (1 +
      X_000125 *
      (30 +
       X_000125 *
       (150 + X_000125 * (200 + X_000125 * (75 + X_000125 * 6))))) |
     X_000125 <= 0 | ln(X_000125) <= X_000126),
    inference(arithmetic, [], [refute_0_237])).

cnf(refute_0_239, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) |
     ln(skoXC1) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_238])).

cnf(refute_0_240, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) |
     skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_239, refute_0_179])).

cnf(refute_0_241, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_240, refute_0_29])).

cnf(refute_0_242, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-71/15 +
      skoXC1 *
      (-71 +
       skoXC1 *
       (-275/2 +
        skoXC1 *
        (200/3 + skoXC1 * (125 + skoXC1 * (107/5 + skoXC1 * 1/6)))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_234, refute_0_241])).

cnf(refute_0_243, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-145408/15 +
      skoXC1 *
      (-2435584/15 +
       skoXC1 *
       (-1616144/3 +
        skoXC1 *
        (-5941394/15 +
         skoXC1 *
         (200130809/480 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (142049251/960 +
            skoXC1 *
            (172001/12 +
             skoXC1 *
             (190483/480 + skoXC1 * (1441/480 + skoXC1 * 1/192)))))))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-145408/15 +
      skoXC1 *
      (-2435584/15 +
       skoXC1 *
       (-1616144/3 +
        skoXC1 *
        (-5941394/15 +
         skoXC1 *
         (200130809/480 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (142049251/960 +
            skoXC1 *
            (172001/12 +
             skoXC1 *
             (190483/480 + skoXC1 * (1441/480 + skoXC1 * 1/192)))))))))) /
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_243, refute_0_178])).

cnf(refute_0_245, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 +
      skoXC1 *
      (30 +
       skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6))))) <
     -145408/15 +
     skoXC1 *
     (-2435584/15 +
      skoXC1 *
      (-1616144/3 +
       skoXC1 *
       (-5941394/15 +
        skoXC1 *
        (200130809/480 +
         skoXC1 *
         (253071991/480 +
          skoXC1 *
          (142049251/960 +
           skoXC1 *
           (172001/12 +
            skoXC1 *
            (190483/480 + skoXC1 * (1441/480 + skoXC1 * 1/192))))))))) |
     1 +
     skoXC1 *
     (30 + skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6)))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_233, refute_0_244])).

cnf(refute_0_246, plain,
    (83968/15 <
     skoXC1 *
     (-626944/15 +
      skoXC1 *
      (14656/3 +
       skoXC1 *
       (404941/15 +
        skoXC1 *
        (3843929/480 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (-13637789/960 +
           skoXC1 *
           (-51259/12 +
            skoXC1 *
            (-165197/480 + skoXC1 * (-1439/480 + skoXC1 * 1/192))))))))) |
     skoXC1 *
     (30 + skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6)))) <=
     -1 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     1/40 +
     skoXC1 *
     (643/240 +
      skoXC1 *
      (63 +
       skoXC1 *
       (12003/40 +
        skoXC1 *
        (23339/60 +
         skoXC1 *
         (-722/5 +
          skoXC1 *
          (-6223/15 +
           skoXC1 *
           (-1399/8 + skoXC1 * (-803/40 + skoXC1 * -131/240)))))))) |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0 |
     (1/40 +
      skoXC1 *
      (643/240 +
       skoXC1 *
       (63 +
        skoXC1 *
        (12003/40 +
         skoXC1 *
         (23339/60 +
          skoXC1 *
          (-722/5 +
           skoXC1 *
           (-6223/15 +
            skoXC1 *
            (-1399/8 + skoXC1 * (-803/40 + skoXC1 * -131/240))))))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <=
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_248, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     (-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_249, plain,
    (~ lgen(0, X_000119, ln(X_000118)) | X_000119 <= ln(X_000118)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_250, plain,
    (~ lgen(0, X_000119,
            1/30 *
            (6 + 481 * X_000118 + 1881 * X_000118 ^ 2 +
             1281 * X_000118 ^ 3 + 131 * X_000118 ^ 4) * (X_000118 - 1) /
            (X_000118 *
             (5 + 40 * X_000118 + 60 * X_000118 ^ 2 + 20 * X_000118 ^ 3 +
              X_000118 ^ 4))) | X_000118 <= 0 |
     lgen(0, X_000119, ln(X_000118))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_251, plain,
    (~ lgen(0, X_000119,
            1/30 *
            (6 + 481 * X_000118 + 1881 * X_000118 ^ 2 +
             1281 * X_000118 ^ 3 + 131 * X_000118 ^ 4) * (X_000118 - 1) /
            (X_000118 *
             (5 + 40 * X_000118 + 60 * X_000118 ^ 2 + 20 * X_000118 ^ 3 +
              X_000118 ^ 4))) | X_000118 <= 0 | X_000119 <= ln(X_000118)),
    inference(resolve, [], [refute_0_250, refute_0_249])).

cnf(refute_0_252, plain,
    ((-1/5 +
      X_000118 *
      (-95/6 +
       X_000118 *
       (-140/3 +
        X_000118 * (20 + X_000118 * (115/3 + X_000118 * 131/30))))) /
     (X_000118 *
      (5 +
       X_000118 * (40 + X_000118 * (60 + X_000118 * (20 + X_000118))))) <
     X_000119 | X_000118 <= 0 | X_000119 <= ln(X_000118)),
    inference(arithmetic, [], [refute_0_251])).

cnf(refute_0_253, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     ln(skoXC1) | skoXC1 <= 0), inference(subst, [], [refute_0_252])).

cnf(refute_0_254, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_253, refute_0_220])).

cnf(refute_0_255, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_254, refute_0_76])).

cnf(refute_0_256, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_248, refute_0_255])).

cnf(refute_0_257, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/40 +
      skoXC1 *
      (643/240 +
       skoXC1 *
       (63 +
        skoXC1 *
        (12003/40 +
         skoXC1 *
         (23339/60 +
          skoXC1 *
          (-722/5 +
           skoXC1 *
           (-6223/15 +
            skoXC1 *
            (-1399/8 + skoXC1 * (-803/40 + skoXC1 * -131/240))))))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) |
     1 <= skoXC1 |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8), inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/40 +
      skoXC1 *
      (643/240 +
       skoXC1 *
       (63 +
        skoXC1 *
        (12003/40 +
         skoXC1 *
         (23339/60 +
          skoXC1 *
          (-722/5 +
           skoXC1 *
           (-6223/15 +
            skoXC1 *
            (-1399/8 + skoXC1 * (-803/40 + skoXC1 * -131/240))))))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_257, refute_0_219])).

cnf(refute_0_259, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     1/40 +
     skoXC1 *
     (643/240 +
      skoXC1 *
      (63 +
       skoXC1 *
       (12003/40 +
        skoXC1 *
        (23339/60 +
         skoXC1 *
         (-722/5 +
          skoXC1 *
          (-6223/15 +
           skoXC1 *
           (-1399/8 + skoXC1 * (-803/40 + skoXC1 * -131/240)))))))) |
     1 <= skoXC1 |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0),
    inference(resolve, [], [refute_0_247, refute_0_258])).

cnf(refute_0_260, plain,
    (-1/40 <
     skoXC1 *
     (-557/240 +
      skoXC1 *
      (-7 +
       skoXC1 *
       (3/40 +
        skoXC1 *
        (2339/60 +
         skoXC1 *
         (-102/5 +
          skoXC1 *
          (-313/15 +
           skoXC1 * (41/8 + skoXC1 * (237/40 + skoXC1 * 109/240)))))))) |
     1 <= skoXC1 |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) <
     -134144/15 +
     skoXC1 *
     (-470784/5 +
      skoXC1 *
      (-180792 +
       skoXC1 *
       (35483/15 +
        skoXC1 *
        (60141823/320 +
         skoXC1 *
         (39844033/480 +
          skoXC1 *
          (409447/40 +
           skoXC1 * (2681/8 + skoXC1 * (3163/960 + skoXC1 * 1/160)))))))) |
     1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5))) <= 0 |
     (-134144/15 +
      skoXC1 *
      (-470784/5 +
       skoXC1 *
       (-180792 +
        skoXC1 *
        (35483/15 +
         skoXC1 *
         (60141823/320 +
          skoXC1 *
          (39844033/480 +
           skoXC1 *
           (409447/40 +
            skoXC1 *
            (2681/8 + skoXC1 * (3163/960 + skoXC1 * 1/160))))))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) <=
     -4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_262, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 |
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_263, plain,
    (~ lgen(0, ln(X_000111), X_000112) | ln(X_000111) <= X_000112),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_264, plain,
    (~ lgen(0,
            1/30 *
            (6 * X_000111 ^ 4 + 481 * X_000111 ^ 3 + 1881 * X_000111 ^ 2 +
             1281 * X_000111 + 131) * (X_000111 - 1) /
            (5 * X_000111 ^ 4 + 40 * X_000111 ^ 3 + 60 * X_000111 ^ 2 +
             20 * X_000111 + 1), X_000112) | X_000111 <= 0 |
     lgen(0, ln(X_000111), X_000112)),
    inference(subst, [], [ln_upper_bound_cf9])).

cnf(refute_0_265, plain,
    (~ lgen(0,
            1/30 *
            (6 * X_000111 ^ 4 + 481 * X_000111 ^ 3 + 1881 * X_000111 ^ 2 +
             1281 * X_000111 + 131) * (X_000111 - 1) /
            (5 * X_000111 ^ 4 + 40 * X_000111 ^ 3 + 60 * X_000111 ^ 2 +
             20 * X_000111 + 1), X_000112) | X_000111 <= 0 |
     ln(X_000111) <= X_000112),
    inference(resolve, [], [refute_0_264, refute_0_263])).

cnf(refute_0_266, plain,
    (X_000112 <
     (-131/30 +
      X_000111 *
      (-115/3 +
       X_000111 *
       (-20 + X_000111 * (140/3 + X_000111 * (95/6 + X_000111 * 1/5))))) /
     (1 +
      X_000111 * (20 + X_000111 * (60 + X_000111 * (40 + X_000111 * 5)))) |
     X_000111 <= 0 | ln(X_000111) <= X_000112),
    inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) |
     ln(skoXC1) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_266])).

cnf(refute_0_268, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) |
     skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_267, refute_0_179])).

cnf(refute_0_269, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_268, refute_0_29])).

cnf(refute_0_270, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-131/30 +
      skoXC1 *
      (-115/3 +
       skoXC1 *
       (-20 + skoXC1 * (140/3 + skoXC1 * (95/6 + skoXC1 * 1/5))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_262, refute_0_269])).

cnf(refute_0_271, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-134144/15 +
      skoXC1 *
      (-470784/5 +
       skoXC1 *
       (-180792 +
        skoXC1 *
        (35483/15 +
         skoXC1 *
         (60141823/320 +
          skoXC1 *
          (39844033/480 +
           skoXC1 *
           (409447/40 +
            skoXC1 *
            (2681/8 + skoXC1 * (3163/960 + skoXC1 * 1/160))))))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_270])).

cnf(refute_0_272, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-134144/15 +
      skoXC1 *
      (-470784/5 +
       skoXC1 *
       (-180792 +
        skoXC1 *
        (35483/15 +
         skoXC1 *
         (60141823/320 +
          skoXC1 *
          (39844033/480 +
           skoXC1 *
           (409447/40 +
            skoXC1 *
            (2681/8 + skoXC1 * (3163/960 + skoXC1 * 1/160))))))))) /
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_271, refute_0_178])).

cnf(refute_0_273, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5)))) <
     -134144/15 +
     skoXC1 *
     (-470784/5 +
      skoXC1 *
      (-180792 +
       skoXC1 *
       (35483/15 +
        skoXC1 *
        (60141823/320 +
         skoXC1 *
         (39844033/480 +
          skoXC1 *
          (409447/40 +
           skoXC1 * (2681/8 + skoXC1 * (3163/960 + skoXC1 * 1/160)))))))) |
     1 + skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5))) <= 0 |
     skoXC1 <= 1), inference(resolve, [], [refute_0_261, refute_0_272])).

cnf(refute_0_274, plain,
    (72704/15 <
     skoXC1 *
     (-72704/5 +
      skoXC1 *
      (17208 +
       skoXC1 *
       (-86182/15 +
        skoXC1 *
        (4237503/320 +
         skoXC1 *
         (-1147967/480 +
          skoXC1 *
          (-106553/40 +
           skoXC1 *
           (-2079/8 + skoXC1 * (-1637/960 + skoXC1 * 1/160)))))))) |
     skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5))) <= -1 |
     skoXC1 <= 1), inference(arithmetic, [], [refute_0_273])).

cnf(refute_0_275, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <
     1/32 +
     skoXC1 *
     (53/24 +
      skoXC1 *
      (1951/48 +
       skoXC1 *
       (2969/24 +
        skoXC1 *
        (477/8 +
         skoXC1 *
         (-967/8 +
          skoXC1 * (-4279/48 + skoXC1 * (-377/24 + skoXC1 * -47/96))))))) |
     skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1))) <= 0 |
     (1/32 +
      skoXC1 *
      (53/24 +
       skoXC1 *
       (1951/48 +
        skoXC1 *
        (2969/24 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (-967/8 +
           skoXC1 *
           (-4279/48 + skoXC1 * (-377/24 + skoXC1 * -47/96)))))))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <=
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_276, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     (-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_277, plain,
    (~ lgen(0, X_000105, ln(X_000104)) | X_000105 <= ln(X_000104)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_278, plain,
    (~ lgen(0, X_000105,
            1/12 *
            (47 * X_000104 ^ 3 + 239 * X_000104 ^ 2 + 131 * X_000104 + 3) *
            (X_000104 - 1) /
            (X_000104 *
             (X_000104 ^ 3 + 12 * X_000104 ^ 2 + 18 * X_000104 + 4))) |
     X_000104 <= 0 | lgen(0, X_000105, ln(X_000104))),
    inference(subst, [], [ln_lower_bound_cf7])).

cnf(refute_0_279, plain,
    (~ lgen(0, X_000105,
            1/12 *
            (47 * X_000104 ^ 3 + 239 * X_000104 ^ 2 + 131 * X_000104 + 3) *
            (X_000104 - 1) /
            (X_000104 *
             (X_000104 ^ 3 + 12 * X_000104 ^ 2 + 18 * X_000104 + 4))) |
     X_000104 <= 0 | X_000105 <= ln(X_000104)),
    inference(resolve, [], [refute_0_278, refute_0_277])).

cnf(refute_0_280, plain,
    ((-1/4 +
      X_000104 *
      (-32/3 + X_000104 * (-9 + X_000104 * (16 + X_000104 * 47/12)))) /
     (X_000104 * (4 + X_000104 * (18 + X_000104 * (12 + X_000104)))) <
     X_000105 | X_000104 <= 0 | X_000105 <= ln(X_000104)),
    inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    ((-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     ln(skoXC1) | skoXC1 <= 0), inference(subst, [], [refute_0_280])).

cnf(refute_0_282, plain,
    ((-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_281, refute_0_220])).

cnf(refute_0_283, plain,
    ((-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_282, refute_0_76])).

cnf(refute_0_284, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/4 +
      skoXC1 * (-32/3 + skoXC1 * (-9 + skoXC1 * (16 + skoXC1 * 47/12)))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_276, refute_0_283])).

cnf(refute_0_285, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/32 +
      skoXC1 *
      (53/24 +
       skoXC1 *
       (1951/48 +
        skoXC1 *
        (2969/24 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (-967/8 +
           skoXC1 *
           (-4279/48 + skoXC1 * (-377/24 + skoXC1 * -47/96)))))))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) |
     1 <= skoXC1 |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8), inference(arithmetic, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/32 +
      skoXC1 *
      (53/24 +
       skoXC1 *
       (1951/48 +
        skoXC1 *
        (2969/24 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (-967/8 +
           skoXC1 *
           (-4279/48 + skoXC1 * (-377/24 + skoXC1 * -47/96)))))))) /
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_285, refute_0_219])).

cnf(refute_0_287, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1)))) <
     1/32 +
     skoXC1 *
     (53/24 +
      skoXC1 *
      (1951/48 +
       skoXC1 *
       (2969/24 +
        skoXC1 *
        (477/8 +
         skoXC1 *
         (-967/8 +
          skoXC1 * (-4279/48 + skoXC1 * (-377/24 + skoXC1 * -47/96))))))) |
     1 <= skoXC1 |
     skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1))) <= 0),
    inference(resolve, [], [refute_0_275, refute_0_286])).

cnf(refute_0_288, plain,
    (-1/32 <
     skoXC1 *
     (-43/24 +
      skoXC1 *
      (-65/48 +
       skoXC1 *
       (89/24 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (-119/8 +
          skoXC1 * (41/48 + skoXC1 * (55/24 + skoXC1 * 49/96))))))) |
     1 <= skoXC1 |
     skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1))) <= 0),
    inference(arithmetic, [], [refute_0_287])).

cnf(refute_0_289, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <
     -24064/3 +
     skoXC1 *
     (-140416/3 +
      skoXC1 *
      (-123316/3 +
       skoXC1 *
       (90173/2 +
        skoXC1 *
        (5575323/128 +
         skoXC1 *
         (41969/6 +
          skoXC1 * (27803/96 + skoXC1 * (23/6 + skoXC1 * 1/128))))))) |
     1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= 0 |
     (-24064/3 +
      skoXC1 *
      (-140416/3 +
       skoXC1 *
       (-123316/3 +
        skoXC1 *
        (90173/2 +
         skoXC1 *
         (5575323/128 +
          skoXC1 *
          (41969/6 +
           skoXC1 * (27803/96 + skoXC1 * (23/6 + skoXC1 * 1/128)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <=
     -4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_290, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 |
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_291, plain,
    (~ lgen(0, ln(X_000097), X_000098) | ln(X_000097) <= X_000098),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_292, plain,
    (~ lgen(0,
            1/12 *
            (3 * X_000097 ^ 3 + 131 * X_000097 ^ 2 + 239 * X_000097 + 47) *
            (X_000097 - 1) /
            (4 * X_000097 ^ 3 + 18 * X_000097 ^ 2 + 12 * X_000097 + 1),
            X_000098) | X_000097 <= 0 | lgen(0, ln(X_000097), X_000098)),
    inference(subst, [], [ln_upper_bound_cf7])).

cnf(refute_0_293, plain,
    (~ lgen(0,
            1/12 *
            (3 * X_000097 ^ 3 + 131 * X_000097 ^ 2 + 239 * X_000097 + 47) *
            (X_000097 - 1) /
            (4 * X_000097 ^ 3 + 18 * X_000097 ^ 2 + 12 * X_000097 + 1),
            X_000098) | X_000097 <= 0 | ln(X_000097) <= X_000098),
    inference(resolve, [], [refute_0_292, refute_0_291])).

cnf(refute_0_294, plain,
    (X_000098 <
     (-47/12 +
      X_000097 *
      (-16 + X_000097 * (9 + X_000097 * (32/3 + X_000097 * 1/4)))) /
     (1 + X_000097 * (12 + X_000097 * (18 + X_000097 * 4))) |
     X_000097 <= 0 | ln(X_000097) <= X_000098),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) |
     ln(skoXC1) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_294])).

cnf(refute_0_296, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 0 |
     skoXC1 <= 1), inference(resolve, [], [refute_0_295, refute_0_179])).

cnf(refute_0_297, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_296, refute_0_29])).

cnf(refute_0_298, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_290, refute_0_297])).

cnf(refute_0_299, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-24064/3 +
      skoXC1 *
      (-140416/3 +
       skoXC1 *
       (-123316/3 +
        skoXC1 *
        (90173/2 +
         skoXC1 *
         (5575323/128 +
          skoXC1 *
          (41969/6 +
           skoXC1 * (27803/96 + skoXC1 * (23/6 + skoXC1 * 1/128)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-24064/3 +
      skoXC1 *
      (-140416/3 +
       skoXC1 *
       (-123316/3 +
        skoXC1 *
        (90173/2 +
         skoXC1 *
         (5575323/128 +
          skoXC1 *
          (41969/6 +
           skoXC1 * (27803/96 + skoXC1 * (23/6 + skoXC1 * 1/128)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_299, refute_0_178])).

cnf(refute_0_301, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <
     -24064/3 +
     skoXC1 *
     (-140416/3 +
      skoXC1 *
      (-123316/3 +
       skoXC1 *
       (90173/2 +
        skoXC1 *
        (5575323/128 +
         skoXC1 *
         (41969/6 +
          skoXC1 * (27803/96 + skoXC1 * (23/6 + skoXC1 * 1/128))))))) |
     1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_289, refute_0_300])).

cnf(refute_0_302, plain,
    (11776/3 <
     skoXC1 *
     (128/3 +
      skoXC1 *
      (9884/3 +
       skoXC1 *
       (-545/2 +
        skoXC1 *
        (354331/128 +
         skoXC1 *
         (-10411/6 +
          skoXC1 * (-16549/96 + skoXC1 * (-1/6 + skoXC1 * 1/128))))))) |
     skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= -1 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_301])).

cnf(refute_0_303, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <
     -47/96 +
     skoXC1 *
     (-377/24 +
      skoXC1 *
      (-4279/48 +
       skoXC1 *
       (-967/8 +
        skoXC1 *
        (477/8 +
         skoXC1 *
         (2969/24 +
          skoXC1 * (1951/48 + skoXC1 * (53/24 + skoXC1 * 1/32))))))) |
     1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= 0 |
     (-47/96 +
      skoXC1 *
      (-377/24 +
       skoXC1 *
       (-4279/48 +
        skoXC1 *
        (-967/8 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (2969/24 +
           skoXC1 * (1951/48 + skoXC1 * (53/24 + skoXC1 * 1/32)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <=
     -1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_304, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 |
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_305, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) |
     ln(skoXC1) <=
     (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_294])).

cnf(refute_0_306, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 0 |
     skoXC1 <= 1), inference(resolve, [], [refute_0_305, refute_0_137])).

cnf(refute_0_307, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) /
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_306, refute_0_29])).

cnf(refute_0_308, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-47/12 +
      skoXC1 * (-16 + skoXC1 * (9 + skoXC1 * (32/3 + skoXC1 * 1/4)))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) *
     (1/8 +
      skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))) |
     1/8 +
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_304, refute_0_307])).

cnf(refute_0_309, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-47/96 +
      skoXC1 *
      (-377/24 +
       skoXC1 *
       (-4279/48 +
        skoXC1 *
        (-967/8 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (2969/24 +
           skoXC1 * (1951/48 + skoXC1 * (53/24 + skoXC1 * 1/32)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_308])).

cnf(refute_0_310, plain,
    (-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1))) <
     (-47/96 +
      skoXC1 *
      (-377/24 +
       skoXC1 *
       (-4279/48 +
        skoXC1 *
        (-967/8 +
         skoXC1 *
         (477/8 +
          skoXC1 *
          (2969/24 +
           skoXC1 * (1951/48 + skoXC1 * (53/24 + skoXC1 * 1/32)))))))) /
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_309, refute_0_136])).

cnf(refute_0_311, plain,
    ((-1 + skoXC1 * (-6 + skoXC1 * (skoXC1 * (6 + skoXC1)))) *
     (1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4))) <
     -47/96 +
     skoXC1 *
     (-377/24 +
      skoXC1 *
      (-4279/48 +
       skoXC1 *
       (-967/8 +
        skoXC1 *
        (477/8 +
         skoXC1 *
         (2969/24 +
          skoXC1 * (1951/48 + skoXC1 * (53/24 + skoXC1 * 1/32))))))) |
     1 + skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_303, refute_0_310])).

cnf(refute_0_312, plain,
    (-49/96 <
     skoXC1 *
     (55/24 +
      skoXC1 *
      (41/48 +
       skoXC1 *
       (-119/8 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (89/24 +
          skoXC1 * (-65/48 + skoXC1 * (-43/24 + skoXC1 * 1/32))))))) |
     skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) <= -1 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_311])).

cnf(refute_0_313, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (6 + skoXC1 * 3)) <
     -20480/3 +
     skoXC1 *
     (-54272/3 +
      skoXC1 *
      (-992/3 +
       skoXC1 *
       (20460 +
        skoXC1 *
        (72201/16 +
         skoXC1 * (25975/96 + skoXC1 * (233/48 + skoXC1 * 1/96)))))) |
     1 + skoXC1 * (6 + skoXC1 * 3) <= 0 |
     (-20480/3 +
      skoXC1 *
      (-54272/3 +
       skoXC1 *
       (-992/3 +
        skoXC1 *
        (20460 +
         skoXC1 *
         (72201/16 +
          skoXC1 * (25975/96 + skoXC1 * (233/48 + skoXC1 * 1/96))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) <=
     -4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_314, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 |
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_315, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) |
     ln(skoXC1) <=
     (-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     skoXC1 <= 0), inference(subst, [], [refute_0_45])).

cnf(refute_0_316, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_315, refute_0_179])).

cnf(refute_0_317, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) /
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_316, refute_0_29])).

cnf(refute_0_318, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-10/3 + skoXC1 * (-3 + skoXC1 * (6 + skoXC1 * 1/3))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) *
     (2048 +
      skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32)))) |
     2048 +
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_314, refute_0_317])).

cnf(refute_0_319, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-20480/3 +
      skoXC1 *
      (-54272/3 +
       skoXC1 *
       (-992/3 +
        skoXC1 *
        (20460 +
         skoXC1 *
         (72201/16 +
          skoXC1 * (25975/96 + skoXC1 * (233/48 + skoXC1 * 1/96))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) |
     skoXC1 * (3584 + skoXC1 * (560 + skoXC1 * (14 + skoXC1 * 1/32))) <=
     -2048 | skoXC1 <= 1), inference(arithmetic, [], [refute_0_318])).

cnf(refute_0_320, plain,
    (-4096 + skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1))) <
     (-20480/3 +
      skoXC1 *
      (-54272/3 +
       skoXC1 *
       (-992/3 +
        skoXC1 *
        (20460 +
         skoXC1 *
         (72201/16 +
          skoXC1 * (25975/96 + skoXC1 * (233/48 + skoXC1 * 1/96))))))) /
     (1 + skoXC1 * (6 + skoXC1 * 3)) | skoXC1 <= 1),
    inference(resolve, [], [refute_0_319, refute_0_178])).

cnf(refute_0_321, plain,
    ((-4096 +
      skoXC1 * (2304 + skoXC1 * (1680 + skoXC1 * (111 + skoXC1)))) *
     (1 + skoXC1 * (6 + skoXC1 * 3)) <
     -20480/3 +
     skoXC1 *
     (-54272/3 +
      skoXC1 *
      (-992/3 +
       skoXC1 *
       (20460 +
        skoXC1 *
        (72201/16 +
         skoXC1 * (25975/96 + skoXC1 * (233/48 + skoXC1 * 1/96)))))) |
     1 + skoXC1 * (6 + skoXC1 * 3) <= 0 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_313, refute_0_320])).

cnf(refute_0_322, plain,
    (8192/3 <
     skoXC1 *
     (12544/3 +
      skoXC1 *
      (-10640/3 +
       skoXC1 *
       (3357 +
        skoXC1 *
        (-19111/16 +
         skoXC1 * (-6569/96 + skoXC1 * (89/48 + skoXC1 * 1/96)))))) |
     skoXC1 * (6 + skoXC1 * 3) <= -1 | skoXC1 <= 1),
    inference(arithmetic, [], [refute_0_321])).

cnf(refute_0_323, plain,
    (-1 < skoXC1 * (6 + skoXC1 * 3) |
     8192/3 <
     skoXC1 *
     (12544/3 +
      skoXC1 *
      (-10640/3 +
       skoXC1 *
       (3357 +
        skoXC1 *
        (-19111/16 +
         skoXC1 * (-6569/96 + skoXC1 * (89/48 + skoXC1 * 1/96)))))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_324, plain,
    (8192/3 <
     skoXC1 *
     (12544/3 +
      skoXC1 *
      (-10640/3 +
       skoXC1 *
       (3357 +
        skoXC1 *
        (-19111/16 +
         skoXC1 * (-6569/96 + skoXC1 * (89/48 + skoXC1 * 1/96)))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_322, refute_0_323])).

cnf(refute_0_325, plain,
    (-49/96 <
     skoXC1 *
     (55/24 +
      skoXC1 *
      (41/48 +
       skoXC1 *
       (-119/8 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (89/24 +
          skoXC1 * (-65/48 + skoXC1 * (-43/24 + skoXC1 * 1/32))))))) |
     -1 < skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_326, plain,
    (-49/96 <
     skoXC1 *
     (55/24 +
      skoXC1 *
      (41/48 +
       skoXC1 *
       (-119/8 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (89/24 +
          skoXC1 * (-65/48 + skoXC1 * (-43/24 + skoXC1 * 1/32))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_312, refute_0_325])).

cnf(refute_0_327, plain,
    (-49/96 <
     skoXC1 *
     (55/24 +
      skoXC1 *
      (41/48 +
       skoXC1 *
       (-119/8 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (89/24 +
          skoXC1 * (-65/48 + skoXC1 * (-43/24 + skoXC1 * 1/32))))))) |
     1 < skoXC1),
    inference(decision, [],
              [refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_328, plain,
    (-49/96 <
     skoXC1 *
     (55/24 +
      skoXC1 *
      (41/48 +
       skoXC1 *
       (-119/8 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (89/24 +
          skoXC1 * (-65/48 + skoXC1 * (-43/24 + skoXC1 * 1/32)))))))),
    inference(resolve, [], [refute_0_326, refute_0_327])).

cnf(refute_0_329, plain,
    (-1 < skoXC1 * (12 + skoXC1 * (18 + skoXC1 * 4)) |
     11776/3 <
     skoXC1 *
     (128/3 +
      skoXC1 *
      (9884/3 +
       skoXC1 *
       (-545/2 +
        skoXC1 *
        (354331/128 +
         skoXC1 *
         (-10411/6 +
          skoXC1 * (-16549/96 + skoXC1 * (-1/6 + skoXC1 * 1/128))))))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_330, plain,
    (11776/3 <
     skoXC1 *
     (128/3 +
      skoXC1 *
      (9884/3 +
       skoXC1 *
       (-545/2 +
        skoXC1 *
        (354331/128 +
         skoXC1 *
         (-10411/6 +
          skoXC1 * (-16549/96 + skoXC1 * (-1/6 + skoXC1 * 1/128))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_302, refute_0_329])).

cnf(refute_0_331, plain,
    (-1/32 <
     skoXC1 *
     (-43/24 +
      skoXC1 *
      (-65/48 +
       skoXC1 *
       (89/24 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (-119/8 +
          skoXC1 * (41/48 + skoXC1 * (55/24 + skoXC1 * 49/96))))))) |
     0 < skoXC1 * (4 + skoXC1 * (18 + skoXC1 * (12 + skoXC1))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_332, plain,
    (-1/32 <
     skoXC1 *
     (-43/24 +
      skoXC1 *
      (-65/48 +
       skoXC1 *
       (89/24 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (-119/8 +
          skoXC1 * (41/48 + skoXC1 * (55/24 + skoXC1 * 49/96))))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_288, refute_0_331])).

cnf(refute_0_333, plain,
    (skoXC1 < 1 |
     -1/32 <
     skoXC1 *
     (-43/24 +
      skoXC1 *
      (-65/48 +
       skoXC1 *
       (89/24 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (-119/8 +
          skoXC1 * (41/48 + skoXC1 * (55/24 + skoXC1 * 49/96)))))))),
    inference(decision, [],
              [refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_334, plain,
    (-1/32 <
     skoXC1 *
     (-43/24 +
      skoXC1 *
      (-65/48 +
       skoXC1 *
       (89/24 +
        skoXC1 *
        (85/8 +
         skoXC1 *
         (-119/8 +
          skoXC1 * (41/48 + skoXC1 * (55/24 + skoXC1 * 49/96)))))))),
    inference(resolve, [], [refute_0_332, refute_0_333])).

cnf(refute_0_335, plain,
    (72704/15 <
     skoXC1 *
     (-72704/5 +
      skoXC1 *
      (17208 +
       skoXC1 *
       (-86182/15 +
        skoXC1 *
        (4237503/320 +
         skoXC1 *
         (-1147967/480 +
          skoXC1 *
          (-106553/40 +
           skoXC1 *
           (-2079/8 + skoXC1 * (-1637/960 + skoXC1 * 1/160)))))))) |
     -1 < skoXC1 * (20 + skoXC1 * (60 + skoXC1 * (40 + skoXC1 * 5))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_336, plain,
    (72704/15 <
     skoXC1 *
     (-72704/5 +
      skoXC1 *
      (17208 +
       skoXC1 *
       (-86182/15 +
        skoXC1 *
        (4237503/320 +
         skoXC1 *
         (-1147967/480 +
          skoXC1 *
          (-106553/40 +
           skoXC1 *
           (-2079/8 + skoXC1 * (-1637/960 + skoXC1 * 1/160)))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_274, refute_0_335])).

cnf(refute_0_337, plain,
    (-1/40 <
     skoXC1 *
     (-557/240 +
      skoXC1 *
      (-7 +
       skoXC1 *
       (3/40 +
        skoXC1 *
        (2339/60 +
         skoXC1 *
         (-102/5 +
          skoXC1 *
          (-313/15 +
           skoXC1 * (41/8 + skoXC1 * (237/40 + skoXC1 * 109/240)))))))) |
     0 <
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_338, plain,
    (-1/40 <
     skoXC1 *
     (-557/240 +
      skoXC1 *
      (-7 +
       skoXC1 *
       (3/40 +
        skoXC1 *
        (2339/60 +
         skoXC1 *
         (-102/5 +
          skoXC1 *
          (-313/15 +
           skoXC1 * (41/8 + skoXC1 * (237/40 + skoXC1 * 109/240)))))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_260, refute_0_337])).

cnf(refute_0_339, plain,
    (skoXC1 < 1 |
     -1/40 <
     skoXC1 *
     (-557/240 +
      skoXC1 *
      (-7 +
       skoXC1 *
       (3/40 +
        skoXC1 *
        (2339/60 +
         skoXC1 *
         (-102/5 +
          skoXC1 *
          (-313/15 +
           skoXC1 * (41/8 + skoXC1 * (237/40 + skoXC1 * 109/240))))))))),
    inference(decision, [],
              [refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_340, plain,
    (-1/40 <
     skoXC1 *
     (-557/240 +
      skoXC1 *
      (-7 +
       skoXC1 *
       (3/40 +
        skoXC1 *
        (2339/60 +
         skoXC1 *
         (-102/5 +
          skoXC1 *
          (-313/15 +
           skoXC1 * (41/8 + skoXC1 * (237/40 + skoXC1 * 109/240))))))))),
    inference(resolve, [], [refute_0_338, refute_0_339])).

cnf(refute_0_341, plain,
    (83968/15 <
     skoXC1 *
     (-626944/15 +
      skoXC1 *
      (14656/3 +
       skoXC1 *
       (404941/15 +
        skoXC1 *
        (3843929/480 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (-13637789/960 +
           skoXC1 *
           (-51259/12 +
            skoXC1 *
            (-165197/480 + skoXC1 * (-1439/480 + skoXC1 * 1/192))))))))) |
     -1 <
     skoXC1 *
     (30 + skoXC1 * (150 + skoXC1 * (200 + skoXC1 * (75 + skoXC1 * 6)))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_342, plain,
    (83968/15 <
     skoXC1 *
     (-626944/15 +
      skoXC1 *
      (14656/3 +
       skoXC1 *
       (404941/15 +
        skoXC1 *
        (3843929/480 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (-13637789/960 +
           skoXC1 *
           (-51259/12 +
            skoXC1 *
            (-165197/480 + skoXC1 * (-1439/480 + skoXC1 * 1/192))))))))) |
     skoXC1 <= 1), inference(resolve, [], [refute_0_246, refute_0_341])).

cnf(refute_0_343, plain,
    (-1/48 <
     skoXC1 *
     (-329/120 +
      skoXC1 *
      (-1141/60 +
       skoXC1 *
       (-49/3 +
        skoXC1 *
        (4249/60 +
         skoXC1 *
         (3563/60 +
          skoXC1 *
          (-13111/120 +
           skoXC1 *
           (-251/15 +
            skoXC1 *
            (1099/48 + skoXC1 * (1267/120 + skoXC1 * 49/120))))))))) |
     0 <
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_342, refute_0_340, refute_0_336, refute_0_334,
               refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_344, plain,
    (-1/48 <
     skoXC1 *
     (-329/120 +
      skoXC1 *
      (-1141/60 +
       skoXC1 *
       (-49/3 +
        skoXC1 *
        (4249/60 +
         skoXC1 *
         (3563/60 +
          skoXC1 *
          (-13111/120 +
           skoXC1 *
           (-251/15 +
            skoXC1 *
            (1099/48 + skoXC1 * (1267/120 + skoXC1 * 49/120))))))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_232, refute_0_343])).

cnf(refute_0_345, plain,
    (skoXC1 < 1 |
     -1/48 <
     skoXC1 *
     (-329/120 +
      skoXC1 *
      (-1141/60 +
       skoXC1 *
       (-49/3 +
        skoXC1 *
        (4249/60 +
         skoXC1 *
         (3563/60 +
          skoXC1 *
          (-13111/120 +
           skoXC1 *
           (-251/15 +
            skoXC1 *
            (1099/48 + skoXC1 * (1267/120 + skoXC1 * 49/120)))))))))),
    inference(decision, [],
              [refute_0_342, refute_0_340, refute_0_336, refute_0_334,
               refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_346, plain,
    (-1/48 <
     skoXC1 *
     (-329/120 +
      skoXC1 *
      (-1141/60 +
       skoXC1 *
       (-49/3 +
        skoXC1 *
        (4249/60 +
         skoXC1 *
         (3563/60 +
          skoXC1 *
          (-13111/120 +
           skoXC1 *
           (-251/15 +
            skoXC1 *
            (1099/48 + skoXC1 * (1267/120 + skoXC1 * 49/120)))))))))),
    inference(resolve, [], [refute_0_344, refute_0_345])).

cnf(refute_0_347, plain,
    (skoXC1 *
     (-83456 +
      skoXC1 *
      (-520696/5 +
       skoXC1 *
       (525492/5 +
        skoXC1 *
        (66863199/2240 +
         skoXC1 *
         (7420609/80 +
          skoXC1 *
          (644707/64 +
           skoXC1 *
           (-41108141/1120 +
            skoXC1 *
            (-2179753/320 +
             skoXC1 *
             (-34613/80 + skoXC1 * (-1327/320 + skoXC1 * 1/224)))))))))) <=
     218112/35 |
     skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) <=
     -1 | skoXC1 <= 1),
    inference(decision, [],
              [refute_0_346, refute_0_342, refute_0_340, refute_0_336,
               refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_348, plain,
    (skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) <=
     -1 | skoXC1 <= 1),
    inference(resolve, [], [refute_0_347, refute_0_191])).

cnf(refute_0_349, plain,
    (-1 <
     skoXC1 *
     (42 +
      skoXC1 *
      (315 +
       skoXC1 * (700 + skoXC1 * (525 + skoXC1 * (126 + skoXC1 * 7))))) |
     skoXC1 <= 1),
    inference(decision, [],
              [refute_0_346, refute_0_342, refute_0_340, refute_0_336,
               refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_350, plain, (skoXC1 <= 1),
    inference(resolve, [], [refute_0_348, refute_0_349])).

cnf(refute_0_351, plain, (skoXC1 < 1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_350, refute_0_346, refute_0_342, refute_0_340,
               refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_352, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_14, refute_0_351])).

cnf(refute_0_353, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (-1 + skoXC1) / sqrt(skoXC1)),
    inference(resolve, [], [refute_0_352, refute_0_199])).

cnf(refute_0_354, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_3, refute_0_353])).

cnf(refute_0_355, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-119 +
       skoXC1 *
       (-399 +
        skoXC1 *
        (-175 +
         skoXC1 * (385 + skoXC1 * (273 + skoXC1 * (41 + skoXC1))))))) /
     sqrt(skoXC1) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0),
    inference(arithmetic, [], [refute_0_354])).

cnf(refute_0_356, plain,
    (0 <
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))),
    inference(decision, [],
              [refute_0_350, refute_0_346, refute_0_342, refute_0_340,
               refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_357, plain,
    (-1/7 +
     skoXC1 *
     (-273/10 +
      skoXC1 *
      (-1302/5 +
       skoXC1 *
       (-875/2 +
        skoXC1 *
        (175 + skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70)))))) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-119 +
       skoXC1 *
       (-399 +
        skoXC1 *
        (-175 +
         skoXC1 * (385 + skoXC1 * (273 + skoXC1 * (41 + skoXC1))))))) /
     sqrt(skoXC1)), inference(resolve, [], [refute_0_355, refute_0_356])).

cnf(refute_0_358, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-7 +
      skoXC1 *
      (-119 +
       skoXC1 *
       (-399 +
        skoXC1 *
        (-175 +
         skoXC1 * (385 + skoXC1 * (273 + skoXC1 * (41 + skoXC1))))))) |
     sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_2, refute_0_357])).

cnf(refute_0_359, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     sqrt(skoXC1) *
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     sqrt(skoXC1) *
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))))),
    inference(resolve, [], [refute_0_359, refute_0_199])).

cnf(refute_0_361, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     sqrt(skoXC1) |
     1/7 +
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))) <=
     0), inference(split, [], [refute_0_360])).

cnf(refute_0_362, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     sqrt(skoXC1) |
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))) <=
     -1/7), inference(arithmetic, [], [refute_0_361])).

cnf(refute_0_363, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     1/56 +
     skoXC1 *
     (313/80 +
      skoXC1 *
      (2587/20 +
       skoXC1 *
       (96437/80 +
        skoXC1 *
        (271841/70 +
         skoXC1 *
         (32613/8 +
          skoXC1 *
          (-29771/20 +
           skoXC1 *
           (-1323593/280 +
            skoXC1 *
            (-102031/40 +
             skoXC1 *
             (-39979/80 + skoXC1 * (-32 + skoXC1 * -353/560)))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0 |
     (1/56 +
      skoXC1 *
      (313/80 +
       skoXC1 *
       (2587/20 +
        skoXC1 *
        (96437/80 +
         skoXC1 *
         (271841/70 +
          skoXC1 *
          (32613/8 +
           skoXC1 *
           (-29771/20 +
            skoXC1 *
            (-1323593/280 +
             skoXC1 *
             (-102031/40 +
              skoXC1 *
              (-39979/80 + skoXC1 * (-32 + skoXC1 * -353/560))))))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <=
     1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_364, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_365, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) <=
     ln(skoXC1) | skoXC1 <= 0), inference(subst, [], [refute_0_12])).

cnf(refute_0_366, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_365, refute_0_220])).

cnf(refute_0_367, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_366, refute_0_351])).

cnf(refute_0_368, plain,
    ((-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) /
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8))))),
    inference(resolve, [], [refute_0_367, refute_0_199])).

cnf(refute_0_369, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (-1/7 +
      skoXC1 *
      (-273/10 +
       skoXC1 *
       (-1302/5 +
        skoXC1 *
        (-875/2 +
         skoXC1 *
         (175 +
          skoXC1 * (861/2 + skoXC1 * (574/5 + skoXC1 * 353/70))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) *
     (-1/8 +
      skoXC1 *
      (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))) |
     0 <=
     -1/8 +
     skoXC1 * (-7/2 + skoXC1 * (-35/4 + skoXC1 * (-7/2 + skoXC1 * -1/8)))),
    inference(resolve, [], [refute_0_364, refute_0_368])).

cnf(refute_0_370, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/56 +
      skoXC1 *
      (313/80 +
       skoXC1 *
       (2587/20 +
        skoXC1 *
        (96437/80 +
         skoXC1 *
         (271841/70 +
          skoXC1 *
          (32613/8 +
           skoXC1 *
           (-29771/20 +
            skoXC1 *
            (-1323593/280 +
             skoXC1 *
             (-102031/40 +
              skoXC1 *
              (-39979/80 + skoXC1 * (-32 + skoXC1 * -353/560))))))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) |
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8))) <=
     -1/8), inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (-1/8 <
     skoXC1 * (7/2 + skoXC1 * (35/4 + skoXC1 * (7/2 + skoXC1 * 1/8)))),
    inference(decision, [],
              [refute_0_350, refute_0_346, refute_0_342, refute_0_340,
               refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_372, plain,
    (1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1))) <
     (1/56 +
      skoXC1 *
      (313/80 +
       skoXC1 *
       (2587/20 +
        skoXC1 *
        (96437/80 +
         skoXC1 *
         (271841/70 +
          skoXC1 *
          (32613/8 +
           skoXC1 *
           (-29771/20 +
            skoXC1 *
            (-1323593/280 +
             skoXC1 *
             (-102031/40 +
              skoXC1 *
              (-39979/80 + skoXC1 * (-32 + skoXC1 * -353/560))))))))))) /
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))))),
    inference(resolve, [], [refute_0_370, refute_0_371])).

cnf(refute_0_373, plain,
    ((1 + skoXC1 * (6 + skoXC1 * (skoXC1 * (-6 + skoXC1 * -1)))) *
     (skoXC1 *
      (7 +
       skoXC1 *
       (126 +
        skoXC1 *
        (525 +
         skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))) <
     1/56 +
     skoXC1 *
     (313/80 +
      skoXC1 *
      (2587/20 +
       skoXC1 *
       (96437/80 +
        skoXC1 *
        (271841/70 +
         skoXC1 *
         (32613/8 +
          skoXC1 *
          (-29771/20 +
           skoXC1 *
           (-1323593/280 +
            skoXC1 *
            (-102031/40 +
             skoXC1 *
             (-39979/80 + skoXC1 * (-32 + skoXC1 * -353/560)))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0),
    inference(resolve, [], [refute_0_363, refute_0_372])).

cnf(refute_0_374, plain,
    (-1/56 <
     skoXC1 *
     (-247/80 +
      skoXC1 *
      (-773/20 +
       skoXC1 *
       (-6043/80 +
        skoXC1 *
        (5281/70 +
         skoXC1 *
         (2597/8 +
          skoXC1 *
          (-2891/20 +
           skoXC1 *
           (-71433/280 +
            skoXC1 *
            (1329/40 +
             skoXC1 *
             (5381/80 + skoXC1 * (16 + skoXC1 * 207/560)))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 +
        skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1)))))) <= 0),
    inference(arithmetic, [], [refute_0_373])).

cnf(refute_0_375, plain,
    (-1/56 <
     skoXC1 *
     (-247/80 +
      skoXC1 *
      (-773/20 +
       skoXC1 *
       (-6043/80 +
        skoXC1 *
        (5281/70 +
         skoXC1 *
         (2597/8 +
          skoXC1 *
          (-2891/20 +
           skoXC1 *
           (-71433/280 +
            skoXC1 *
            (1329/40 +
             skoXC1 *
             (5381/80 + skoXC1 * (16 + skoXC1 * 207/560)))))))))) |
     0 <
     skoXC1 *
     (7 +
      skoXC1 *
      (126 +
       skoXC1 *
       (525 + skoXC1 * (700 + skoXC1 * (315 + skoXC1 * (42 + skoXC1))))))),
    inference(decision, [],
              [refute_0_350, refute_0_346, refute_0_342, refute_0_340,
               refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_376, plain,
    (-1/56 <
     skoXC1 *
     (-247/80 +
      skoXC1 *
      (-773/20 +
       skoXC1 *
       (-6043/80 +
        skoXC1 *
        (5281/70 +
         skoXC1 *
         (2597/8 +
          skoXC1 *
          (-2891/20 +
           skoXC1 *
           (-71433/280 +
            skoXC1 *
            (1329/40 +
             skoXC1 *
             (5381/80 + skoXC1 * (16 + skoXC1 * 207/560))))))))))),
    inference(resolve, [], [refute_0_374, refute_0_375])).

cnf(refute_0_377, plain,
    (-1/7 <
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))),
    inference(decision, [],
              [refute_0_376, refute_0_350, refute_0_346, refute_0_342,
               refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_378, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     sqrt(skoXC1)), inference(resolve, [], [refute_0_362, refute_0_377])).

cnf(refute_0_379, plain,
    (~ lgen(0, sqrt(X_000177), X_000178) | sqrt(X_000177) <= X_000178),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_380, plain,
    (X_000177 < 0 |
     ~ lgen(0,
            1/32 *
            (65536 * X_000177 ^ 4 + 114688 * X_000177 ^ 3 +
             17920 * X_000177 ^ 2 + 448 * X_000177 + 1) /
            ((16 * X_000177 + 1) *
             (256 * X_000177 ^ 2 + 96 * X_000177 + 1)), X_000178) |
     1 <= X_000177 | lgen(0, sqrt(X_000177), X_000178)),
    inference(subst, [], [sqrt_upper_bound_2_small])).

cnf(refute_0_381, plain,
    (X_000177 < 0 |
     ~ lgen(0,
            1/32 *
            (65536 * X_000177 ^ 4 + 114688 * X_000177 ^ 3 +
             17920 * X_000177 ^ 2 + 448 * X_000177 + 1) /
            ((16 * X_000177 + 1) *
             (256 * X_000177 ^ 2 + 96 * X_000177 + 1)), X_000178) |
     1 <= X_000177 | sqrt(X_000177) <= X_000178),
    inference(resolve, [], [refute_0_380, refute_0_379])).

cnf(refute_0_382, plain,
    (X_000177 < 0 |
     X_000178 <
     (1/32 +
      X_000177 *
      (14 + X_000177 * (560 + X_000177 * (3584 + X_000177 * 2048)))) /
     (1 + X_000177 * (112 + X_000177 * (1792 + X_000177 * 4096))) |
     1 <= X_000177 | sqrt(X_000177) <= X_000178),
    inference(arithmetic, [], [refute_0_381])).

cnf(refute_0_383, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1 |
     sqrt(skoXC1) <=
     skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))))),
    inference(subst, [], [refute_0_382])).

cnf(refute_0_384, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_383, refute_0_378])).

cnf(refute_0_385, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_384])).

cnf(refute_0_386, plain, (skoXC1 < 1),
    inference(decision, [],
              [refute_0_376, refute_0_350, refute_0_346, refute_0_342,
               refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_387, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) /
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_385, refute_0_386])).

cnf(refute_0_388, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/7 +
      skoXC1 *
      (273/10 +
       skoXC1 *
       (1302/5 +
        skoXC1 *
        (875/2 +
         skoXC1 *
         (-175 +
          skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70))))))) |
     1/7 +
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))) <=
     0), inference(resolve, [], [refute_0_1, refute_0_387])).

cnf(refute_0_389, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     (1/224 +
      skoXC1 *
      (913/320 +
       skoXC1 *
       (37627/80 +
        skoXC1 *
        (6226967/320 +
         skoXC1 *
         (280088819/1120 +
          skoXC1 *
          (78830115/64 +
           skoXC1 *
           (159781489/80 +
            skoXC1 *
            (58492319/2240 +
             skoXC1 *
             (-9828353/5 +
              skoXC1 *
              (-6479656/5 +
               skoXC1 * (-253184 + skoXC1 * -361472/35))))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) |
     skoXC1 *
     (273/10 +
      skoXC1 *
      (1302/5 +
       skoXC1 *
       (875/2 +
        skoXC1 *
        (-175 +
         skoXC1 * (-861/2 + skoXC1 * (-574/5 + skoXC1 * -353/70)))))) <=
     -1/7), inference(arithmetic, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) <
     (1/224 +
      skoXC1 *
      (913/320 +
       skoXC1 *
       (37627/80 +
        skoXC1 *
        (6226967/320 +
         skoXC1 *
         (280088819/1120 +
          skoXC1 *
          (78830115/64 +
           skoXC1 *
           (159781489/80 +
            skoXC1 *
            (58492319/2240 +
             skoXC1 *
             (-9828353/5 +
              skoXC1 *
              (-6479656/5 +
               skoXC1 * (-253184 + skoXC1 * -361472/35))))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_389, refute_0_377])).

cnf(refute_0_391, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (119 +
       skoXC1 *
       (399 +
        skoXC1 *
        (175 +
         skoXC1 *
         (-385 + skoXC1 * (-273 + skoXC1 * (-41 + skoXC1 * -1))))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/224 +
     skoXC1 *
     (913/320 +
      skoXC1 *
      (37627/80 +
       skoXC1 *
       (6226967/320 +
        skoXC1 *
        (280088819/1120 +
         skoXC1 *
         (78830115/64 +
          skoXC1 *
          (159781489/80 +
           skoXC1 *
           (58492319/2240 +
            skoXC1 *
            (-9828353/5 +
             skoXC1 *
             (-6479656/5 +
              skoXC1 * (-253184 + skoXC1 * -361472/35)))))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_390])).

cnf(refute_0_392, plain,
    (-1/224 <
     skoXC1 *
     (-1327/320 +
      skoXC1 *
      (-34613/80 +
       skoXC1 *
       (-2179753/320 +
        skoXC1 *
        (-41108141/1120 +
         skoXC1 *
         (644707/64 +
          skoXC1 *
          (7420609/80 +
           skoXC1 *
           (66863199/2240 +
            skoXC1 *
            (525492/5 +
             skoXC1 *
             (-520696/5 +
              skoXC1 * (-83456 + skoXC1 * -218112/35)))))))))) |
     skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= -1),
    inference(arithmetic, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/192 +
     skoXC1 *
     (1441/480 +
      skoXC1 *
      (190483/480 +
       skoXC1 *
       (172001/12 +
        skoXC1 *
        (142049251/960 +
         skoXC1 *
         (253071991/480 +
          skoXC1 *
          (200130809/480 +
           skoXC1 *
           (-5941394/15 +
            skoXC1 *
            (-1616144/3 +
             skoXC1 * (-2435584/15 + skoXC1 * -145408/15))))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0 |
     (1/192 +
      skoXC1 *
      (1441/480 +
       skoXC1 *
       (190483/480 +
        skoXC1 *
        (172001/12 +
         skoXC1 *
         (142049251/960 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (200130809/480 +
            skoXC1 *
            (-5941394/15 +
             skoXC1 *
             (-1616144/3 +
              skoXC1 * (-2435584/15 + skoXC1 * -145408/15)))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_394, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) |
     1/6 +
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) <=
     0 |
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_395, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-69 +
       skoXC1 *
       (-125 + skoXC1 * (50 + skoXC1 * (120 + skoXC1 * (29 + skoXC1)))))) |
     skoXC1 *
     (-6 +
      skoXC1 *
      (-69 +
       skoXC1 *
       (-125 + skoXC1 * (50 + skoXC1 * (120 + skoXC1 * (29 + skoXC1)))))) /
     sqrt(skoXC1) <=
     -1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) |
     sqrt(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_396, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0 |
     (-1 + skoXC1) / sqrt(skoXC1) <=
     (-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_397, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (-1 + skoXC1) / sqrt(skoXC1) |
     (-1 + skoXC1) / sqrt(skoXC1) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_224])).

cnf(refute_0_398, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_397, refute_0_8])).

cnf(refute_0_399, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_398, refute_0_76])).

cnf(refute_0_400, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1)))))) |
     1 <= skoXC1 |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0), inference(resolve, [], [refute_0_396, refute_0_399])).

cnf(refute_0_401, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-69 +
       skoXC1 *
       (-125 + skoXC1 * (50 + skoXC1 * (120 + skoXC1 * (29 + skoXC1)))))) /
     sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) <=
     0), inference(arithmetic, [], [refute_0_400])).

cnf(refute_0_402, plain,
    (0 <
     skoXC1 *
     (6 +
      skoXC1 *
      (75 + skoXC1 * (200 + skoXC1 * (150 + skoXC1 * (30 + skoXC1))))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_342, refute_0_340, refute_0_336, refute_0_334,
               refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_403, plain,
    (-1/6 +
     skoXC1 *
     (-107/5 +
      skoXC1 *
      (-125 +
       skoXC1 *
       (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15))))) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-69 +
       skoXC1 *
       (-125 + skoXC1 * (50 + skoXC1 * (120 + skoXC1 * (29 + skoXC1)))))) /
     sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_401, refute_0_402])).

cnf(refute_0_404, plain,
    ((-1/6 +
      skoXC1 *
      (-107/5 +
       skoXC1 *
       (-125 +
        skoXC1 *
        (-200/3 + skoXC1 * (275/2 + skoXC1 * (71 + skoXC1 * 71/15)))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-69 +
       skoXC1 *
       (-125 + skoXC1 * (50 + skoXC1 * (120 + skoXC1 * (29 + skoXC1)))))) |
     1 <= skoXC1 | sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_395, refute_0_403])).

cnf(refute_0_405, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     sqrt(skoXC1) *
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_404])).

cnf(refute_0_406, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     sqrt(skoXC1) *
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_405, refute_0_76])).

cnf(refute_0_407, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     sqrt(skoXC1) | 1 <= skoXC1 |
     1/6 +
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) <=
     0), inference(split, [], [refute_0_406])).

cnf(refute_0_408, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) <=
     -1/6), inference(arithmetic, [], [refute_0_407])).

cnf(refute_0_409, plain,
    (-1/6 <
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_346, refute_0_342, refute_0_340, refute_0_336,
               refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_410, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_408, refute_0_409])).

cnf(refute_0_411, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1 |
     sqrt(skoXC1) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))))),
    inference(subst, [], [refute_0_382])).

cnf(refute_0_412, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_411, refute_0_410])).

cnf(refute_0_413, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_412])).

cnf(refute_0_414, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) /
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_413, refute_0_386])).

cnf(refute_0_415, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/6 +
      skoXC1 *
      (107/5 +
       skoXC1 *
       (125 +
        skoXC1 *
        (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))) |
     1/6 +
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) <=
     0), inference(resolve, [], [refute_0_394, refute_0_414])).

cnf(refute_0_416, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     (1/192 +
      skoXC1 *
      (1441/480 +
       skoXC1 *
       (190483/480 +
        skoXC1 *
        (172001/12 +
         skoXC1 *
         (142049251/960 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (200130809/480 +
            skoXC1 *
            (-5941394/15 +
             skoXC1 *
             (-1616144/3 +
              skoXC1 * (-2435584/15 + skoXC1 * -145408/15)))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) |
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15))))) <=
     -1/6), inference(arithmetic, [], [refute_0_415])).

cnf(refute_0_417, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/160 +
     skoXC1 *
     (3163/960 +
      skoXC1 *
      (2681/8 +
       skoXC1 *
       (409447/40 +
        skoXC1 *
        (39844033/480 +
         skoXC1 *
         (60141823/320 +
          skoXC1 *
          (35483/15 +
           skoXC1 *
           (-180792 + skoXC1 * (-470784/5 + skoXC1 * -134144/15)))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0 |
     (1/160 +
      skoXC1 *
      (3163/960 +
       skoXC1 *
       (2681/8 +
        skoXC1 *
        (409447/40 +
         skoXC1 *
         (39844033/480 +
          skoXC1 *
          (60141823/320 +
           skoXC1 *
           (35483/15 +
            skoXC1 *
            (-180792 + skoXC1 * (-470784/5 + skoXC1 * -134144/15))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_418, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) |
     1/5 +
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) <=
     0 |
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <=
     skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_419, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-5 +
      skoXC1 *
      (-35 + skoXC1 * (-20 + skoXC1 * (40 + skoXC1 * (19 + skoXC1))))) |
     skoXC1 *
     (-5 +
      skoXC1 *
      (-35 + skoXC1 * (-20 + skoXC1 * (40 + skoXC1 * (19 + skoXC1))))) /
     sqrt(skoXC1) <=
     -1/5 +
     skoXC1 *
     (-95/6 +
      skoXC1 *
      (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30)))) |
     sqrt(skoXC1) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_420, plain,
    (-1/5 +
     skoXC1 *
     (-95/6 +
      skoXC1 *
      (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30)))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0 |
     (-1 + skoXC1) / sqrt(skoXC1) <=
     (-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_421, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (-1 + skoXC1) / sqrt(skoXC1) |
     (-1 + skoXC1) / sqrt(skoXC1) <= ln(skoXC1) | skoXC1 <= 0),
    inference(subst, [], [refute_0_252])).

cnf(refute_0_422, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_421, refute_0_8])).

cnf(refute_0_423, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) /
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) <
     (-1 + skoXC1) / sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_422, refute_0_76])).

cnf(refute_0_424, plain,
    (-1/5 +
     skoXC1 *
     (-95/6 +
      skoXC1 *
      (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30)))) <
     (-1 + skoXC1) / sqrt(skoXC1) *
     (skoXC1 *
      (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1))))) |
     1 <= skoXC1 |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0),
    inference(resolve, [], [refute_0_420, refute_0_423])).

cnf(refute_0_425, plain,
    (-1/5 +
     skoXC1 *
     (-95/6 +
      skoXC1 *
      (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30)))) <
     skoXC1 *
     (-5 +
      skoXC1 *
      (-35 + skoXC1 * (-20 + skoXC1 * (40 + skoXC1 * (19 + skoXC1))))) /
     sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) <= 0),
    inference(arithmetic, [], [refute_0_424])).

cnf(refute_0_426, plain,
    (0 <
     skoXC1 *
     (5 + skoXC1 * (40 + skoXC1 * (60 + skoXC1 * (20 + skoXC1)))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_336, refute_0_334, refute_0_330, refute_0_328,
               refute_0_324, refute_0_166, refute_0_162, refute_0_135,
               refute_0_131, refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_427, plain,
    (-1/5 +
     skoXC1 *
     (-95/6 +
      skoXC1 *
      (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30)))) <
     skoXC1 *
     (-5 +
      skoXC1 *
      (-35 + skoXC1 * (-20 + skoXC1 * (40 + skoXC1 * (19 + skoXC1))))) /
     sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_425, refute_0_426])).

cnf(refute_0_428, plain,
    ((-1/5 +
      skoXC1 *
      (-95/6 +
       skoXC1 *
       (-140/3 + skoXC1 * (20 + skoXC1 * (115/3 + skoXC1 * 131/30))))) *
     sqrt(skoXC1) <
     skoXC1 *
     (-5 +
      skoXC1 *
      (-35 + skoXC1 * (-20 + skoXC1 * (40 + skoXC1 * (19 + skoXC1))))) |
     1 <= skoXC1 | sqrt(skoXC1) <= 0),
    inference(resolve, [], [refute_0_419, refute_0_427])).

cnf(refute_0_429, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     sqrt(skoXC1) *
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) |
     1 <= skoXC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     sqrt(skoXC1) *
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) |
     1 <= skoXC1), inference(resolve, [], [refute_0_429, refute_0_76])).

cnf(refute_0_431, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     sqrt(skoXC1) | 1 <= skoXC1 |
     1/5 +
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) <=
     0), inference(split, [], [refute_0_430])).

cnf(refute_0_432, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     sqrt(skoXC1) | 1 <= skoXC1 |
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) <=
     -1/5), inference(arithmetic, [], [refute_0_431])).

cnf(refute_0_433, plain,
    (-1/5 <
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) |
     1 <= skoXC1),
    inference(decision, [],
              [refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_434, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     sqrt(skoXC1) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_432, refute_0_433])).

cnf(refute_0_435, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1 |
     sqrt(skoXC1) <=
     skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))))),
    inference(subst, [], [refute_0_382])).

cnf(refute_0_436, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_435, refute_0_434])).

cnf(refute_0_437, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) | 1 <= skoXC1),
    inference(resolve, [], [refute_0_91, refute_0_436])).

cnf(refute_0_438, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) /
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_437, refute_0_386])).

cnf(refute_0_439, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     (1/32 +
      skoXC1 * (14 + skoXC1 * (560 + skoXC1 * (3584 + skoXC1 * 2048)))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) *
     (1/5 +
      skoXC1 *
      (95/6 +
       skoXC1 *
       (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))) |
     1/5 +
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) <=
     0), inference(resolve, [], [refute_0_418, refute_0_438])).

cnf(refute_0_440, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     (1/160 +
      skoXC1 *
      (3163/960 +
       skoXC1 *
       (2681/8 +
        skoXC1 *
        (409447/40 +
         skoXC1 *
         (39844033/480 +
          skoXC1 *
          (60141823/320 +
           skoXC1 *
           (35483/15 +
            skoXC1 *
            (-180792 + skoXC1 * (-470784/5 + skoXC1 * -134144/15))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) |
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30)))) <=
     -1/5), inference(arithmetic, [], [refute_0_439])).

cnf(refute_0_441, plain,
    (-1/5 <
     skoXC1 *
     (95/6 +
      skoXC1 *
      (140/3 + skoXC1 * (-20 + skoXC1 * (-115/3 + skoXC1 * -131/30))))),
    inference(decision, [],
              [refute_0_376, refute_0_350, refute_0_346, refute_0_342,
               refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_442, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) <
     (1/160 +
      skoXC1 *
      (3163/960 +
       skoXC1 *
       (2681/8 +
        skoXC1 *
        (409447/40 +
         skoXC1 *
         (39844033/480 +
          skoXC1 *
          (60141823/320 +
           skoXC1 *
           (35483/15 +
            skoXC1 *
            (-180792 + skoXC1 * (-470784/5 + skoXC1 * -134144/15))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_440, refute_0_441])).

cnf(refute_0_443, plain,
    (skoXC1 *
     (5 +
      skoXC1 *
      (35 +
       skoXC1 * (20 + skoXC1 * (-40 + skoXC1 * (-19 + skoXC1 * -1))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/160 +
     skoXC1 *
     (3163/960 +
      skoXC1 *
      (2681/8 +
       skoXC1 *
       (409447/40 +
        skoXC1 *
        (39844033/480 +
         skoXC1 *
         (60141823/320 +
          skoXC1 *
          (35483/15 +
           skoXC1 *
           (-180792 + skoXC1 * (-470784/5 + skoXC1 * -134144/15)))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0),
    inference(resolve, [], [refute_0_417, refute_0_442])).

cnf(refute_0_444, plain,
    (-1/160 <
     skoXC1 *
     (-1637/960 +
      skoXC1 *
      (-2079/8 +
       skoXC1 *
       (-106553/40 +
        skoXC1 *
        (-1147967/480 +
         skoXC1 *
         (4237503/320 +
          skoXC1 *
          (-86182/15 +
           skoXC1 *
           (17208 + skoXC1 * (-72704/5 + skoXC1 * -72704/15)))))))) |
     skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= -1),
    inference(arithmetic, [], [refute_0_443])).

cnf(refute_0_445, plain,
    (-1/160 <
     skoXC1 *
     (-1637/960 +
      skoXC1 *
      (-2079/8 +
       skoXC1 *
       (-106553/40 +
        skoXC1 *
        (-1147967/480 +
         skoXC1 *
         (4237503/320 +
          skoXC1 *
          (-86182/15 +
           skoXC1 *
           (17208 + skoXC1 * (-72704/5 + skoXC1 * -72704/15)))))))) |
     -1 < skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))),
    inference(decision, [],
              [refute_0_376, refute_0_350, refute_0_346, refute_0_342,
               refute_0_340, refute_0_336, refute_0_334, refute_0_330,
               refute_0_328, refute_0_324, refute_0_166, refute_0_162,
               refute_0_135, refute_0_131, refute_0_26, refute_0_27,
               refute_0_28])).

cnf(refute_0_446, plain,
    (-1/160 <
     skoXC1 *
     (-1637/960 +
      skoXC1 *
      (-2079/8 +
       skoXC1 *
       (-106553/40 +
        skoXC1 *
        (-1147967/480 +
         skoXC1 *
         (4237503/320 +
          skoXC1 *
          (-86182/15 +
           skoXC1 *
           (17208 + skoXC1 * (-72704/5 + skoXC1 * -72704/15))))))))),
    inference(resolve, [], [refute_0_444, refute_0_445])).

cnf(refute_0_447, plain,
    (-1/6 <
     skoXC1 *
     (107/5 +
      skoXC1 *
      (125 +
       skoXC1 *
       (200/3 + skoXC1 * (-275/2 + skoXC1 * (-71 + skoXC1 * -71/15)))))),
    inference(decision, [],
              [refute_0_446, refute_0_376, refute_0_350, refute_0_346,
               refute_0_342, refute_0_340, refute_0_336, refute_0_334,
               refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_448, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) <
     (1/192 +
      skoXC1 *
      (1441/480 +
       skoXC1 *
       (190483/480 +
        skoXC1 *
        (172001/12 +
         skoXC1 *
         (142049251/960 +
          skoXC1 *
          (253071991/480 +
           skoXC1 *
           (200130809/480 +
            skoXC1 *
            (-5941394/15 +
             skoXC1 *
             (-1616144/3 +
              skoXC1 * (-2435584/15 + skoXC1 * -145408/15)))))))))) /
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)))),
    inference(resolve, [], [refute_0_416, refute_0_447])).

cnf(refute_0_449, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (69 +
       skoXC1 *
       (125 +
        skoXC1 *
        (-50 + skoXC1 * (-120 + skoXC1 * (-29 + skoXC1 * -1)))))) *
     (1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))) <
     1/192 +
     skoXC1 *
     (1441/480 +
      skoXC1 *
      (190483/480 +
       skoXC1 *
       (172001/12 +
        skoXC1 *
        (142049251/960 +
         skoXC1 *
         (253071991/480 +
          skoXC1 *
          (200130809/480 +
           skoXC1 *
           (-5941394/15 +
            skoXC1 *
            (-1616144/3 +
             skoXC1 * (-2435584/15 + skoXC1 * -145408/15))))))))) |
     1 + skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= 0),
    inference(resolve, [], [refute_0_393, refute_0_448])).

cnf(refute_0_450, plain,
    (-1/192 <
     skoXC1 *
     (-1439/480 +
      skoXC1 *
      (-165197/480 +
       skoXC1 *
       (-51259/12 +
        skoXC1 *
        (-13637789/960 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (3843929/480 +
           skoXC1 *
           (404941/15 +
            skoXC1 *
            (14656/3 + skoXC1 * (-626944/15 + skoXC1 * -83968/15))))))))) |
     skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= -1),
    inference(arithmetic, [], [refute_0_449])).

cnf(refute_0_451, plain,
    (-1/192 <
     skoXC1 *
     (-1439/480 +
      skoXC1 *
      (-165197/480 +
       skoXC1 *
       (-51259/12 +
        skoXC1 *
        (-13637789/960 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (3843929/480 +
           skoXC1 *
           (404941/15 +
            skoXC1 *
            (14656/3 + skoXC1 * (-626944/15 + skoXC1 * -83968/15))))))))) |
     -1 < skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))),
    inference(decision, [],
              [refute_0_446, refute_0_376, refute_0_350, refute_0_346,
               refute_0_342, refute_0_340, refute_0_336, refute_0_334,
               refute_0_330, refute_0_328, refute_0_324, refute_0_166,
               refute_0_162, refute_0_135, refute_0_131, refute_0_26,
               refute_0_27, refute_0_28])).

cnf(refute_0_452, plain,
    (-1/192 <
     skoXC1 *
     (-1439/480 +
      skoXC1 *
      (-165197/480 +
       skoXC1 *
       (-51259/12 +
        skoXC1 *
        (-13637789/960 +
         skoXC1 *
         (12638071/480 +
          skoXC1 *
          (3843929/480 +
           skoXC1 *
           (404941/15 +
            skoXC1 *
            (14656/3 + skoXC1 * (-626944/15 + skoXC1 * -83968/15)))))))))),
    inference(resolve, [], [refute_0_450, refute_0_451])).

cnf(refute_0_453, plain,
    (skoXC1 *
     (-1327/320 +
      skoXC1 *
      (-34613/80 +
       skoXC1 *
       (-2179753/320 +
        skoXC1 *
        (-41108141/1120 +
         skoXC1 *
         (644707/64 +
          skoXC1 *
          (7420609/80 +
           skoXC1 *
           (66863199/2240 +
            skoXC1 *
            (525492/5 +
             skoXC1 *
             (-520696/5 +
              skoXC1 * (-83456 + skoXC1 * -218112/35)))))))))) <= -1/224 |
     skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= -1),
    inference(decision, [],
              [refute_0_452, refute_0_446, refute_0_376, refute_0_350,
               refute_0_346, refute_0_342, refute_0_340, refute_0_336,
               refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_454, plain,
    (skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096)) <= -1),
    inference(resolve, [], [refute_0_453, refute_0_392])).

cnf(refute_0_455, plain,
    (-1 < skoXC1 * (112 + skoXC1 * (1792 + skoXC1 * 4096))),
    inference(decision, [],
              [refute_0_452, refute_0_446, refute_0_376, refute_0_350,
               refute_0_346, refute_0_342, refute_0_340, refute_0_336,
               refute_0_334, refute_0_330, refute_0_328, refute_0_324,
               refute_0_166, refute_0_162, refute_0_135, refute_0_131,
               refute_0_26, refute_0_27, refute_0_28])).

cnf(refute_0_456, plain, ($false),
    inference(resolve, [], [refute_0_454, refute_0_455])).
% SZS output end CNFRefutation for log-3615-sqrt.tptp
