% SZS output start CNFRefutation for log-upper-21-sqrt-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf11, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (5 * X ^ 5 + 647 * X ^ 4 + 4397 * X ^ 3 + 6397 * X ^ 2 +
             2272 * X + 142) * (X - 1) /
            (6 * X ^ 5 + 75 * X ^ 4 + 200 * X ^ 3 + 150 * X ^ 2 + 30 * X +
             1), Y) | lgen(R, ln(X), Y))).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8 * (X ^ 4 + 28 * X ^ 3 + 70 * X ^ 2 + 28 * X + 1) /
            ((X + 1) * (X ^ 2 + 6 * X + 1)), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_upper_bound_2_large, axiom,
    (X < 1 |
     ~ lgen(R,
            1/32 *
            (65536 + 114688 * X + 17920 * X ^ 2 + 448 * X ^ 3 + X ^ 4) /
            ((X + 16) * (256 + 96 * X + X ^ 2)), Y) |
     lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

fof(log_upper_bound_problem_21, conjecture,
    (! [X] :
       (1/100 <= X => ln(1 + 1 / X) <= 1/100 + 1 / sqrt(X ^ 2 + X)))).

fof(subgoal_0, plain,
    (! [X] : (1/100 <= X => ln(1 + 1 / X) <= 1/100 + 1 / sqrt(X ^ 2 + X))),
    inference(strip, [], [log_upper_bound_problem_21])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (1/100 <= X => ln(1 + 1 / X) <= 1/100 + 1 / sqrt(X ^ 2 + X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (1/100 + 1 / sqrt(X ^ 2 + X) < ln(1 + 1 / X) & 1/100 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (1/100 + 1 / sqrt(skoXC1 ^ 2 + skoXC1) < ln(1 + 1 / skoXC1) &
     1/100 <= skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/100 + 1 / sqrt(skoXC1 ^ 2 + skoXC1) < ln(1 + 1 / skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (1/100 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (-1 < 1 / skoXC1 | 1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/48 +
     skoXC1 *
     (4051/1200 +
      skoXC1 *
      (262591/2400 +
       skoXC1 *
       (59577/50 +
        skoXC1 *
        (3710149/600 +
         skoXC1 *
         (4436663/240 +
          skoXC1 *
          (16866799/480 +
           skoXC1 *
           (5349343/120 +
            skoXC1 *
            (4636717/120 +
             skoXC1 *
             (9287289/400 +
              skoXC1 *
              (306943/32 +
               skoXC1 *
               (246463/100 +
                skoXC1 *
                (66423/200 +
                 skoXC1 * (10773/200 + skoXC1 * -231/400))))))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     (1/48 +
      skoXC1 *
      (4051/1200 +
       skoXC1 *
       (262591/2400 +
        skoXC1 *
        (59577/50 +
         skoXC1 *
         (3710149/600 +
          skoXC1 *
          (4436663/240 +
           skoXC1 *
           (16866799/480 +
            skoXC1 *
            (5349343/120 +
             skoXC1 *
             (4636717/120 +
              skoXC1 *
              (9287289/400 +
               skoXC1 *
               (306943/32 +
                skoXC1 *
                (246463/100 +
                 skoXC1 *
                 (66423/200 +
                  skoXC1 * (10773/200 + skoXC1 * -231/400)))))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_2, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) |
     1/6 +
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))) <= 0 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_3, plain,
    (~ lgen(0, sqrt(X_000069), X_000070) | sqrt(X_000069) <= X_000070),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_4, plain,
    (X_000069 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000069 ^ 4 + 28 * X_000069 ^ 3 + 70 * X_000069 ^ 2 +
             28 * X_000069 + 1) /
            ((X_000069 + 1) * (X_000069 ^ 2 + 6 * X_000069 + 1)),
            X_000070) | lgen(0, sqrt(X_000069), X_000070)),
    inference(subst, [], [sqrt_upper_bound_2])).

cnf(refute_0_5, plain,
    (X_000069 < 0 |
     ~ lgen(0,
            1/8 *
            (X_000069 ^ 4 + 28 * X_000069 ^ 3 + 70 * X_000069 ^ 2 +
             28 * X_000069 + 1) /
            ((X_000069 + 1) * (X_000069 ^ 2 + 6 * X_000069 + 1)),
            X_000070) | sqrt(X_000069) <= X_000070),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (X_000069 < 0 |
     X_000070 <
     (1/8 +
      X_000069 *
      (7/2 + X_000069 * (35/4 + X_000069 * (7/2 + X_000069 * 1/8)))) /
     (1 + X_000069 * (7 + X_000069 * (7 + X_000069))) |
     sqrt(X_000069) <= X_000070), inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))))),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_8, plain,
    ((-1/6 +
      skoXC1 *
      (-112/5 +
       skoXC1 *
       (-469/2 + skoXC1 * (-784 + skoXC1 * (-1029 + skoXC1 * -462))))) *
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-105 +
       skoXC1 *
       (-560 + skoXC1 * (-1260 + skoXC1 * (-1260 + skoXC1 * -462))))) +
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (-3/50 +
       skoXC1 *
       (-21/20 +
        skoXC1 *
        (-28/5 +
         skoXC1 * (-63/5 + skoXC1 * (-63/5 + skoXC1 * -231/50)))))) |
     (skoXC1 *
      (-6 +
       skoXC1 *
       (-105 +
        skoXC1 *
        (-560 + skoXC1 * (-1260 + skoXC1 * (-1260 + skoXC1 * -462))))) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 *
       (-3/50 +
        skoXC1 *
        (-21/20 +
         skoXC1 *
         (-28/5 +
          skoXC1 * (-63/5 + skoXC1 * (-63/5 + skoXC1 * -231/50))))))) /
     sqrt(skoXC1 * (1 + skoXC1)) <=
     -1/6 +
     skoXC1 *
     (-112/5 +
      skoXC1 *
      (-469/2 + skoXC1 * (-784 + skoXC1 * (-1029 + skoXC1 * -462)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_9, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) <
     1/6 +
     skoXC1 *
     (112/5 +
      skoXC1 * (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462)))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <= 0 |
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_10, plain,
    (1/100 + 1 / sqrt(skoXC1 ^ 2 + skoXC1) < ln(1 + 1 / skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_11, plain,
    (1 / sqrt(skoXC1 * (1 + skoXC1)) < -1/100 + ln(1 + 1 / skoXC1)),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 < (-1/100 + ln(1 + 1 / skoXC1)) * sqrt(skoXC1 * (1 + skoXC1)) |
     -1/100 + ln(1 + 1 / skoXC1) <= 1 / sqrt(skoXC1 * (1 + skoXC1)) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_13, plain,
    (1 < (-1/100 + ln(1 + 1 / skoXC1)) * sqrt(skoXC1 * (1 + skoXC1)) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_11])).

cnf(refute_0_14, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     ln(1 + 1 / skoXC1) * sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (1 + skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (1/100 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain, (0 < skoXC1 * (1 + skoXC1)),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     ln(1 + 1 / skoXC1) * sqrt(skoXC1 * (1 + skoXC1))),
    inference(resolve, [], [refute_0_14, refute_0_16])).

cnf(refute_0_18, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < ln(1 + 1 / skoXC1) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(split, [], [refute_0_17])).

cnf(refute_0_19, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < ln(1 + 1 / skoXC1) |
     skoXC1 * (1 + skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < ln(1 + 1 / skoXC1)),
    inference(resolve, [], [refute_0_19, refute_0_16])).

cnf(refute_0_21, plain,
    (~ lgen(0, ln(X_000111), X_000112) | ln(X_000111) <= X_000112),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_22, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000111 ^ 5 + 647 * X_000111 ^ 4 + 4397 * X_000111 ^ 3 +
             6397 * X_000111 ^ 2 + 2272 * X_000111 + 142) *
            (X_000111 - 1) /
            (6 * X_000111 ^ 5 + 75 * X_000111 ^ 4 + 200 * X_000111 ^ 3 +
             150 * X_000111 ^ 2 + 30 * X_000111 + 1), X_000112) |
     X_000111 <= 0 | lgen(0, ln(X_000111), X_000112)),
    inference(subst, [], [ln_upper_bound_cf11])).

cnf(refute_0_23, plain,
    (~ lgen(0,
            1/30 *
            (5 * X_000111 ^ 5 + 647 * X_000111 ^ 4 + 4397 * X_000111 ^ 3 +
             6397 * X_000111 ^ 2 + 2272 * X_000111 + 142) *
            (X_000111 - 1) /
            (6 * X_000111 ^ 5 + 75 * X_000111 ^ 4 + 200 * X_000111 ^ 3 +
             150 * X_000111 ^ 2 + 30 * X_000111 + 1), X_000112) |
     X_000111 <= 0 | ln(X_000111) <= X_000112),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (X_000112 <
     (-71/15 +
      X_000111 *
      (-71 +
       X_000111 *
       (-275/2 +
        X_000111 *
        (200/3 +
         X_000111 * (125 + X_000111 * (107/5 + X_000111 * 1/6)))))) /
     (1 +
      X_000111 *
      (30 +
       X_000111 *
       (150 + X_000111 * (200 + X_000111 * (75 + X_000111 * 6))))) |
     X_000111 <= 0 | ln(X_000111) <= X_000112),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-71/15 +
      (1 + 1 / skoXC1) *
      (-71 +
       (1 + 1 / skoXC1) *
       (-275/2 +
        (1 + 1 / skoXC1) *
        (200/3 +
         (1 + 1 / skoXC1) *
         (125 + (1 + 1 / skoXC1) * (107/5 + (1 + 1 / skoXC1) * 1/6)))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (30 +
       (1 + 1 / skoXC1) *
       (150 +
        (1 + 1 / skoXC1) *
        (200 + (1 + 1 / skoXC1) * (75 + (1 + 1 / skoXC1) * 6))))) |
     1 + 1 / skoXC1 <= 0 |
     ln(1 + 1 / skoXC1) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))), inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-71/15 +
      (1 + 1 / skoXC1) *
      (-71 +
       (1 + 1 / skoXC1) *
       (-275/2 +
        (1 + 1 / skoXC1) *
        (200/3 +
         (1 + 1 / skoXC1) *
         (125 + (1 + 1 / skoXC1) * (107/5 + (1 + 1 / skoXC1) * 1/6)))))) /
     (1 +
      (1 + 1 / skoXC1) *
      (30 +
       (1 + 1 / skoXC1) *
       (150 +
        (1 + 1 / skoXC1) *
        (200 + (1 + 1 / skoXC1) * (75 + (1 + 1 / skoXC1) * 6))))) |
     1 + 1 / skoXC1 <= 0),
    inference(resolve, [], [refute_0_25, refute_0_20])).

cnf(refute_0_27, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 * (2 + skoXC1 * 3) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     1024 +
     skoXC1 *
     (197376/25 +
      skoXC1 *
      (317272/25 +
       skoXC1 *
       (64043/5 +
        skoXC1 *
        (227585/64 +
         skoXC1 *
         (2849753/1600 +
          skoXC1 *
          (366609/3200 +
           skoXC1 *
           (16431/400 +
            skoXC1 *
            (-3/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200))))))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 |
     (1024 +
      skoXC1 *
      (197376/25 +
       skoXC1 *
       (317272/25 +
        skoXC1 *
        (64043/5 +
         skoXC1 *
         (227585/64 +
          skoXC1 *
          (2849753/1600 +
           skoXC1 *
           (366609/3200 +
            skoXC1 *
            (16431/400 +
             skoXC1 *
             (-3/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * 3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_29, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_30, plain,
    ((-1/2 + skoXC1 * -3) * sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-2 + skoXC1 * -3) +
     sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (-1/50 + skoXC1 * -3/100)) |
     (skoXC1 * (-2 + skoXC1 * -3) +
      sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (-1/50 + skoXC1 * -3/100))) /
     sqrt(skoXC1 * (1 + skoXC1)) <= -1/2 + skoXC1 * -3 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_31, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (2 + skoXC1 * 3)) <
     1/2 + skoXC1 * 3 | skoXC1 * (2 + skoXC1 * 3) <= 0 |
     (1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3)) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_32, plain,
    (~ lgen(0, ln(X_000055), X_000056) | ln(X_000055) <= X_000056),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_33, plain,
    (~ lgen(0, 1/2 * (X_000055 + 5) * (X_000055 - 1) / (2 * X_000055 + 1),
            X_000056) | X_000055 <= 0 | lgen(0, ln(X_000055), X_000056)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_34, plain,
    (~ lgen(0, 1/2 * (X_000055 + 5) * (X_000055 - 1) / (2 * X_000055 + 1),
            X_000056) | X_000055 <= 0 | ln(X_000055) <= X_000056),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (X_000056 <
     (-5/2 + X_000055 * (2 + X_000055 * 1/2)) / (1 + X_000055 * 2) |
     X_000055 <= 0 | ln(X_000055) <= X_000056),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-5/2 + (1 + 1 / skoXC1) * (2 + (1 + 1 / skoXC1) * 1/2)) /
     (1 + (1 + 1 / skoXC1) * 2) | 1 + 1 / skoXC1 <= 0 |
     ln(1 + 1 / skoXC1) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))), inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-5/2 + (1 + 1 / skoXC1) * (2 + (1 + 1 / skoXC1) * 1/2)) /
     (1 + (1 + 1 / skoXC1) * 2) | 1 + 1 / skoXC1 <= 0),
    inference(resolve, [], [refute_0_36, refute_0_20])).

cnf(refute_0_38, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3)) | 1 / skoXC1 <= -1 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 * (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 * (2 + skoXC1 * (99/50 + skoXC1 * -1/50)) |
     1 + skoXC1 * (1 + skoXC1) <= 0 |
     skoXC1 * (2 + skoXC1 * (99/50 + skoXC1 * -1/50)) /
     (1 + skoXC1 * (1 + skoXC1)) <= skoXC1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_40, plain,
    (skoXC1 <
     skoXC1 * (2 + skoXC1 * 2) / (1 + skoXC1 * (1 + skoXC1)) *
     (1 + skoXC1 * -1/100) | 0 <= 1 + skoXC1 * -1/100 |
     skoXC1 / (1 + skoXC1 * -1/100) <=
     skoXC1 * (2 + skoXC1 * 2) / (1 + skoXC1 * (1 + skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_41, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) * skoXC1 <
     sqrt(skoXC1 * (1 + skoXC1)) |
     sqrt(skoXC1 * (1 + skoXC1)) / skoXC1 <=
     1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 | skoXC1 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_42, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     1 / skoXC1 * sqrt(skoXC1 * (1 + skoXC1)) |
     1 / skoXC1 <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_43, plain,
    (~ lgen(0, ln(X_000034), X_000035) | ln(X_000034) <= X_000035),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_44, plain,
    (~ lgen(0, X_000034 - 1, X_000035) | X_000034 <= 0 |
     lgen(0, ln(X_000034), X_000035)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_45, plain,
    (~ lgen(0, X_000034 - 1, X_000035) | X_000034 <= 0 |
     ln(X_000034) <= X_000035),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (X_000035 < -1 + X_000034 | X_000034 <= 0 | ln(X_000034) <= X_000035),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < -1 + (1 + 1 / skoXC1) |
     1 + 1 / skoXC1 <= 0 |
     ln(1 + 1 / skoXC1) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))), inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < -1 + (1 + 1 / skoXC1) |
     1 + 1 / skoXC1 <= 0),
    inference(resolve, [], [refute_0_47, refute_0_20])).

cnf(refute_0_49, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) < 1 / skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     1 / skoXC1 * sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_42, refute_0_49])).

cnf(refute_0_51, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     sqrt(skoXC1 * (1 + skoXC1)) / skoXC1 | skoXC1 * (1 + skoXC1) <= 0 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100 <
     sqrt(skoXC1 * (1 + skoXC1)) / skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_51, refute_0_16])).

cnf(refute_0_53, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) * skoXC1 <
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_41, refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < sqrt(skoXC1 * (1 + skoXC1)) * (1 + skoXC1 * -1/100) |
     1 / skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain, (0 < skoXC1),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_56, plain,
    (skoXC1 < sqrt(skoXC1 * (1 + skoXC1)) * (1 + skoXC1 * -1/100) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) < skoXC1 / (1 + skoXC1 * -1/100) |
     0 <= 1 + skoXC1 * -1/100 | 1 / skoXC1 <= -1),
    inference(split, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (sqrt(skoXC1 * (1 + skoXC1)) < skoXC1 / (1 + skoXC1 * -1/100) |
     1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (~ lgen(0, X_000049, sqrt(X_000048)) | X_000049 <= sqrt(X_000048)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_60, plain,
    (X_000048 < 0 | ~ lgen(0, X_000049, 2 * X_000048 / (X_000048 + 1)) |
     lgen(0, X_000049, sqrt(X_000048))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_61, plain,
    (X_000048 < 0 | ~ lgen(0, X_000049, 2 * X_000048 / (X_000048 + 1)) |
     X_000049 <= sqrt(X_000048)),
    inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (X_000048 * 2 / (1 + X_000048) < X_000049 | X_000048 < 0 |
     X_000049 <= sqrt(X_000048)),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (1 + skoXC1) * 2 / (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 / (1 + skoXC1 * -1/100) |
     skoXC1 / (1 + skoXC1 * -1/100) <= sqrt(skoXC1 * (1 + skoXC1))),
    inference(subst, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (1 + skoXC1) * 2 / (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 / (1 + skoXC1 * -1/100) | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_63, refute_0_58])).

cnf(refute_0_65, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 2) / (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 / (1 + skoXC1 * -1/100) | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) * (1 + skoXC1 * -1/100) |
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) <=
     skoXC1 / (1 + skoXC1 * -1/100) | 1 + skoXC1 * -1/100 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_67, plain,
    (skoXC1 / (1 + skoXC1 * -1/100) < sqrt(skoXC1 * (1 + skoXC1)) |
     1 + skoXC1 * -1/100 <= 0 | 1 / skoXC1 <= -1),
    inference(split, [], [refute_0_56])).

cnf(refute_0_68, plain,
    (skoXC1 / (1 + skoXC1 * -1/100) < sqrt(skoXC1 * (1 + skoXC1)) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (~ lgen(0, sqrt(X_000041), X_000042) | sqrt(X_000041) <= X_000042),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_70, plain,
    (X_000041 < 0 | ~ lgen(0, (X_000041 + 1) / 2, X_000042) |
     lgen(0, sqrt(X_000041), X_000042)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_71, plain,
    (X_000041 < 0 | ~ lgen(0, (X_000041 + 1) / 2, X_000042) |
     sqrt(X_000041) <= X_000042),
    inference(resolve, [], [refute_0_70, refute_0_69])).

cnf(refute_0_72, plain,
    (X_000041 < 0 | X_000042 < 1/2 + X_000041 * 1/2 |
     sqrt(X_000041) <= X_000042),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 / (1 + skoXC1 * -1/100) < 1/2 + skoXC1 * (1 + skoXC1) * 1/2 |
     sqrt(skoXC1 * (1 + skoXC1)) <= skoXC1 / (1 + skoXC1 * -1/100)),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 / (1 + skoXC1 * -1/100) < 1/2 + skoXC1 * (1 + skoXC1) * 1/2 |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_73, refute_0_68])).

cnf(refute_0_75, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 / (1 + skoXC1 * -1/100) < 1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (1 + skoXC1 * -1/100 < 0 | skoXC1 < 0 | 0 < 1 + skoXC1 * -1/100 |
     1 / skoXC1 <= -1), inference(split, [], [refute_0_56])).

cnf(refute_0_77, plain,
    (skoXC1 < 0 | skoXC1 < 100 | 100 < skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain, (skoXC1 < 100 | 100 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_79, plain, (skoXC1 < 100 | 100 < skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_78, refute_0_77])).

cnf(refute_0_80, plain,
    (skoXC1 < 100 | 100 < skoXC1 | 1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_79, refute_0_0])).

cnf(refute_0_81, plain,
    (skoXC1 < 100 | 100 < skoXC1 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (skoXC1 < 100 | -1 < skoXC1 | 100 < skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_83, plain, (skoXC1 < 100 | 100 < skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain, (skoXC1 < 100 | 0 < skoXC1 | 100 < skoXC1),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_85, plain, (skoXC1 < 100 | 100 < skoXC1),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain, (100 <= skoXC1 | skoXC1 * (-1 + skoXC1 * -1) <= 0),
    inference(decision, [], [refute_0_85, refute_0_15])).

cnf(refute_0_87, plain,
    (skoXC1 / (1 + skoXC1 * -1/100) < 1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_86, refute_0_75])).

cnf(refute_0_88, plain,
    (skoXC1 <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) * (1 + skoXC1 * -1/100) |
     100 <= skoXC1 | 1 + skoXC1 * -1/100 <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_66, refute_0_87])).

cnf(refute_0_89, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     1 <= -1 * skoXC1 | 100 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_89, refute_0_0])).

cnf(refute_0_91, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     100 <= skoXC1 | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     -1 < skoXC1 | 100 <= skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_85, refute_0_15])).

cnf(refute_0_93, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     100 <= skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     0 < skoXC1 | 100 <= skoXC1),
    inference(decision, [], [refute_0_85, refute_0_15])).

cnf(refute_0_95, plain,
    (-1/2 < skoXC1 * (-101/200 + skoXC1 * (99/200 + skoXC1 * -1/200)) |
     100 <= skoXC1), inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain, (skoXC1 * (-1 + skoXC1 * -1) <= 0 | skoXC1 <= 100),
    inference(decision, [], [refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_97, plain,
    (skoXC1 * (2 + skoXC1 * 2) / (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 / (1 + skoXC1 * -1/100) | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_96, refute_0_65])).

cnf(refute_0_98, plain,
    (skoXC1 <
     skoXC1 * (2 + skoXC1 * 2) / (1 + skoXC1 * (1 + skoXC1)) *
     (1 + skoXC1 * -1/100) | 0 <= 1 + skoXC1 * -1/100 | 1 / skoXC1 <= -1 |
     skoXC1 <= 100), inference(resolve, [], [refute_0_40, refute_0_97])).

cnf(refute_0_99, plain,
    (skoXC1 <
     skoXC1 * (2 + skoXC1 * (99/50 + skoXC1 * -1/50)) /
     (1 + skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 * (1 + skoXC1 * (1 + skoXC1)) <
     skoXC1 * (2 + skoXC1 * (99/50 + skoXC1 * -1/50)) |
     1 + skoXC1 * (1 + skoXC1) <= 0 | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_39, refute_0_99])).

cnf(refute_0_101, plain,
    (0 < skoXC1 * (1 + skoXC1 * (49/50 + skoXC1 * -51/50)) |
     skoXC1 * (1 + skoXC1) <= -1 | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (skoXC1 * (1 + skoXC1 * (49/50 + skoXC1 * -51/50)) <= 0 |
     skoXC1 * (1 + skoXC1) <= -1 | skoXC1 <= 100),
    inference(decision, [], [refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_103, plain,
    (skoXC1 * (1 + skoXC1) <= -1 | 1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_102, refute_0_101])).

cnf(refute_0_104, plain, (-1 < skoXC1 * (1 + skoXC1) | skoXC1 <= 100),
    inference(decision, [], [refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_105, plain, (1 / skoXC1 <= -1 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain, (1 <= -1 * skoXC1 | skoXC1 <= 0 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_105, refute_0_0])).

cnf(refute_0_107, plain, (skoXC1 <= -1 | skoXC1 <= 0 | skoXC1 <= 100),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain, (-1 < skoXC1 | skoXC1 <= 0 | skoXC1 <= 100),
    inference(decision, [], [refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_109, plain, (skoXC1 <= 0 | skoXC1 <= 100),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain, (0 < skoXC1 | skoXC1 <= 100),
    inference(decision, [], [refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_111, plain, (skoXC1 <= 100),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain, (skoXC1 != 0),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_113, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/2 + skoXC1 * 3) / (skoXC1 * (2 + skoXC1 * 3)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_38, refute_0_112])).

cnf(refute_0_114, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (2 + skoXC1 * 3)) <
     1/2 + skoXC1 * 3 | skoXC1 * (2 + skoXC1 * 3) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_31, refute_0_113])).

cnf(refute_0_115, plain,
    (-1/2 + skoXC1 * -3 <
     (skoXC1 * (-2 + skoXC1 * -3) +
      sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (-1/50 + skoXC1 * -3/100))) /
     sqrt(skoXC1 * (1 + skoXC1)) | skoXC1 * (2 + skoXC1 * 3) <= 0 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain, (0 < skoXC1 * (2 + skoXC1 * 3)),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_117, plain,
    (-1/2 + skoXC1 * -3 <
     (skoXC1 * (-2 + skoXC1 * -3) +
      sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (-1/50 + skoXC1 * -3/100))) /
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    ((-1/2 + skoXC1 * -3) * sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-2 + skoXC1 * -3) +
     sqrt(skoXC1 * (1 + skoXC1)) * (skoXC1 * (-1/50 + skoXC1 * -3/100)) |
     1 / skoXC1 <= -1 | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_30, refute_0_117])).

cnf(refute_0_119, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     skoXC1 * (1 + skoXC1) <= 0 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_119, refute_0_16])).

cnf(refute_0_121, plain,
    (skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 | 1 / skoXC1 <= -1),
    inference(split, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (149/50 + skoXC1 * -3/100) <= -1/2 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain, (-1/2 < skoXC1 * (149/50 + skoXC1 * -3/100)),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_124, plain,
    (skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (~ lgen(0, sqrt(X_000076), X_000077) | sqrt(X_000076) <= X_000077),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_126, plain,
    (X_000076 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000076 + 17920 * X_000076 ^ 2 +
             448 * X_000076 ^ 3 + X_000076 ^ 4) /
            ((X_000076 + 16) * (256 + 96 * X_000076 + X_000076 ^ 2)),
            X_000077) | lgen(0, sqrt(X_000076), X_000077)),
    inference(subst, [], [sqrt_upper_bound_2_large])).

cnf(refute_0_127, plain,
    (X_000076 < 1 |
     ~ lgen(0,
            1/32 *
            (65536 + 114688 * X_000076 + 17920 * X_000076 ^ 2 +
             448 * X_000076 ^ 3 + X_000076 ^ 4) /
            ((X_000076 + 16) * (256 + 96 * X_000076 + X_000076 ^ 2)),
            X_000077) | sqrt(X_000076) <= X_000077),
    inference(resolve, [], [refute_0_126, refute_0_125])).

cnf(refute_0_128, plain,
    (X_000076 < 1 |
     X_000077 <
     (2048 +
      X_000076 *
      (3584 + X_000076 * (560 + X_000076 * (14 + X_000076 * 1/32)))) /
     (4096 + X_000076 * (1792 + X_000076 * (112 + X_000076))) |
     sqrt(X_000076) <= X_000077),
    inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100))),
    inference(subst, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_129, refute_0_124])).

cnf(refute_0_131, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_29, refute_0_131])).

cnf(refute_0_133, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) <
     (1024 +
      skoXC1 *
      (197376/25 +
       skoXC1 *
       (317272/25 +
        skoXC1 *
        (64043/5 +
         skoXC1 *
         (227585/64 +
          skoXC1 *
          (2849753/1600 +
           skoXC1 *
           (366609/3200 +
            skoXC1 *
            (16431/400 +
             skoXC1 *
             (-3/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     skoXC1 * (149/50 + skoXC1 * -3/100) <= -1/2 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/24 +
     skoXC1 *
     (4891/2400 +
      skoXC1 *
      (4457/150 +
       skoXC1 *
       (2539/20 +
        skoXC1 *
        (32821/120 +
         skoXC1 *
         (827917/2400 +
          skoXC1 *
          (160447/600 +
           skoXC1 *
           (81919/600 +
            skoXC1 *
            (533/12 +
             skoXC1 *
             (4309/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     (1/24 +
      skoXC1 *
      (4891/2400 +
       skoXC1 *
       (4457/150 +
        skoXC1 *
        (2539/20 +
         skoXC1 *
         (32821/120 +
          skoXC1 *
          (827917/2400 +
           skoXC1 *
           (160447/600 +
            skoXC1 *
            (81919/600 +
             skoXC1 *
             (533/12 +
              skoXC1 *
              (4309/800 + skoXC1 * (237/200 + skoXC1 * -1/80))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_135, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= 0 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_136, plain,
    ((-1/3 + skoXC1 * (-7 + skoXC1 * -10)) * sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) +
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (-3/100 + skoXC1 * (-3/25 + skoXC1 * -1/10))) |
     (skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 * (-3/100 + skoXC1 * (-3/25 + skoXC1 * -1/10)))) /
     sqrt(skoXC1 * (1 + skoXC1)) <= -1/3 + skoXC1 * (-7 + skoXC1 * -10) |
     sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_137, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <
     1/3 + skoXC1 * (7 + skoXC1 * 10) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0 |
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_138, plain,
    (~ lgen(0, ln(X_000062), X_000063) | ln(X_000062) <= X_000063),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_139, plain,
    (~ lgen(0,
            1/3 * (X_000062 ^ 2 + 19 * X_000062 + 10) * (X_000062 - 1) /
            (3 * X_000062 ^ 2 + 6 * X_000062 + 1), X_000063) |
     X_000062 <= 0 | lgen(0, ln(X_000062), X_000063)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_140, plain,
    (~ lgen(0,
            1/3 * (X_000062 ^ 2 + 19 * X_000062 + 10) * (X_000062 - 1) /
            (3 * X_000062 ^ 2 + 6 * X_000062 + 1), X_000063) |
     X_000062 <= 0 | ln(X_000062) <= X_000063),
    inference(resolve, [], [refute_0_139, refute_0_138])).

cnf(refute_0_141, plain,
    (X_000063 <
     (-10/3 + X_000062 * (-3 + X_000062 * (6 + X_000062 * 1/3))) /
     (1 + X_000062 * (6 + X_000062 * 3)) | X_000062 <= 0 |
     ln(X_000062) <= X_000063), inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-10/3 +
      (1 + 1 / skoXC1) *
      (-3 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 1/3))) /
     (1 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 3)) |
     1 + 1 / skoXC1 <= 0 |
     ln(1 + 1 / skoXC1) <=
     (1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1))), inference(subst, [], [refute_0_141])).

cnf(refute_0_143, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (-10/3 +
      (1 + 1 / skoXC1) *
      (-3 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 1/3))) /
     (1 + (1 + 1 / skoXC1) * (6 + (1 + 1 / skoXC1) * 3)) |
     1 + 1 / skoXC1 <= 0),
    inference(resolve, [], [refute_0_142, refute_0_20])).

cnf(refute_0_144, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) | 1 / skoXC1 <= -1 |
     skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_146, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     1/2 + skoXC1 * (1 + skoXC1) * 1/2 |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100))),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_147, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     1/2 + skoXC1 * (1 + skoXC1) * 1/2 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_146, refute_0_124])).

cnf(refute_0_148, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain, (skoXC1 * (-1 + skoXC1 * -1) <= 0),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_150, plain,
    (skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_148])).

cnf(refute_0_151, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_145, refute_0_150])).

cnf(refute_0_152, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     skoXC1 * (149/50 + skoXC1 * -3/100) <= -1/2 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     -1/2 < skoXC1 * (149/50 + skoXC1 * -3/100)),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_154, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_154, refute_0_0])).

cnf(refute_0_156, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_158, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_160, plain,
    (-1/4 <
     skoXC1 *
     (-13/50 + skoXC1 * (-51/40 + skoXC1 * (59/40 + skoXC1 * -3/200)))),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain, (skoXC1 * skoXC1 != 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_162, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) | 1 / skoXC1 <= -1 |
     skoXC1 = 0), inference(resolve, [], [refute_0_144, refute_0_161])).

cnf(refute_0_163, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/3 + skoXC1 * (7 + skoXC1 * 10)) /
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_162, refute_0_112])).

cnf(refute_0_164, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))) <
     1/3 + skoXC1 * (7 + skoXC1 * 10) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_137, refute_0_163])).

cnf(refute_0_165, plain,
    (-1/3 + skoXC1 * (-7 + skoXC1 * -10) <
     (skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 * (-3/100 + skoXC1 * (-3/25 + skoXC1 * -1/10)))) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <= 0 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_164])).

cnf(refute_0_166, plain, (0 < skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10))),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_167, plain,
    (-1/3 + skoXC1 * (-7 + skoXC1 * -10) <
     (skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 * (-3/100 + skoXC1 * (-3/25 + skoXC1 * -1/10)))) /
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    ((-1/3 + skoXC1 * (-7 + skoXC1 * -10)) * sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 * (-3 + skoXC1 * (-12 + skoXC1 * -10)) +
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 * (-3/100 + skoXC1 * (-3/25 + skoXC1 * -1/10))) |
     1 / skoXC1 <= -1 | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_136, refute_0_167])).

cnf(refute_0_169, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     skoXC1 * (1 + skoXC1) <= 0 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_169, refute_0_16])).

cnf(refute_0_171, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= 0 |
     1 / skoXC1 <= -1), inference(split, [], [refute_0_170])).

cnf(refute_0_172, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= -1/3 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_171])).

cnf(refute_0_173, plain,
    (-1/3 < skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_174, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)))),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_176, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_175, refute_0_174])).

cnf(refute_0_177, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_177])).

cnf(refute_0_179, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_135, refute_0_178])).

cnf(refute_0_180, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/24 +
      skoXC1 *
      (4891/2400 +
       skoXC1 *
       (4457/150 +
        skoXC1 *
        (2539/20 +
         skoXC1 *
         (32821/120 +
          skoXC1 *
          (827917/2400 +
           skoXC1 *
           (160447/600 +
            skoXC1 *
            (81919/600 +
             skoXC1 *
             (533/12 +
              skoXC1 *
              (4309/800 + skoXC1 * (237/200 + skoXC1 * -1/80))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= -1/3 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/24 +
      skoXC1 *
      (4891/2400 +
       skoXC1 *
       (4457/150 +
        skoXC1 *
        (2539/20 +
         skoXC1 *
         (32821/120 +
          skoXC1 *
          (827917/2400 +
           skoXC1 *
           (160447/600 +
            skoXC1 *
            (81919/600 +
             skoXC1 *
             (533/12 +
              skoXC1 *
              (4309/800 + skoXC1 * (237/200 + skoXC1 * -1/80))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_180, refute_0_173])).

cnf(refute_0_182, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/24 +
     skoXC1 *
     (4891/2400 +
      skoXC1 *
      (4457/150 +
       skoXC1 *
       (2539/20 +
        skoXC1 *
        (32821/120 +
         skoXC1 *
         (827917/2400 +
          skoXC1 *
          (160447/600 +
           skoXC1 *
           (81919/600 +
            skoXC1 *
            (533/12 +
             skoXC1 *
             (4309/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_134, refute_0_181])).

cnf(refute_0_183, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_182])).

cnf(refute_0_184, plain,
    (skoXC1 * (2 + skoXC1 * 3) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/16 +
     skoXC1 *
     (849/400 +
      skoXC1 *
      (13241/800 +
       skoXC1 *
       (469/10 +
        skoXC1 *
        (719/10 +
         skoXC1 *
         (25043/400 +
          skoXC1 *
          (27459/800 +
           skoXC1 *
           (2517/200 +
            skoXC1 *
            (57/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     (1/16 +
      skoXC1 *
      (849/400 +
       skoXC1 *
       (13241/800 +
        skoXC1 *
        (469/10 +
         skoXC1 *
         (719/10 +
          skoXC1 *
          (25043/400 +
           skoXC1 *
           (27459/800 +
            skoXC1 *
            (2517/200 +
             skoXC1 *
             (57/40 + skoXC1 * (143/400 + skoXC1 * -3/800)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * 3)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_185, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_186, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100))),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_187, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_186, refute_0_124])).

cnf(refute_0_188, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (skoXC1 * (2 + skoXC1 * 3) /
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_188])).

cnf(refute_0_190, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/2 + skoXC1 * (149/50 + skoXC1 * -3/100)) |
     1/2 + skoXC1 * (149/50 + skoXC1 * -3/100) <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_185, refute_0_189])).

cnf(refute_0_191, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/16 +
      skoXC1 *
      (849/400 +
       skoXC1 *
       (13241/800 +
        skoXC1 *
        (469/10 +
         skoXC1 *
         (719/10 +
          skoXC1 *
          (25043/400 +
           skoXC1 *
           (27459/800 +
            skoXC1 *
            (2517/200 +
             skoXC1 *
             (57/40 + skoXC1 * (143/400 + skoXC1 * -3/800)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     skoXC1 * (149/50 + skoXC1 * -3/100) <= -1/2 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_190])).

cnf(refute_0_192, plain,
    (skoXC1 * (2 + skoXC1 * 3) <
     (1/16 +
      skoXC1 *
      (849/400 +
       skoXC1 *
       (13241/800 +
        skoXC1 *
        (469/10 +
         skoXC1 *
         (719/10 +
          skoXC1 *
          (25043/400 +
           skoXC1 *
           (27459/800 +
            skoXC1 *
            (2517/200 +
             skoXC1 *
             (57/40 + skoXC1 * (143/400 + skoXC1 * -3/800)))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_191, refute_0_123])).

cnf(refute_0_193, plain,
    (skoXC1 * (2 + skoXC1 * 3) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/16 +
     skoXC1 *
     (849/400 +
      skoXC1 *
      (13241/800 +
       skoXC1 *
       (469/10 +
        skoXC1 *
        (719/10 +
         skoXC1 *
         (25043/400 +
          skoXC1 *
          (27459/800 +
           skoXC1 *
           (2517/200 +
            skoXC1 *
            (57/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_184, refute_0_192])).

cnf(refute_0_194, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/8 +
     skoXC1 *
     (2799/800 +
      skoXC1 *
      (2443/200 +
       skoXC1 *
       (8351/400 +
        skoXC1 *
        (3833/200 +
         skoXC1 *
         (1729/160 +
          skoXC1 *
          (207/50 +
           skoXC1 * (183/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     (1/8 +
      skoXC1 *
      (2799/800 +
       skoXC1 *
       (2443/200 +
        skoXC1 *
        (8351/400 +
         skoXC1 *
         (3833/200 +
          skoXC1 *
          (1729/160 +
           skoXC1 *
           (207/50 +
            skoXC1 * (183/400 + skoXC1 * (3/25 + skoXC1 * -1/800))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_196, plain,
    (skoXC1 <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1 + skoXC1 * -1/100) | 1 + skoXC1 * -1/100 <= 0 |
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 / (1 + skoXC1 * -1/100)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_197, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <= skoXC1 / (1 + skoXC1 * -1/100)),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_198, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_197, refute_0_68])).

cnf(refute_0_199, plain, (skoXC1 < 100),
    inference(resolve, [], [refute_0_111, refute_0_85])).

cnf(refute_0_200, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (skoXC1 / (1 + skoXC1 * -1/100) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_201])).

cnf(refute_0_203, plain,
    (skoXC1 <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1 + skoXC1 * -1/100) | 1 + skoXC1 * -1/100 <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_196, refute_0_202])).

cnf(refute_0_204, plain,
    (skoXC1 <
     (1/8 +
      skoXC1 *
      (2799/800 +
       skoXC1 *
       (2443/200 +
        skoXC1 *
        (8351/400 +
         skoXC1 *
         (3833/200 +
          skoXC1 *
          (1729/160 +
           skoXC1 *
           (207/50 +
            skoXC1 * (183/400 + skoXC1 * (3/25 + skoXC1 * -1/800))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_203])).

cnf(refute_0_205, plain,
    (skoXC1 <
     (1/8 +
      skoXC1 *
      (2799/800 +
       skoXC1 *
       (2443/200 +
        skoXC1 *
        (8351/400 +
         skoXC1 *
         (3833/200 +
          skoXC1 *
          (1729/160 +
           skoXC1 *
           (207/50 +
            skoXC1 * (183/400 + skoXC1 * (3/25 + skoXC1 * -1/800))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_204, refute_0_199])).

cnf(refute_0_206, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/8 +
     skoXC1 *
     (2799/800 +
      skoXC1 *
      (2443/200 +
       skoXC1 *
       (8351/400 +
        skoXC1 *
        (3833/200 +
         skoXC1 *
         (1729/160 +
          skoXC1 *
          (207/50 +
           skoXC1 * (183/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_195, refute_0_205])).

cnf(refute_0_207, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_206])).

cnf(refute_0_208, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= 0 |
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) <=
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_209, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     1/2 + skoXC1 * (1 + skoXC1) * 1/2 |
     sqrt(skoXC1 * (1 + skoXC1)) <=
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)))),
    inference(subst, [], [refute_0_72])).

cnf(refute_0_210, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     1/2 + skoXC1 * (1 + skoXC1) * 1/2 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_209, refute_0_174])).

cnf(refute_0_211, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) /
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) <
     1/2 + skoXC1 * (1/2 + skoXC1 * 1/2) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_211])).

cnf(refute_0_213, plain,
    (skoXC1 * (3 + skoXC1 * (12 + skoXC1 * 10)) <
     (1/2 + skoXC1 * (1/2 + skoXC1 * 1/2)) *
     (1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))) |
     1/3 + skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_208, refute_0_212])).

cnf(refute_0_214, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10)) <= -1/3 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     -1/3 < skoXC1 * (697/100 + skoXC1 * (247/25 + skoXC1 * -1/10))),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_216, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_216, refute_0_0])).

cnf(refute_0_218, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_220, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20)))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_222, plain,
    (-1/6 <
     skoXC1 *
     (391/600 +
      skoXC1 *
      (-409/120 +
       skoXC1 * (-13/8 + skoXC1 * (489/100 + skoXC1 * -1/20))))),
    inference(resolve, [], [refute_0_220, refute_0_221])).

cnf(refute_0_223, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     -1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_224, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_207, refute_0_223])).

cnf(refute_0_225, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_224, refute_0_0])).

cnf(refute_0_226, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_225])).

cnf(refute_0_227, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_228, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800)))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_230, plain,
    (-1/8 <
     skoXC1 *
     (1999/800 +
      skoXC1 *
      (1043/200 +
       skoXC1 *
       (2751/400 +
        skoXC1 *
        (833/200 +
         skoXC1 *
         (129/160 +
          skoXC1 *
          (57/50 +
           skoXC1 * (-217/400 + skoXC1 * (3/25 + skoXC1 * -1/800))))))))),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     -1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_232, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_194, refute_0_231])).

cnf(refute_0_233, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_232, refute_0_0])).

cnf(refute_0_234, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_236, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800))))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_238, plain,
    (-1/16 <
     skoXC1 *
     (49/400 +
      skoXC1 *
      (-359/800 +
       skoXC1 *
       (-21/10 +
        skoXC1 *
        (-1/10 +
         skoXC1 *
         (-957/400 +
          skoXC1 *
          (-1341/800 +
           skoXC1 *
           (317/200 +
            skoXC1 *
            (-63/40 + skoXC1 * (143/400 + skoXC1 * -3/800)))))))))),
    inference(resolve, [], [refute_0_236, refute_0_237])).

cnf(refute_0_239, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     -1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_238, refute_0_230, refute_0_222, refute_0_160,
               refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_240, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_183, refute_0_239])).

cnf(refute_0_241, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_240, refute_0_0])).

cnf(refute_0_242, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_238, refute_0_230, refute_0_222, refute_0_160,
               refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_244, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_242, refute_0_243])).

cnf(refute_0_245, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80)))))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_238, refute_0_230, refute_0_222, refute_0_160,
               refute_0_111, refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_246, plain,
    (-1/24 <
     skoXC1 *
     (-2309/2400 +
      skoXC1 *
      (-493/150 +
       skoXC1 *
       (-181/20 +
        skoXC1 *
        (-1139/120 +
         skoXC1 *
         (-12083/2400 +
          skoXC1 *
          (-6953/600 +
           skoXC1 *
           (-1481/600 +
            skoXC1 *
            (29/12 +
             skoXC1 *
             (-3691/800 + skoXC1 * (237/200 + skoXC1 * -1/80))))))))))),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -1/2 < skoXC1 * (149/50 + skoXC1 * -3/100)),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_248, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) <
     (1024 +
      skoXC1 *
      (197376/25 +
       skoXC1 *
       (317272/25 +
        skoXC1 *
        (64043/5 +
         skoXC1 *
         (227585/64 +
          skoXC1 *
          (2849753/1600 +
           skoXC1 *
           (366609/3200 +
            skoXC1 *
            (16431/400 +
             skoXC1 *
             (-3/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200)))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_133, refute_0_247])).

cnf(refute_0_249, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 * (2 + skoXC1 * 3) *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     1024 +
     skoXC1 *
     (197376/25 +
      skoXC1 *
      (317272/25 +
       skoXC1 *
       (64043/5 +
        skoXC1 *
        (227585/64 +
         skoXC1 *
         (2849753/1600 +
          skoXC1 *
          (366609/3200 +
           skoXC1 *
           (16431/400 +
            skoXC1 *
            (-3/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200))))))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_28, refute_0_248])).

cnf(refute_0_250, plain,
    (-1024 <
     skoXC1 *
     (-7424/25 +
      skoXC1 *
      (-79528/25 +
       skoXC1 *
       (18123/5 +
        skoXC1 *
        (-166783/64 +
         skoXC1 *
         (1401753/1600 +
          skoXC1 *
          (-756591/3200 +
           skoXC1 *
           (12031/400 +
            skoXC1 *
            (-243/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200))))))))) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (skoXC1 *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     2048 +
     skoXC1 *
     (89088/25 +
      skoXC1 *
      (102704/25 +
       skoXC1 *
       (27314/25 +
        skoXC1 *
        (472553/800 +
         skoXC1 *
         (23107/640 +
          skoXC1 *
          (11013/800 +
           skoXC1 *
           (-27/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 |
     (2048 +
      skoXC1 *
      (89088/25 +
       skoXC1 *
       (102704/25 +
        skoXC1 *
        (27314/25 +
         skoXC1 *
         (472553/800 +
          skoXC1 *
          (23107/640 +
           skoXC1 *
           (11013/800 +
            skoXC1 *
            (-27/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_252, plain,
    (skoXC1 <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (1 + skoXC1 * -1/100) | 1 + skoXC1 * -1/100 <= 0 |
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 +
        skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <=
     skoXC1 / (1 + skoXC1 * -1/100)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_253, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1)))) |
     sqrt(skoXC1 * (1 + skoXC1)) <= skoXC1 / (1 + skoXC1 * -1/100)),
    inference(subst, [], [refute_0_128])).

cnf(refute_0_254, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1)))) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_253, refute_0_68])).

cnf(refute_0_255, plain,
    (skoXC1 * (1 + skoXC1) < 1 |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (2048 +
      skoXC1 * (1 + skoXC1) *
      (3584 +
       skoXC1 * (1 + skoXC1) *
       (560 +
        skoXC1 * (1 + skoXC1) * (14 + skoXC1 * (1 + skoXC1) * 1/32)))) /
     (4096 +
      skoXC1 * (1 + skoXC1) *
      (1792 + skoXC1 * (1 + skoXC1) * (112 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_254, refute_0_199])).

cnf(refute_0_256, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 / (1 + skoXC1 * -1/100) <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_255])).

cnf(refute_0_257, plain,
    (skoXC1 <
     (2048 +
      skoXC1 *
      (3584 +
       skoXC1 *
       (4144 +
        skoXC1 *
        (1134 +
         skoXC1 *
         (19265/32 +
          skoXC1 *
          (337/8 +
           skoXC1 * (227/16 + skoXC1 * (1/8 + skoXC1 * 1/32)))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) *
     (1 + skoXC1 * -1/100) | -1 < skoXC1 * (-1 + skoXC1 * -1) |
     1 + skoXC1 * -1/100 <= 0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_252, refute_0_256])).

cnf(refute_0_258, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 <
     (2048 +
      skoXC1 *
      (89088/25 +
       skoXC1 *
       (102704/25 +
        skoXC1 *
        (27314/25 +
         skoXC1 *
         (472553/800 +
          skoXC1 *
          (23107/640 +
           skoXC1 *
           (11013/800 +
            skoXC1 *
            (-27/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     100 <= skoXC1 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain, (skoXC1 < 100 | -1 < skoXC1 * (-1 + skoXC1 * -1)),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_260, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 <
     (2048 +
      skoXC1 *
      (89088/25 +
       skoXC1 *
       (102704/25 +
        skoXC1 *
        (27314/25 +
         skoXC1 *
         (472553/800 +
          skoXC1 *
          (23107/640 +
           skoXC1 *
           (11013/800 +
            skoXC1 *
            (-27/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200))))))))) /
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_258, refute_0_259])).

cnf(refute_0_261, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (4096 +
      skoXC1 *
      (1792 +
       skoXC1 *
       (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))) <
     2048 +
     skoXC1 *
     (89088/25 +
      skoXC1 *
      (102704/25 +
       skoXC1 *
       (27314/25 +
        skoXC1 *
        (472553/800 +
         skoXC1 *
         (23107/640 +
          skoXC1 *
          (11013/800 +
           skoXC1 *
           (-27/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     4096 +
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     0 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_251, refute_0_260])).

cnf(refute_0_262, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_261])).

cnf(refute_0_263, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) |
     -4096 <
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_264, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     -1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     skoXC1 * (-1 + skoXC1 * -1) <= -1),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_266, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_265, refute_0_264])).

cnf(refute_0_267, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_266, refute_0_0])).

cnf(refute_0_268, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_267])).

cnf(refute_0_269, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_270, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_268, refute_0_269])).

cnf(refute_0_271, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200)))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_246, refute_0_238, refute_0_230, refute_0_222,
               refute_0_160, refute_0_111, refute_0_95, refute_0_85,
               refute_0_15])).

cnf(refute_0_272, plain,
    (-2048 <
     skoXC1 *
     (-13312/25 +
      skoXC1 *
      (57904/25 +
       skoXC1 *
       (-20286/25 +
        skoXC1 *
        (292553/800 +
         skoXC1 *
         (-50493/640 +
          skoXC1 *
          (8613/800 +
           skoXC1 *
           (-1627/1600 + skoXC1 * (3/100 + skoXC1 * -1/3200))))))))),
    inference(resolve, [], [refute_0_270, refute_0_271])).

cnf(refute_0_273, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (-7424/25 +
      skoXC1 *
      (-79528/25 +
       skoXC1 *
       (18123/5 +
        skoXC1 *
        (-166783/64 +
         skoXC1 *
         (1401753/1600 +
          skoXC1 *
          (-756591/3200 +
           skoXC1 *
           (12031/400 +
            skoXC1 *
            (-243/80 + skoXC1 * (143/1600 + skoXC1 * -3/3200))))))))) <=
     -1024 |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096),
    inference(decision, [],
              [refute_0_272, refute_0_246, refute_0_238, refute_0_230,
               refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_274, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1))))) <=
     -4096 | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_273, refute_0_250])).

cnf(refute_0_275, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) |
     -4096 <
     skoXC1 *
     (1792 +
      skoXC1 *
      (1904 + skoXC1 * (225 + skoXC1 * (115 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_272, refute_0_246, refute_0_238, refute_0_230,
               refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_276, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_274, refute_0_275])).

cnf(refute_0_277, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | 1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_276, refute_0_0])).

cnf(refute_0_278, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_277])).

cnf(refute_0_279, plain,
    (-1 < skoXC1 * (-1 + skoXC1 * -1) | -1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_272, refute_0_246, refute_0_238, refute_0_230,
               refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_280, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_278, refute_0_279])).

cnf(refute_0_281, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_272, refute_0_246, refute_0_238, refute_0_230,
               refute_0_222, refute_0_160, refute_0_111, refute_0_95,
               refute_0_85, refute_0_15])).

cnf(refute_0_282, plain, (-1 < skoXC1 * (-1 + skoXC1 * -1)),
    inference(resolve, [], [refute_0_280, refute_0_281])).

cnf(refute_0_283, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_284, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1 | skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_27, refute_0_283])).

cnf(refute_0_285, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) != 0 |
     skoXC1 * (skoXC1 * skoXC1) = 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_286, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1 | skoXC1 * (skoXC1 * skoXC1) = 0 |
     skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_284, refute_0_285])).

cnf(refute_0_287, plain,
    (skoXC1 * (skoXC1 * skoXC1) != 0 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_288, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1 | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_286, refute_0_287])).

cnf(refute_0_289, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1 | skoXC1 = 0),
    inference(resolve, [], [refute_0_288, refute_0_161])).

cnf(refute_0_290, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) <
     (1/6 +
      skoXC1 *
      (112/5 +
       skoXC1 *
       (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462))))) /
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_289, refute_0_112])).

cnf(refute_0_291, plain,
    ((1 + sqrt(skoXC1 * (1 + skoXC1)) * 1/100) /
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (6 +
       skoXC1 *
       (105 +
        skoXC1 *
        (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))) <
     1/6 +
     skoXC1 *
     (112/5 +
      skoXC1 * (469/2 + skoXC1 * (784 + skoXC1 * (1029 + skoXC1 * 462)))) |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_9, refute_0_290])).

cnf(refute_0_292, plain,
    (-1/6 +
     skoXC1 *
     (-112/5 +
      skoXC1 *
      (-469/2 + skoXC1 * (-784 + skoXC1 * (-1029 + skoXC1 * -462)))) <
     (skoXC1 *
      (-6 +
       skoXC1 *
       (-105 +
        skoXC1 *
        (-560 + skoXC1 * (-1260 + skoXC1 * (-1260 + skoXC1 * -462))))) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 *
       (-3/50 +
        skoXC1 *
        (-21/20 +
         skoXC1 *
         (-28/5 +
          skoXC1 * (-63/5 + skoXC1 * (-63/5 + skoXC1 * -231/50))))))) /
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <= 0 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_291])).

cnf(refute_0_293, plain,
    (0 <
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462)))))),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_294, plain,
    (-1/6 +
     skoXC1 *
     (-112/5 +
      skoXC1 *
      (-469/2 + skoXC1 * (-784 + skoXC1 * (-1029 + skoXC1 * -462)))) <
     (skoXC1 *
      (-6 +
       skoXC1 *
       (-105 +
        skoXC1 *
        (-560 + skoXC1 * (-1260 + skoXC1 * (-1260 + skoXC1 * -462))))) +
      sqrt(skoXC1 * (1 + skoXC1)) *
      (skoXC1 *
       (-3/50 +
        skoXC1 *
        (-21/20 +
         skoXC1 *
         (-28/5 +
          skoXC1 * (-63/5 + skoXC1 * (-63/5 + skoXC1 * -231/50))))))) /
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_292, refute_0_293])).

cnf(refute_0_295, plain,
    ((-1/6 +
      skoXC1 *
      (-112/5 +
       skoXC1 *
       (-469/2 + skoXC1 * (-784 + skoXC1 * (-1029 + skoXC1 * -462))))) *
     sqrt(skoXC1 * (1 + skoXC1)) <
     skoXC1 *
     (-6 +
      skoXC1 *
      (-105 +
       skoXC1 *
       (-560 + skoXC1 * (-1260 + skoXC1 * (-1260 + skoXC1 * -462))))) +
     sqrt(skoXC1 * (1 + skoXC1)) *
     (skoXC1 *
      (-3/50 +
       skoXC1 *
       (-21/20 +
        skoXC1 *
        (-28/5 +
         skoXC1 * (-63/5 + skoXC1 * (-63/5 + skoXC1 * -231/50)))))) |
     1 / skoXC1 <= -1 | sqrt(skoXC1 * (1 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_8, refute_0_294])).

cnf(refute_0_296, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) |
     skoXC1 * (1 + skoXC1) <= 0 | 1 / skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_295])).

cnf(refute_0_297, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     sqrt(skoXC1 * (1 + skoXC1)) *
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_296, refute_0_16])).

cnf(refute_0_298, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     1/6 +
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))) <= 0 |
     1 / skoXC1 <= -1), inference(split, [], [refute_0_297])).

cnf(refute_0_299, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     sqrt(skoXC1 * (1 + skoXC1)) |
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))) <=
     -1/6 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (-1/6 <
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_301, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     sqrt(skoXC1 * (1 + skoXC1)) | 1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_299, refute_0_300])).

cnf(refute_0_302, plain,
    (skoXC1 * (1 + skoXC1) < 0 |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     (1/8 +
      skoXC1 * (1 + skoXC1) *
      (7/2 +
       skoXC1 * (1 + skoXC1) *
       (35/4 +
        skoXC1 * (1 + skoXC1) * (7/2 + skoXC1 * (1 + skoXC1) * 1/8)))) /
     (1 +
      skoXC1 * (1 + skoXC1) *
      (7 + skoXC1 * (1 + skoXC1) * (7 + skoXC1 * (1 + skoXC1)))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_7, refute_0_301])).

cnf(refute_0_303, plain,
    (0 < skoXC1 * (-1 + skoXC1 * -1) |
     skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_302])).

cnf(refute_0_304, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) /
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_149, refute_0_303])).

cnf(refute_0_305, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     (1/8 +
      skoXC1 *
      (7/2 +
       skoXC1 *
       (49/4 +
        skoXC1 *
        (21 +
         skoXC1 *
         (155/8 +
          skoXC1 *
          (11 + skoXC1 * (17/4 + skoXC1 * (1/2 + skoXC1 * 1/8)))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) *
     (1/6 +
      skoXC1 *
      (1117/50 +
       skoXC1 *
       (4669/20 +
        skoXC1 *
        (3892/5 +
         skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50)))))) |
     1/6 +
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_2, refute_0_304])).

cnf(refute_0_306, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     (1/48 +
      skoXC1 *
      (4051/1200 +
       skoXC1 *
       (262591/2400 +
        skoXC1 *
        (59577/50 +
         skoXC1 *
         (3710149/600 +
          skoXC1 *
          (4436663/240 +
           skoXC1 *
           (16866799/480 +
            skoXC1 *
            (5349343/120 +
             skoXC1 *
             (4636717/120 +
              skoXC1 *
              (9287289/400 +
               skoXC1 *
               (306943/32 +
                skoXC1 *
                (246463/100 +
                 skoXC1 *
                 (66423/200 +
                  skoXC1 * (10773/200 + skoXC1 * -231/400)))))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     skoXC1 *
     (1117/50 +
      skoXC1 *
      (4669/20 +
       skoXC1 *
       (3892/5 +
        skoXC1 * (5082/5 + skoXC1 * (2247/5 + skoXC1 * -231/50))))) <=
     -1/6 | 1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) <
     (1/48 +
      skoXC1 *
      (4051/1200 +
       skoXC1 *
       (262591/2400 +
        skoXC1 *
        (59577/50 +
         skoXC1 *
         (3710149/600 +
          skoXC1 *
          (4436663/240 +
           skoXC1 *
           (16866799/480 +
            skoXC1 *
            (5349343/120 +
             skoXC1 *
             (4636717/120 +
              skoXC1 *
              (9287289/400 +
               skoXC1 *
               (306943/32 +
                skoXC1 *
                (246463/100 +
                 skoXC1 *
                 (66423/200 +
                  skoXC1 * (10773/200 + skoXC1 * -231/400)))))))))))))) /
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_306, refute_0_300])).

cnf(refute_0_308, plain,
    (skoXC1 *
     (6 +
      skoXC1 *
      (105 +
       skoXC1 *
       (560 + skoXC1 * (1260 + skoXC1 * (1260 + skoXC1 * 462))))) *
     (1 +
      skoXC1 *
      (7 +
       skoXC1 *
       (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))) <
     1/48 +
     skoXC1 *
     (4051/1200 +
      skoXC1 *
      (262591/2400 +
       skoXC1 *
       (59577/50 +
        skoXC1 *
        (3710149/600 +
         skoXC1 *
         (4436663/240 +
          skoXC1 *
          (16866799/480 +
           skoXC1 *
           (5349343/120 +
            skoXC1 *
            (4636717/120 +
             skoXC1 *
             (9287289/400 +
              skoXC1 *
              (306943/32 +
               skoXC1 *
               (246463/100 +
                skoXC1 *
                (66423/200 +
                 skoXC1 * (10773/200 + skoXC1 * -231/400))))))))))))) |
     1 +
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= 0 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_1, refute_0_307])).

cnf(refute_0_309, plain,
    (-1/48 <
     skoXC1 *
     (-3149/1200 +
      skoXC1 *
      (-90209/2400 +
       skoXC1 *
       (-9373/50 +
        skoXC1 *
        (-333851/600 +
         skoXC1 *
         (-256537/240 +
          skoXC1 *
          (-600401/480 +
           skoXC1 *
           (-134057/120 +
            skoXC1 *
            (-133643/120 +
             skoXC1 *
             (-260711/400 +
              skoXC1 *
              (-2177/32 +
               skoXC1 *
               (-18137/100 +
                skoXC1 *
                (-25977/200 +
                 skoXC1 * (10773/200 + skoXC1 * -231/400))))))))))))) |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1), inference(arithmetic, [], [refute_0_308])).

cnf(refute_0_310, plain,
    (skoXC1 *
     (-3149/1200 +
      skoXC1 *
      (-90209/2400 +
       skoXC1 *
       (-9373/50 +
        skoXC1 *
        (-333851/600 +
         skoXC1 *
         (-256537/240 +
          skoXC1 *
          (-600401/480 +
           skoXC1 *
           (-134057/120 +
            skoXC1 *
            (-133643/120 +
             skoXC1 *
             (-260711/400 +
              skoXC1 *
              (-2177/32 +
               skoXC1 *
               (-18137/100 +
                skoXC1 *
                (-25977/200 +
                 skoXC1 * (10773/200 + skoXC1 * -231/400))))))))))))) <=
     -1/48 |
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_311, plain,
    (skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1))))) <= -1 |
     1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_310, refute_0_309])).

cnf(refute_0_312, plain,
    (-1 <
     skoXC1 *
     (7 +
      skoXC1 *
      (14 + skoXC1 * (15 + skoXC1 * (10 + skoXC1 * (3 + skoXC1)))))),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_313, plain, (1 / skoXC1 <= -1),
    inference(resolve, [], [refute_0_311, refute_0_312])).

cnf(refute_0_314, plain, (1 <= -1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_313, refute_0_0])).

cnf(refute_0_315, plain, (skoXC1 <= -1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain, (-1 < skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_282, refute_0_272, refute_0_246, refute_0_238,
               refute_0_230, refute_0_222, refute_0_160, refute_0_111,
               refute_0_95, refute_0_85, refute_0_15])).

cnf(refute_0_317, plain, (skoXC1 <= 0),
    inference(resolve, [], [refute_0_315, refute_0_316])).

cnf(refute_0_318, plain, ($false),
    inference(resolve, [], [refute_0_317, refute_0_55])).
% SZS output end CNFRefutation for log-upper-21-sqrt-weak.tptp
