% SZS output start CNFRefutation for log-error-analysis-3-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/3 * (10 * X ^ 2 + 19 * X + 1) * (X - 1) /
            (X * (X ^ 2 + 6 * X + 3))) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf7, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/12 * (47 * X ^ 3 + 239 * X ^ 2 + 131 * X + 3) * (X - 1) /
            (X * (X ^ 3 + 12 * X ^ 2 + 18 * X + 4))) | lgen(R, Y, ln(X)))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(log_error_analysis_3, conjecture,
    (! [X] :
       ((0 <= X & X <= 8/25) =>
        abs(ln(1 + X) -
            (24987389/25000000 * X - 3074431/6250000 * X ^ 2 +
             14473739/50000000 * X ^ 3 - 544251/4000000 * X ^ 4 +
             643169/20000000 * X ^ 5)) <= 10 ^ -5))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 8/25) =>
        abs(ln(1 + X) -
            (24987389/25000000 * X - 3074431/6250000 * X ^ 2 +
             14473739/50000000 * X ^ 3 - 544251/4000000 * X ^ 4 +
             643169/20000000 * X ^ 5)) <= 10 ^ -5)),
    inference(strip, [], [log_error_analysis_3])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 8/25) =>
          abs(ln(1 + X) -
              (24987389/25000000 * X - 3074431/6250000 * X ^ 2 +
               14473739/50000000 * X ^ 3 - 544251/4000000 * X ^ 4 +
               643169/20000000 * X ^ 5)) <= 10 ^ -5)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (10 ^ -5 <
        abs(ln(1 + X) -
            (24987389/25000000 * X - 3074431/6250000 * X ^ 2 +
             14473739/50000000 * X ^ 3 - 544251/4000000 * X ^ 4 +
             643169/20000000 * X ^ 5)) & X <= 8/25 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (10 ^ -5 <
     abs(ln(1 + skoXC1) -
         (24987389/25000000 * skoXC1 - 3074431/6250000 * skoXC1 ^ 2 +
          14473739/50000000 * skoXC1 ^ 3 - 544251/4000000 * skoXC1 ^ 4 +
          643169/20000000 * skoXC1 ^ 5)) & 0 <= skoXC1 & skoXC1 <= 8/25),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (10 ^ -5 <
     abs(ln(1 + skoXC1) -
         (24987389/25000000 * skoXC1 - 3074431/6250000 * skoXC1 ^ 2 +
          14473739/50000000 * skoXC1 ^ 3 - 544251/4000000 * skoXC1 ^ 4 +
          643169/20000000 * skoXC1 ^ 5))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 8/25),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (10 + skoXC1 * (12 + skoXC1 * 3)) <
     skoXC1 * (10 + skoXC1 * (7 + skoXC1 * 1/3)) |
     10 + skoXC1 * (12 + skoXC1 * 3) <= 0 |
     skoXC1 * (10 + skoXC1 * (7 + skoXC1 * 1/3)) /
     (10 + skoXC1 * (12 + skoXC1 * 3)) <=
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, ln(X_000063), X_000064) | ln(X_000063) <= X_000064),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0 | lgen(0, ln(X_000063), X_000064)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_3, plain,
    (~ lgen(0,
            1/3 * (X_000063 ^ 2 + 19 * X_000063 + 10) * (X_000063 - 1) /
            (3 * X_000063 ^ 2 + 6 * X_000063 + 1), X_000064) |
     X_000063 <= 0 | ln(X_000063) <= X_000064),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000064 <
     (-10/3 + X_000063 * (-3 + X_000063 * (6 + X_000063 * 1/3))) /
     (1 + X_000063 * (6 + X_000063 * 3)) | X_000063 <= 0 |
     ln(X_000063) <= X_000064), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     (-10/3 +
      (1 + skoXC1) * (-3 + (1 + skoXC1) * (6 + (1 + skoXC1) * 1/3))) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1) * 3)) | 1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <=
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (10 ^ -5 <
     abs(ln(1 + skoXC1) -
         (24987389/25000000 * skoXC1 - 3074431/6250000 * skoXC1 ^ 2 +
          14473739/50000000 * skoXC1 ^ 3 - 544251/4000000 * skoXC1 ^ 4 +
          643169/20000000 * skoXC1 ^ 5))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (1/100000 <
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1))), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) < 0 |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) =
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1)), inference(subst, [], [abs_nonnegative])).

cnf(refute_0_9, plain,
    (1/100000 <
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) !=
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) <= 1/100000), introduced(tautology, [equality])).

cnf(refute_0_10, plain,
    (skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) < 0 |
     1/100000 <
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) <= 1/100000),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) < 0 |
     1/100000 <
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_10, refute_0_7])).

cnf(refute_0_12, plain,
    (ln(1 + skoXC1) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     ln(1 + skoXC1)), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (skoXC1 * (35 + skoXC1 * (125/2 + skoXC1 * (95/3 + skoXC1 * 47/12))) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1)))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 * (35 + skoXC1 * (125/2 + skoXC1 * (95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1)))) |
     35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_14, plain,
    (0 <=
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) =
     -(skoXC1 *
      (-24987389/25000000 +
       skoXC1 *
       (3074431/6250000 +
        skoXC1 *
        (-14473739/50000000 +
         skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
      ln(1 + skoXC1))), inference(subst, [], [abs_negative])).

cnf(refute_0_15, plain,
    (1/100000 <
     -(skoXC1 *
      (-24987389/25000000 +
       skoXC1 *
       (3074431/6250000 +
        skoXC1 *
        (-14473739/50000000 +
         skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
      ln(1 + skoXC1)) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) !=
     -(skoXC1 *
      (-24987389/25000000 +
       skoXC1 *
       (3074431/6250000 +
        skoXC1 *
        (-14473739/50000000 +
         skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
      ln(1 + skoXC1)) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) <= 1/100000), introduced(tautology, [equality])).

cnf(refute_0_16, plain,
    (1/100000 <
     -(skoXC1 *
      (-24987389/25000000 +
       skoXC1 *
       (3074431/6250000 +
        skoXC1 *
        (-14473739/50000000 +
         skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
      ln(1 + skoXC1)) |
     0 <=
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1) |
     abs(skoXC1 *
         (-24987389/25000000 +
          skoXC1 *
          (3074431/6250000 +
           skoXC1 *
           (-14473739/50000000 +
            skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
         ln(1 + skoXC1)) <= 1/100000),
    inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain,
    (1/100000 <
     -(skoXC1 *
      (-24987389/25000000 +
       skoXC1 *
       (3074431/6250000 +
        skoXC1 *
        (-14473739/50000000 +
         skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
      ln(1 + skoXC1)) |
     0 <=
     skoXC1 *
     (-24987389/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) +
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_16, refute_0_7])).

cnf(refute_0_18, plain,
    (ln(1 + skoXC1) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000085, ln(X_000084)) | X_000085 <= ln(X_000084)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000085,
            1/12 *
            (47 * X_000084 ^ 3 + 239 * X_000084 ^ 2 + 131 * X_000084 + 3) *
            (X_000084 - 1) /
            (X_000084 *
             (X_000084 ^ 3 + 12 * X_000084 ^ 2 + 18 * X_000084 + 4))) |
     X_000084 <= 0 | lgen(0, X_000085, ln(X_000084))),
    inference(subst, [], [ln_lower_bound_cf7])).

cnf(refute_0_21, plain,
    (~ lgen(0, X_000085,
            1/12 *
            (47 * X_000084 ^ 3 + 239 * X_000084 ^ 2 + 131 * X_000084 + 3) *
            (X_000084 - 1) /
            (X_000084 *
             (X_000084 ^ 3 + 12 * X_000084 ^ 2 + 18 * X_000084 + 4))) |
     X_000084 <= 0 | X_000085 <= ln(X_000084)),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    ((-1/4 +
      X_000084 *
      (-32/3 + X_000084 * (-9 + X_000084 * (16 + X_000084 * 47/12)))) /
     (X_000084 * (4 + X_000084 * (18 + X_000084 * (12 + X_000084)))) <
     X_000085 | X_000084 <= 0 | X_000085 <= ln(X_000084)),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    ((-1/4 +
      (1 + skoXC1) *
      (-32/3 +
       (1 + skoXC1) * (-9 + (1 + skoXC1) * (16 + (1 + skoXC1) * 47/12)))) /
     ((1 + skoXC1) *
      (4 + (1 + skoXC1) * (18 + (1 + skoXC1) * (12 + (1 + skoXC1))))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    ((-1/4 +
      (1 + skoXC1) *
      (-32/3 +
       (1 + skoXC1) * (-9 + (1 + skoXC1) * (16 + (1 + skoXC1) * 47/12)))) /
     ((1 + skoXC1) *
      (4 + (1 + skoXC1) * (18 + (1 + skoXC1) * (12 + (1 + skoXC1))))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_23, refute_0_18])).

cnf(refute_0_25, plain,
    (skoXC1 * (35 + skoXC1 * (125/2 + skoXC1 * (95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1)))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (1 + skoXC1) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 / (1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000043, ln(X_000042)) | X_000043 <= ln(X_000042)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     lgen(0, X_000043, ln(X_000042))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_29, plain,
    (~ lgen(0, X_000043, (X_000042 - 1) / X_000042) | X_000042 <= 0 |
     X_000043 <= ln(X_000042)),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    ((-1 + X_000042) / X_000042 < X_000043 | X_000042 <= 0 |
     X_000043 <= ln(X_000042)), inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((-1 + (1 + skoXC1)) / (1 + skoXC1) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((-1 + (1 + skoXC1)) / (1 + skoXC1) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_31, refute_0_18])).

cnf(refute_0_33, plain,
    (skoXC1 / (1 + skoXC1) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_35, plain, (skoXC1 <= 8/25),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_36, plain, (-1 < skoXC1),
    inference(decision, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (skoXC1 / (1 + skoXC1) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_33, refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (1 + skoXC1) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_26, refute_0_37])).

cnf(refute_0_39, plain,
    (1/100000 <
     skoXC1 *
     (-12861/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (skoXC1 <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) *
     (1 + skoXC1) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 / (1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_41, plain,
    ((-1 + (1 + skoXC1)) / (1 + skoXC1) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_42, plain,
    (skoXC1 / (1 + skoXC1) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (skoXC1 / (1 + skoXC1) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_42, refute_0_36])).

cnf(refute_0_44, plain,
    (skoXC1 <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) *
     (1 + skoXC1) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_40, refute_0_43])).

cnf(refute_0_45, plain,
    (0 <
     skoXC1 *
     (-12611/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (0 <
     skoXC1 *
     (-12611/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     -1 < skoXC1), inference(decision, [], [refute_0_34, refute_0_35])).

cnf(refute_0_47, plain,
    (0 <
     skoXC1 *
     (-12611/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain,
    (1/100000 <
     skoXC1 *
     (-12861/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     -1 < skoXC1 |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [], [refute_0_47, refute_0_34, refute_0_35])).

cnf(refute_0_49, plain,
    (1/100000 <
     skoXC1 *
     (-12861/25000000 +
      skoXC1 *
      (2537933/5000000 +
       skoXC1 *
       (-10121709/50000000 +
        skoXC1 *
        (15341203/100000000 +
         skoXC1 * (-1039043/10000000 + skoXC1 * 643169/20000000))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_39, refute_0_48])).

cnf(refute_0_50, plain,
    (-1 < skoXC1 |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_49, refute_0_47, refute_0_34, refute_0_35])).

cnf(refute_0_51, plain,
    (skoXC1 * (35 + skoXC1 * (125/2 + skoXC1 * (95/3 + skoXC1 * 47/12))) /
     (35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1)))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_25, refute_0_50])).

cnf(refute_0_52, plain,
    (skoXC1 * (35 + skoXC1 * (125/2 + skoXC1 * (95/3 + skoXC1 * 47/12))) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) |
     35 + skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) <= 0),
    inference(resolve, [], [refute_0_13, refute_0_51])).

cnf(refute_0_53, plain,
    (7/20000 <
     skoXC1 *
     (-92277/5000000 +
      skoXC1 *
      (302789/1250000 +
       skoXC1 *
       (-27545777/30000000 +
        skoXC1 *
        (286950253/300000000 +
         skoXC1 *
         (2951859/4000000 +
          skoXC1 *
          (-36285037/25000000 +
           skoXC1 *
           (2098889/50000000 +
            skoXC1 * (7569449/20000000 + skoXC1 * 643169/20000000)))))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) |
     skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) <= -35),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 * (10 + skoXC1 * (13 + skoXC1 * 10/3)) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 * (10 + skoXC1 * (13 + skoXC1 * 10/3)) /
     (10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1))) |
     10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_55, plain,
    (~ lgen(0, X_000071, ln(X_000070)) | X_000071 <= ln(X_000070)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_56, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | lgen(0, X_000071, ln(X_000070))),
    inference(subst, [], [ln_lower_bound_cf5])).

cnf(refute_0_57, plain,
    (~ lgen(0, X_000071,
            1/3 * (10 * X_000070 ^ 2 + 19 * X_000070 + 1) *
            (X_000070 - 1) /
            (X_000070 * (X_000070 ^ 2 + 6 * X_000070 + 3))) |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(resolve, [], [refute_0_56, refute_0_55])).

cnf(refute_0_58, plain,
    ((-1/3 + X_000070 * (-6 + X_000070 * (3 + X_000070 * 10/3))) /
     (X_000070 * (3 + X_000070 * (6 + X_000070))) < X_000071 |
     X_000070 <= 0 | X_000071 <= ln(X_000070)),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    ((-1/3 +
      (1 + skoXC1) * (-6 + (1 + skoXC1) * (3 + (1 + skoXC1) * 10/3))) /
     ((1 + skoXC1) * (3 + (1 + skoXC1) * (6 + (1 + skoXC1)))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_58])).

cnf(refute_0_60, plain,
    ((-1/3 +
      (1 + skoXC1) * (-6 + (1 + skoXC1) * (3 + (1 + skoXC1) * 10/3))) /
     ((1 + skoXC1) * (3 + (1 + skoXC1) * (6 + (1 + skoXC1)))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_59, refute_0_18])).

cnf(refute_0_61, plain,
    (skoXC1 * (10 + skoXC1 * (13 + skoXC1 * 10/3)) /
     (10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 * (10 + skoXC1 * (13 + skoXC1 * 10/3)) /
     (10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_61, refute_0_50])).

cnf(refute_0_63, plain,
    (skoXC1 * (10 + skoXC1 * (13 + skoXC1 * 10/3)) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 10 + skoXC1 * (18 + skoXC1 * (9 + skoXC1)) <= 0),
    inference(resolve, [], [refute_0_54, refute_0_62])).

cnf(refute_0_64, plain,
    (1/10000 <
     skoXC1 *
     (-13061/2500000 +
      skoXC1 *
      (224189/3125000 +
       skoXC1 *
       (-11156129/37500000 +
        skoXC1 *
        (21111673/50000000 +
         skoXC1 *
         (-709047/50000000 +
          skoXC1 *
          (-35623787/100000000 +
           skoXC1 * (1533633/10000000 + skoXC1 * 643169/20000000))))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 * (18 + skoXC1 * (9 + skoXC1)) <= -10),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (3 + skoXC1 * (4 + skoXC1)) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) |
     3 + skoXC1 * (4 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_66, plain,
    (~ lgen(0, X_000057, ln(X_000056)) | X_000057 <= ln(X_000056)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_67, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     lgen(0, X_000057, ln(X_000056))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_68, plain,
    (~ lgen(0, X_000057,
            1/2 * (1 + 5 * X_000056) * (X_000056 - 1) /
            (X_000056 * (2 + X_000056))) | X_000056 <= 0 |
     X_000057 <= ln(X_000056)),
    inference(resolve, [], [refute_0_67, refute_0_66])).

cnf(refute_0_69, plain,
    ((-1/2 + X_000056 * (-2 + X_000056 * 5/2)) /
     (X_000056 * (2 + X_000056)) < X_000057 | X_000056 <= 0 |
     X_000057 <= ln(X_000056)), inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    ((-1/2 + (1 + skoXC1) * (-2 + (1 + skoXC1) * 5/2)) /
     ((1 + skoXC1) * (2 + (1 + skoXC1))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_69])).

cnf(refute_0_71, plain,
    ((-1/2 + (1 + skoXC1) * (-2 + (1 + skoXC1) * 5/2)) /
     ((1 + skoXC1) * (2 + (1 + skoXC1))) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_70, refute_0_18])).

cnf(refute_0_72, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) <
     -1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_72, refute_0_50])).

cnf(refute_0_74, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) <
     (-1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (3 + skoXC1 * (4 + skoXC1)) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 3 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_65, refute_0_73])).

cnf(refute_0_75, plain,
    (3/100000 <
     skoXC1 *
     (-38833/25000000 +
      skoXC1 *
      (278067/12500000 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 * (4 + skoXC1) <= -3),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) *
     (3 + skoXC1 * (4 + skoXC1)) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) |
     3 + skoXC1 * (4 + skoXC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_77, plain,
    ((-1/2 + (1 + skoXC1) * (-2 + (1 + skoXC1) * 5/2)) /
     ((1 + skoXC1) * (2 + (1 + skoXC1))) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(subst, [], [refute_0_69])).

cnf(refute_0_78, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) / (3 + skoXC1 * (4 + skoXC1)) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_78, refute_0_50])).

cnf(refute_0_80, plain,
    (skoXC1 * (3 + skoXC1 * 5/2) <
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) *
     (3 + skoXC1 * (4 + skoXC1)) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | 3 + skoXC1 * (4 + skoXC1) <= 0),
    inference(resolve, [], [refute_0_76, refute_0_79])).

cnf(refute_0_81, plain,
    (0 <
     skoXC1 *
     (-37833/25000000 +
      skoXC1 *
      (17387/781250 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) | skoXC1 * (4 + skoXC1) <= -3),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (0 <
     skoXC1 *
     (-37833/25000000 +
      skoXC1 *
      (17387/781250 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     -3 < skoXC1 * (4 + skoXC1) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_49, refute_0_47, refute_0_34, refute_0_35])).

cnf(refute_0_83, plain,
    (0 <
     skoXC1 *
     (-37833/25000000 +
      skoXC1 *
      (17387/781250 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (3/100000 <
     skoXC1 *
     (-38833/25000000 +
      skoXC1 *
      (278067/12500000 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     -3 < skoXC1 * (4 + skoXC1) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_83, refute_0_49, refute_0_47, refute_0_34,
               refute_0_35])).

cnf(refute_0_85, plain,
    (3/100000 <
     skoXC1 *
     (-38833/25000000 +
      skoXC1 *
      (278067/12500000 +
       skoXC1 *
       (-4985797/50000000 +
        skoXC1 *
        (25780191/100000000 +
         skoXC1 *
         (-15830087/100000000 +
          skoXC1 * (-148579/20000000 + skoXC1 * 643169/20000000)))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_75, refute_0_84])).

cnf(refute_0_86, plain,
    (1/10000 <
     skoXC1 *
     (-13061/2500000 +
      skoXC1 *
      (224189/3125000 +
       skoXC1 *
       (-11156129/37500000 +
        skoXC1 *
        (21111673/50000000 +
         skoXC1 *
         (-709047/50000000 +
          skoXC1 *
          (-35623787/100000000 +
           skoXC1 * (1533633/10000000 + skoXC1 * 643169/20000000))))))) |
     -10 < skoXC1 * (18 + skoXC1 * (9 + skoXC1)) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_85, refute_0_83, refute_0_49, refute_0_47,
               refute_0_34, refute_0_35])).

cnf(refute_0_87, plain,
    (1/10000 <
     skoXC1 *
     (-13061/2500000 +
      skoXC1 *
      (224189/3125000 +
       skoXC1 *
       (-11156129/37500000 +
        skoXC1 *
        (21111673/50000000 +
         skoXC1 *
         (-709047/50000000 +
          skoXC1 *
          (-35623787/100000000 +
           skoXC1 * (1533633/10000000 + skoXC1 * 643169/20000000))))))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_64, refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 *
     (-92277/5000000 +
      skoXC1 *
      (302789/1250000 +
       skoXC1 *
       (-27545777/30000000 +
        skoXC1 *
        (286950253/300000000 +
         skoXC1 *
         (2951859/4000000 +
          skoXC1 *
          (-36285037/25000000 +
           skoXC1 *
           (2098889/50000000 +
            skoXC1 *
            (7569449/20000000 + skoXC1 * 643169/20000000)))))))) <=
     7/20000 |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) |
     skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) <= -35),
    inference(decision, [],
              [refute_0_87, refute_0_85, refute_0_83, refute_0_49,
               refute_0_47, refute_0_34, refute_0_35])).

cnf(refute_0_89, plain,
    (skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1) |
     skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) <= -35),
    inference(resolve, [], [refute_0_88, refute_0_53])).

cnf(refute_0_90, plain,
    (-35 < skoXC1 * (80 + skoXC1 * (60 + skoXC1 * (16 + skoXC1))) |
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)),
    inference(decision, [],
              [refute_0_87, refute_0_85, refute_0_83, refute_0_49,
               refute_0_47, refute_0_34, refute_0_35])).

cnf(refute_0_91, plain,
    (skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <=
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     ln(1 + skoXC1)), inference(resolve, [], [refute_0_91, refute_0_12])).

cnf(refute_0_93, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     (-10/3 +
      (1 + skoXC1) * (-3 + (1 + skoXC1) * (6 + (1 + skoXC1) * 1/3))) /
     (1 + (1 + skoXC1) * (6 + (1 + skoXC1) * 3)) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_92])).

cnf(refute_0_94, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     skoXC1 * (10 + skoXC1 * (7 + skoXC1 * 1/3)) /
     (10 + skoXC1 * (12 + skoXC1 * 3)) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     skoXC1 * (10 + skoXC1 * (7 + skoXC1 * 1/3)) /
     (10 + skoXC1 * (12 + skoXC1 * 3))),
    inference(resolve, [], [refute_0_94, refute_0_36])).

cnf(refute_0_96, plain,
    ((1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (10 + skoXC1 * (12 + skoXC1 * 3)) <
     skoXC1 * (10 + skoXC1 * (7 + skoXC1 * 1/3)) |
     10 + skoXC1 * (12 + skoXC1 * 3) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_95])).

cnf(refute_0_97, plain,
    (1/10000 <
     skoXC1 *
     (12311/2500000 +
      skoXC1 *
      (-936089/12500000 +
       skoXC1 *
       (12862739/37500000 +
        skoXC1 *
        (-31867149/50000000 +
         skoXC1 *
         (2767151/6250000 +
          skoXC1 * (445737/20000000 + skoXC1 * -1929507/20000000)))))) |
     skoXC1 * (12 + skoXC1 * 3) <= -10),
    inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    ((1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (3 + skoXC1 * 2) < skoXC1 * (3 + skoXC1 * 1/2) | 3 + skoXC1 * 2 <= 0 |
     skoXC1 * (3 + skoXC1 * 1/2) / (3 + skoXC1 * 2) <=
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_99, plain,
    (~ lgen(0, ln(X_000049), X_000050) | ln(X_000049) <= X_000050),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_100, plain,
    (~ lgen(0, 1/2 * (X_000049 + 5) * (X_000049 - 1) / (2 * X_000049 + 1),
            X_000050) | X_000049 <= 0 | lgen(0, ln(X_000049), X_000050)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_101, plain,
    (~ lgen(0, 1/2 * (X_000049 + 5) * (X_000049 - 1) / (2 * X_000049 + 1),
            X_000050) | X_000049 <= 0 | ln(X_000049) <= X_000050),
    inference(resolve, [], [refute_0_100, refute_0_99])).

cnf(refute_0_102, plain,
    (X_000050 <
     (-5/2 + X_000049 * (2 + X_000049 * 1/2)) / (1 + X_000049 * 2) |
     X_000049 <= 0 | ln(X_000049) <= X_000050),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     (-5/2 + (1 + skoXC1) * (2 + (1 + skoXC1) * 1/2)) /
     (1 + (1 + skoXC1) * 2) | 1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <=
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))),
    inference(subst, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     (-5/2 + (1 + skoXC1) * (2 + (1 + skoXC1) * 1/2)) /
     (1 + (1 + skoXC1) * 2) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_103, refute_0_92])).

cnf(refute_0_105, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     skoXC1 * (3 + skoXC1 * 1/2) / (3 + skoXC1 * 2) | skoXC1 <= -1),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     skoXC1 * (3 + skoXC1 * 1/2) / (3 + skoXC1 * 2)),
    inference(resolve, [], [refute_0_105, refute_0_36])).

cnf(refute_0_107, plain,
    ((1/100000 +
      skoXC1 *
      (24987389/25000000 +
       skoXC1 *
       (-3074431/6250000 +
        skoXC1 *
        (14473739/50000000 +
         skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))) *
     (3 + skoXC1 * 2) < skoXC1 * (3 + skoXC1 * 1/2) | 3 + skoXC1 * 2 <= 0),
    inference(resolve, [], [refute_0_98, refute_0_106])).

cnf(refute_0_108, plain,
    (3/100000 <
     skoXC1 *
     (37333/25000000 +
      skoXC1 *
      (-290803/12500000 +
       skoXC1 *
       (5769679/50000000 +
        skoXC1 *
        (-17076131/100000000 +
         skoXC1 * (3513003/20000000 + skoXC1 * -643169/10000000))))) |
     skoXC1 <= -3/2), inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (~ lgen(0, ln(X_000035), X_000036) | ln(X_000035) <= X_000036),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_110, plain,
    (~ lgen(0, X_000035 - 1, X_000036) | X_000035 <= 0 |
     lgen(0, ln(X_000035), X_000036)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_111, plain,
    (~ lgen(0, X_000035 - 1, X_000036) | X_000035 <= 0 |
     ln(X_000035) <= X_000036),
    inference(resolve, [], [refute_0_110, refute_0_109])).

cnf(refute_0_112, plain,
    (X_000036 < -1 + X_000035 | X_000035 <= 0 | ln(X_000035) <= X_000036),
    inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     -1 + (1 + skoXC1) | 1 + skoXC1 <= 0 |
     ln(1 + skoXC1) <=
     1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000))))),
    inference(subst, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (1/100000 +
     skoXC1 *
     (24987389/25000000 +
      skoXC1 *
      (-3074431/6250000 +
       skoXC1 *
       (14473739/50000000 +
        skoXC1 * (-544251/4000000 + skoXC1 * 643169/20000000)))) <
     -1 + (1 + skoXC1) | 1 + skoXC1 <= 0),
    inference(resolve, [], [refute_0_113, refute_0_92])).

cnf(refute_0_115, plain,
    (1/100000 <
     skoXC1 *
     (12611/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) |
     skoXC1 <= -1), inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (1/100000 <
     skoXC1 *
     (12611/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000)))) |
     -1 < skoXC1), inference(decision, [], [refute_0_35, refute_0_34])).

cnf(refute_0_117, plain,
    (1/100000 <
     skoXC1 *
     (12611/25000000 +
      skoXC1 *
      (3074431/6250000 +
       skoXC1 *
       (-14473739/50000000 +
        skoXC1 * (544251/4000000 + skoXC1 * -643169/20000000))))),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (3/100000 <
     skoXC1 *
     (37333/25000000 +
      skoXC1 *
      (-290803/12500000 +
       skoXC1 *
       (5769679/50000000 +
        skoXC1 *
        (-17076131/100000000 +
         skoXC1 * (3513003/20000000 + skoXC1 * -643169/10000000))))) |
     -3/2 < skoXC1),
    inference(decision, [], [refute_0_117, refute_0_35, refute_0_34])).

cnf(refute_0_119, plain,
    (3/100000 <
     skoXC1 *
     (37333/25000000 +
      skoXC1 *
      (-290803/12500000 +
       skoXC1 *
       (5769679/50000000 +
        skoXC1 *
        (-17076131/100000000 +
         skoXC1 * (3513003/20000000 + skoXC1 * -643169/10000000)))))),
    inference(resolve, [], [refute_0_108, refute_0_118])).

cnf(refute_0_120, plain,
    (skoXC1 *
     (12311/2500000 +
      skoXC1 *
      (-936089/12500000 +
       skoXC1 *
       (12862739/37500000 +
        skoXC1 *
        (-31867149/50000000 +
         skoXC1 *
         (2767151/6250000 +
          skoXC1 * (445737/20000000 + skoXC1 * -1929507/20000000)))))) <=
     1/10000 | skoXC1 * (12 + skoXC1 * 3) <= -10),
    inference(decision, [],
              [refute_0_119, refute_0_117, refute_0_35, refute_0_34])).

cnf(refute_0_121, plain, (skoXC1 * (12 + skoXC1 * 3) <= -10),
    inference(resolve, [], [refute_0_120, refute_0_97])).

cnf(refute_0_122, plain, (-10 < skoXC1 * (12 + skoXC1 * 3)),
    inference(decision, [],
              [refute_0_119, refute_0_117, refute_0_35, refute_0_34])).

cnf(refute_0_123, plain, ($false),
    inference(resolve, [], [refute_0_121, refute_0_122])).
% SZS output end CNFRefutation for log-error-analysis-3-weak.tptp
