% SZS output start CNFRefutation for Mastellone2c.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(sin_lower_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, Y, X) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(log_def, axiom, (log(X) = 868589/2000000 * ln(X))).

fof(lyapunov_1, conjecture,
    (! [X1, X2, X3] :
       ((6/5 < X3 & X3 < 77/25) =>
        0 <=
        X1 ^ 2 * sin(X3) + X2 ^ 4 - X1 * X2 * cos(X3) +
        X2 ^ 2 * X3 ^ 2 * log(X3)))).

fof(subgoal_0, plain,
    (! [X1, X2, X3] :
       ((6/5 < X3 & X3 < 77/25) =>
        0 <=
        X1 ^ 2 * sin(X3) + X2 ^ 4 - X1 * X2 * cos(X3) +
        X2 ^ 2 * X3 ^ 2 * log(X3))), inference(strip, [], [lyapunov_1])).

fof(negate_0_0, plain,
    (~ ! [X1, X2, X3] :
         ((6/5 < X3 & X3 < 77/25) =>
          0 <=
          X1 ^ 2 * sin(X3) + X2 ^ 4 - X1 * X2 * cos(X3) +
          X2 ^ 2 * X3 ^ 2 * log(X3))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X1, X2, X3] :
       (6/5 < X3 &
        X1 ^ 2 * sin(X3) + X2 ^ 4 - X1 * X2 * cos(X3) +
        X2 ^ 2 * X3 ^ 2 * log(X3) < 0 & X3 < 77/25)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) < 0 & skoX3C1 < 77/25 &
     6/5 < skoX3C1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) < 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (6/5 < skoX3C1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoX3C1 < 77/25),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) *
     (2 + skoX3C1) <
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) |
     2 + skoX3C1 <= 0 |
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) /
     (2 + skoX3C1) <=
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     (-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1)) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     (-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000122, ln(X_000121)) | X_000122 <= ln(X_000121)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (~ lgen(0, X_000122,
            1/2 * (1 + 5 * X_000121) * (X_000121 - 1) /
            (X_000121 * (2 + X_000121))) | X_000121 <= 0 |
     lgen(0, X_000122, ln(X_000121))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_4, plain,
    (~ lgen(0, X_000122,
            1/2 * (1 + 5 * X_000121) * (X_000121 - 1) /
            (X_000121 * (2 + X_000121))) | X_000121 <= 0 |
     X_000122 <= ln(X_000121)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    ((-1/2 + X_000121 * (-2 + X_000121 * 5/2)) /
     (X_000121 * (2 + X_000121)) < X_000122 | X_000121 <= 0 |
     X_000122 <= ln(X_000121)), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1)) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     ln(skoX3C1) | skoX3C1 <= 0), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/2 +
        skoX3C1 *
        (skoX3C1 *
         (1/24 +
          skoX3C1 *
          (skoX3C1 *
           (-1/720 +
            skoX3C1 *
            (skoX3C1 *
             (1/40320 + skoX3C1 * (skoX3C1 * -1/3628800)))))))))) *
     (skoX2C1 * skoX1C1) | 0 <= skoX2C1 * skoX1C1 |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <=
     1 +
     skoX3C1 *
     (skoX3C1 *
      (-1/2 +
       skoX3C1 *
       (skoX3C1 *
        (1/24 +
         skoX3C1 *
         (skoX3C1 *
          (-1/720 +
           skoX3C1 *
           (skoX3C1 * (1/40320 + skoX3C1 * (skoX3C1 * -1/3628800)))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000025, cos(X_000024)) | X_000025 <= cos(X_000024)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000025,
            1 - X_000024 ^ 2 / 2 + X_000024 ^ 4 / 24 - X_000024 ^ 6 / 720 +
            X_000024 ^ 8 / 40320 - X_000024 ^ 10 / 3628800) |
     lgen(0, X_000025, cos(X_000024))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_10, plain,
    (~ lgen(0, X_000025,
            1 - X_000024 ^ 2 / 2 + X_000024 ^ 4 / 24 - X_000024 ^ 6 / 720 +
            X_000024 ^ 8 / 40320 - X_000024 ^ 10 / 3628800) |
     X_000025 <= cos(X_000024)),
    inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (1 +
     X_000024 *
     (X_000024 *
      (-1/2 +
       X_000024 *
       (X_000024 *
        (1/24 +
         X_000024 *
         (X_000024 *
          (-1/720 +
           X_000024 *
           (X_000024 *
            (1/40320 + X_000024 * (X_000024 * -1/3628800))))))))) <
     X_000025 | X_000025 <= cos(X_000024)),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 +
     skoX3C1 *
     (skoX3C1 *
      (-1/2 +
       skoX3C1 *
       (skoX3C1 *
        (1/24 +
         skoX3C1 *
         (skoX3C1 *
          (-1/720 +
           skoX3C1 *
           (skoX3C1 * (1/40320 + skoX3C1 * (skoX3C1 * -1/3628800))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <= cos(skoX3C1)),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) *
     (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <=
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_14, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) < 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_15, plain, (log(skoX3C1) = 868589/2000000 * ln(skoX3C1)),
    inference(subst, [], [log_def])).

cnf(refute_0_16, plain,
    (skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1)),
    introduced(tautology, [refl])).

cnf(refute_0_17, plain,
    (skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) !=
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) |
     log(skoX3C1) != 868589/2000000 * ln(skoX3C1) |
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    introduced(tautology, [equality])).

cnf(refute_0_18, plain,
    (log(skoX3C1) != 868589/2000000 * ln(skoX3C1) |
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    inference(resolve, [], [refute_0_15, refute_0_18])).

cnf(refute_0_20, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1)),
    introduced(tautology, [refl])).

cnf(refute_0_21, plain,
    (skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) !=
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) |
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) !=
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) |
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    introduced(tautology, [equality])).

cnf(refute_0_22, plain,
    (skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) !=
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) |
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) =
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1))),
    inference(resolve, [], [refute_0_19, refute_0_22])).

cnf(refute_0_24, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) < 0 |
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1) !=
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) |
     0 <=
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1)),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) < 0 |
     0 <=
     skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * log(skoX3C1)),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (skoX1C1 ^ 2 * sin(skoX3C1) + skoX2C1 ^ 4 -
     skoX1C1 * skoX2C1 * cos(skoX3C1) +
     skoX2C1 ^ 2 * skoX3C1 ^ 2 * (868589/2000000 * ln(skoX3C1)) < 0),
    inference(resolve, [], [refute_0_25, refute_0_14])).

cnf(refute_0_27, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     sin(skoX3C1) * (skoX1C1 * (skoX1C1 * -1))),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (sin(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1)),
    inference(split, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (sin(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | skoX1C1 * skoX1C1 <= 0),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (~ lgen(0, X_000164, sin(X_000163)) | X_000164 <= sin(X_000163)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_31, plain,
    (X_000163 < 0 |
     ~ lgen(0, X_000164,
            X_000163 - X_000163 ^ 3 / 6 + X_000163 ^ 5 / 120 -
            X_000163 ^ 7 / 5040 + X_000163 ^ 9 / 362880 -
            X_000163 ^ 11 / 39916800) | lgen(0, X_000164, sin(X_000163))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_32, plain,
    (X_000163 < 0 |
     ~ lgen(0, X_000164,
            X_000163 - X_000163 ^ 3 / 6 + X_000163 ^ 5 / 120 -
            X_000163 ^ 7 / 5040 + X_000163 ^ 9 / 362880 -
            X_000163 ^ 11 / 39916800) | X_000164 <= sin(X_000163)),
    inference(resolve, [], [refute_0_31, refute_0_30])).

cnf(refute_0_33, plain,
    (X_000163 *
     (1 +
      X_000163 *
      (X_000163 *
       (-1/6 +
        X_000163 *
        (X_000163 *
         (1/120 +
          X_000163 *
          (X_000163 *
           (-1/5040 +
            X_000163 *
            (X_000163 *
             (1/362880 + X_000163 * (X_000163 * -1/39916800)))))))))) <
     X_000164 | X_000163 < 0 | X_000164 <= sin(X_000163)),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoX3C1 < 0 |
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <= sin(skoX3C1)),
    inference(subst, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoX3C1 < 0 |
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | skoX1C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_34, refute_0_29])).

cnf(refute_0_36, plain,
    (skoX2C1 * skoX2C1 <
     (-1 + skoX3C1) / skoX3C1 * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 <= skoX3C1 * (skoX3C1 * -868589/2000000) |
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) <=
     (-1 + skoX3C1) / skoX3C1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_37, plain,
    (skoX1C1 * (skoX1C1 * -1) < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) < 0 |
     0 < skoX1C1 * (skoX1C1 * -1)), inference(split, [], [refute_0_27])).

cnf(refute_0_38, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | 0 < skoX1C1 * (skoX1C1 * -1) |
     0 < skoX1C1 * skoX1C1), inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain, (6/5 < skoX3C1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_40, plain, (skoX3C1 < 77/25),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_41, plain,
    (0 < skoX1C1 * skoX1C1 | skoX1C1 * (skoX1C1 * -1) <= 0),
    inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_42, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | 0 < skoX1C1 * skoX1C1),
    inference(resolve, [], [refute_0_41, refute_0_38])).

cnf(refute_0_43, plain,
    (skoX2C1 * skoX1C1 < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) < 0 |
     0 < skoX1C1 * skoX1C1 | 0 < skoX2C1 * skoX1C1),
    inference(split, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX1C1 * skoX1C1 | 0 < skoX2C1 * (skoX1C1 * -1) |
     0 < skoX2C1 * skoX1C1), inference(arithmetic, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (0 < skoX1C1 * skoX1C1 | 0 < skoX2C1 * skoX1C1 |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_46, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX1C1 * skoX1C1 | 0 < skoX2C1 * skoX1C1),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain, (0 < skoX1C1 * skoX1C1 | skoX2C1 * skoX1C1 <= 0),
    inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_48, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX1C1 * skoX1C1),
    inference(resolve, [], [refute_0_47, refute_0_46])).

cnf(refute_0_49, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(split, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX2C1 = 0), inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 = 0), inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_52, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     skoX2C1 = 0), inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    (skoX2C1 * skoX2C1 <
     ln(skoX3C1) * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 <= skoX3C1 * (skoX3C1 * -868589/2000000) |
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) <=
     ln(skoX3C1)), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_54, plain,
    (skoX2C1 * skoX2C1 <
     ln(skoX3C1) * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 < skoX1C1 * skoX1C1 | 0 <= skoX3C1 * (skoX3C1 * -868589/2000000) |
     skoX2C1 = 0), inference(resolve, [], [refute_0_53, refute_0_52])).

cnf(refute_0_55, plain,
    (skoX2C1 * skoX2C1 <
     ln(skoX3C1) * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 < skoX1C1 * skoX1C1 | skoX3C1 * (skoX3C1 * 868589/2000000) <= 0 |
     skoX2C1 = 0), inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) < 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(split, [], [refute_0_48])).

cnf(refute_0_57, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX2C1 * (skoX2C1 * (skoX2C1 * (skoX2C1 * -1))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX2C1 * (skoX2C1 * (skoX2C1 * -1))) <= 0),
    inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_59, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [], [refute_0_39, refute_0_40])).

cnf(refute_0_61, plain,
    (0 < skoX1C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (0 < skoX1C1 * skoX1C1 | 0 < skoX3C1 * (skoX3C1 * 868589/2000000) |
     skoX2C1 = 0),
    inference(decision, [], [refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_63, plain,
    (skoX2C1 * skoX2C1 <
     ln(skoX3C1) * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 < skoX1C1 * skoX1C1 | skoX2C1 = 0),
    inference(resolve, [], [refute_0_55, refute_0_62])).

cnf(refute_0_64, plain, (0 < skoX1C1 * skoX1C1 | skoX2C1 != 0),
    inference(decision, [], [refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_65, plain,
    (skoX2C1 * skoX2C1 <
     ln(skoX3C1) * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 < skoX1C1 * skoX1C1),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 <= skoX3C1 * (skoX3C1 * -868589/2000000)),
    inference(split, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     skoX3C1 * (skoX3C1 * 868589/2000000) <= 0),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (0 < skoX1C1 * skoX1C1 | 0 < skoX3C1 * (skoX3C1 * 868589/2000000)),
    inference(decision, [], [refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_69, plain,
    (0 < skoX1C1 * skoX1C1 |
     ln(skoX3C1) <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000))),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (~ lgen(0, X_000090, ln(X_000089)) | X_000090 <= ln(X_000089)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_71, plain,
    (~ lgen(0, X_000090, (X_000089 - 1) / X_000089) | X_000089 <= 0 |
     lgen(0, X_000090, ln(X_000089))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_72, plain,
    (~ lgen(0, X_000090, (X_000089 - 1) / X_000089) | X_000089 <= 0 |
     X_000090 <= ln(X_000089)),
    inference(resolve, [], [refute_0_71, refute_0_70])).

cnf(refute_0_73, plain,
    ((-1 + X_000089) / X_000089 < X_000090 | X_000089 <= 0 |
     X_000090 <= ln(X_000089)), inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) <=
     ln(skoX3C1) | skoX3C1 <= 0), inference(subst, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (0 < skoX1C1 * skoX1C1 |
     (-1 + skoX3C1) / skoX3C1 <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     skoX3C1 <= 0), inference(resolve, [], [refute_0_74, refute_0_69])).

cnf(refute_0_76, plain,
    (0 < skoX3C1 * (skoX3C1 * 868589/2000000) | skoX1C1 * skoX1C1 <= 0),
    inference(decision, [], [refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_77, plain, (0 < skoX3C1 * (skoX3C1 * 868589/2000000)),
    inference(resolve, [], [refute_0_76, refute_0_68])).

cnf(refute_0_78, plain, (0 < skoX1C1 * skoX1C1 | 0 < skoX3C1),
    inference(decision, [],
              [refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_79, plain,
    (0 < skoX1C1 * skoX1C1 |
     (-1 + skoX3C1) / skoX3C1 <
     skoX2C1 * skoX2C1 / (skoX3C1 * (skoX3C1 * -868589/2000000))),
    inference(resolve, [], [refute_0_75, refute_0_78])).

cnf(refute_0_80, plain,
    (skoX2C1 * skoX2C1 <
     (-1 + skoX3C1) / skoX3C1 * (skoX3C1 * (skoX3C1 * -868589/2000000)) |
     0 < skoX1C1 * skoX1C1 | 0 <= skoX3C1 * (skoX3C1 * -868589/2000000)),
    inference(resolve, [], [refute_0_36, refute_0_79])).

cnf(refute_0_81, plain,
    (0 < skoX1C1 * skoX1C1 |
     skoX2C1 * skoX2C1 <
     skoX3C1 * (868589/2000000 + skoX3C1 * -868589/2000000) |
     skoX3C1 * (skoX3C1 * 868589/2000000) <= 0 | skoX3C1 = 0),
    inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | skoX2C1 * skoX1C1 <= 0 |
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_83, plain,
    (~ lgen(0, cos(X_000017), X_000018) | cos(X_000017) <= X_000018),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_84, plain,
    (~ lgen(0, 1 - X_000017 ^ 2 / 2 + X_000017 ^ 4 / 24, X_000018) |
     lgen(0, cos(X_000017), X_000018)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_85, plain,
    (~ lgen(0, 1 - X_000017 ^ 2 / 2 + X_000017 ^ 4 / 24, X_000018) |
     cos(X_000017) <= X_000018),
    inference(resolve, [], [refute_0_84, refute_0_83])).

cnf(refute_0_86, plain,
    (X_000018 <
     1 + X_000017 * (X_000017 * (-1/2 + X_000017 * (X_000017 * 1/24))) |
     cos(X_000017) <= X_000018), inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     cos(skoX3C1) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)), inference(subst, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 * (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <= skoX3C1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_89, plain,
    (~ lgen(0, X_000076, sin(X_000075)) | X_000076 <= sin(X_000075)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_90, plain,
    (0 < X_000075 | ~ lgen(0, X_000076, X_000075) |
     lgen(0, X_000076, sin(X_000075))),
    inference(subst, [], [sin_lower_bound_0_neg])).

cnf(refute_0_91, plain,
    (0 < X_000075 | ~ lgen(0, X_000076, X_000075) |
     X_000076 <= sin(X_000075)),
    inference(resolve, [], [refute_0_90, refute_0_89])).

cnf(refute_0_92, plain,
    (0 < X_000075 | X_000075 < X_000076 | X_000076 <= sin(X_000075)),
    inference(arithmetic, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (skoX3C1 <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | 0 < skoX3C1 |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <= sin(skoX3C1)),
    inference(subst, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (skoX3C1 <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | 0 < skoX3C1 | skoX1C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_93, refute_0_29])).

cnf(refute_0_95, plain,
    (skoX3C1 <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | 0 < skoX3C1),
    inference(resolve, [], [refute_0_94, refute_0_78])).

cnf(refute_0_96, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 * (skoX1C1 * (skoX1C1 * -1)) | 0 < skoX3C1 |
     0 <= skoX1C1 * (skoX1C1 * -1)),
    inference(resolve, [], [refute_0_88, refute_0_95])).

cnf(refute_0_97, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | 0 < skoX3C1 |
     skoX1C1 * skoX1C1 <= 0), inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | 0 < skoX3C1),
    inference(resolve, [], [refute_0_97, refute_0_78])).

cnf(refute_0_99, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) < cos(skoX3C1) | 0 < skoX3C1 |
     skoX2C1 * skoX1C1 <= 0), inference(split, [], [refute_0_98])).

cnf(refute_0_100, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_87, refute_0_99])).

cnf(refute_0_101, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | 0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_82, refute_0_100])).

cnf(refute_0_102, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24))))) < 0 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(split, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2) +
       skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) | 0 < skoX3C1 |
     skoX2C1 * skoX1C1 <= 0), inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * -1/2)) * (skoX2C1 * skoX1C1) |
     0 <= skoX2C1 * skoX1C1 |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <= 1 + skoX3C1 * (skoX3C1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_106, plain,
    (~ lgen(0, X_000011, cos(X_000010)) | X_000011 <= cos(X_000010)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_107, plain,
    (~ lgen(0, X_000011, 1 - X_000010 ^ 2 / 2) |
     lgen(0, X_000011, cos(X_000010))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_108, plain,
    (~ lgen(0, X_000011, 1 - X_000010 ^ 2 / 2) |
     X_000011 <= cos(X_000010)),
    inference(resolve, [], [refute_0_107, refute_0_106])).

cnf(refute_0_109, plain,
    (1 + X_000010 * (X_000010 * -1/2) < X_000011 |
     X_000011 <= cos(X_000010)),
    inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (1 + skoX3C1 * (skoX3C1 * -1/2) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <= cos(skoX3C1)),
    inference(subst, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | 0 < skoX3C1 | 0 <= skoX2C1 * skoX1C1),
    inference(split, [], [refute_0_98])).

cnf(refute_0_112, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | 0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (1 + skoX3C1 * (skoX3C1 * -1/2) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 * (skoX1C1 * skoX1C1) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | 0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_110, refute_0_112])).

cnf(refute_0_114, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * -1/2)) * (skoX2C1 * skoX1C1) | 0 < skoX3C1 |
     0 <= skoX2C1 * skoX1C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_105, refute_0_113])).

cnf(refute_0_115, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 + skoX3C1 * (skoX2C1 * (skoX1C1 * 1/2))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 + skoX3C1 * (skoX2C1 * (skoX1C1 * 1/2))) < 0 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(split, [], [refute_0_115])).

cnf(refute_0_117, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX3C1 * (skoX2C1 * (skoX1C1 * -1/2))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (skoX2C1 * skoX1C1 < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) < 0 |
     0 < skoX2C1 * skoX1C1 | 0 < skoX3C1),
    inference(split, [], [refute_0_98])).

cnf(refute_0_119, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 | 0 < skoX3C1),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 * (skoX1C1 * skoX1C1) < 0 | 0 < skoX2C1 * (skoX1C1 * -1) |
     0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1), inference(split, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 * (skoX1C1 * (skoX1C1 * -1)) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1), inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 |
     skoX3C1 * (skoX1C1 * (skoX1C1 * -1)) <=
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1))),
    inference(decision, [],
              [refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_123, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1), inference(resolve, [], [refute_0_122, refute_0_121])).

cnf(refute_0_124, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_125, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1), inference(resolve, [], [refute_0_124, refute_0_123])).

cnf(refute_0_126, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 | 0 < skoX3C1 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_127, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 | 0 < skoX3C1),
    inference(resolve, [], [refute_0_126, refute_0_125])).

cnf(refute_0_128, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0 |
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX3C1 * (skoX2C1 * (skoX1C1 * -1/2))) <=
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1))),
    inference(decision, [],
              [refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_129, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_128, refute_0_117])).

cnf(refute_0_130, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_131, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_130, refute_0_129])).

cnf(refute_0_132, plain,
    (0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_133, plain, (0 < skoX3C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_132, refute_0_131])).

cnf(refute_0_134, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2) +
       skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) <=
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1))),
    inference(decision, [],
              [refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_135, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_134, refute_0_104])).

cnf(refute_0_136, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_137, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_136, refute_0_135])).

cnf(refute_0_138, plain,
    (0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_139, plain, (0 < skoX3C1 | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_138, refute_0_137])).

cnf(refute_0_140, plain, (0 < skoX2C1 * skoX1C1 | 0 < skoX3C1),
    inference(decision, [],
              [refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_141, plain, (0 < skoX3C1),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (0 < skoX1C1 * skoX1C1 | skoX3C1 * (skoX3C1 * 868589/2000000) <= 0 |
     skoX3C1 * (868589/2000000 + skoX3C1 * -868589/2000000) <=
     skoX2C1 * skoX2C1 | skoX3C1 = 0),
    inference(decision, [],
              [refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_143, plain,
    (0 < skoX1C1 * skoX1C1 | skoX3C1 * (skoX3C1 * 868589/2000000) <= 0 |
     skoX3C1 = 0), inference(resolve, [], [refute_0_142, refute_0_81])).

cnf(refute_0_144, plain,
    (0 < skoX1C1 * skoX1C1 | 0 < skoX3C1 * (skoX3C1 * 868589/2000000) |
     skoX3C1 = 0),
    inference(decision, [],
              [refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_145, plain, (0 < skoX1C1 * skoX1C1 | skoX3C1 = 0),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain, (0 < skoX1C1 * skoX1C1 | skoX3C1 != 0),
    inference(decision, [],
              [refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_147, plain, (0 < skoX1C1 * skoX1C1),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain, (skoX3C1 < 0 | 0 < skoX1C1 * skoX1C1),
    inference(decision, [],
              [refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_149, plain,
    (skoX3C1 < 0 |
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1))),
    inference(resolve, [], [refute_0_35, refute_0_148])).

cnf(refute_0_150, plain, (0 <= skoX3C1),
    inference(decision, [],
              [refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_151, plain,
    (skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1))),
    inference(resolve, [], [refute_0_150, refute_0_149])).

cnf(refute_0_152, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 *
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/6 +
        skoX3C1 *
        (skoX3C1 *
         (1/120 +
          skoX3C1 *
          (skoX3C1 *
           (-1/5040 +
            skoX3C1 *
            (skoX3C1 *
             (1/362880 + skoX3C1 * (skoX3C1 * -1/39916800)))))))))) *
     (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1)),
    inference(resolve, [], [refute_0_13, refute_0_151])).

cnf(refute_0_153, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | skoX1C1 * skoX1C1 <= 0),
    inference(arithmetic, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1)),
    inference(resolve, [], [refute_0_153, refute_0_147])).

cnf(refute_0_155, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | 0 <= skoX2C1 * skoX1C1),
    inference(split, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_157, plain,
    (1 +
     skoX3C1 *
     (skoX3C1 *
      (-1/2 +
       skoX3C1 *
       (skoX3C1 *
        (1/24 +
         skoX3C1 *
         (skoX3C1 *
          (-1/720 +
           skoX3C1 *
           (skoX3C1 * (1/40320 + skoX3C1 * (skoX3C1 * -1/3628800))))))))) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_156])).

cnf(refute_0_158, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 +
      skoX3C1 *
      (skoX3C1 *
       (-1/2 +
        skoX3C1 *
        (skoX3C1 *
         (1/24 +
          skoX3C1 *
          (skoX3C1 *
           (-1/720 +
            skoX3C1 *
            (skoX3C1 *
             (1/40320 + skoX3C1 * (skoX3C1 * -1/3628800)))))))))) *
     (skoX2C1 * skoX1C1) | 0 <= skoX2C1 * skoX1C1 |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_7, refute_0_157])).

cnf(refute_0_159, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_158])).

cnf(refute_0_160, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     (-1 + skoX3C1) / skoX3C1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_161, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     ln(skoX3C1) | skoX3C1 <= 0), inference(subst, [], [refute_0_73])).

cnf(refute_0_162, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | skoX2C1 * skoX1C1 <= 0 |
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_163, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     cos(skoX3C1) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)), inference(subst, [], [refute_0_86])).

cnf(refute_0_164, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) < cos(skoX3C1) | skoX2C1 * skoX1C1 <= 0),
    inference(split, [], [refute_0_154])).

cnf(refute_0_165, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_163, refute_0_164])).

cnf(refute_0_166, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_162, refute_0_165])).

cnf(refute_0_167, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0), inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0), inference(split, [], [refute_0_167])).

cnf(refute_0_169, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) *
     (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <=
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_171, plain,
    (~ lgen(0, X_000104, sin(X_000103)) | X_000104 <= sin(X_000103)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_172, plain,
    (X_000103 < 0 | ~ lgen(0, X_000104, X_000103 - X_000103 ^ 3 / 6) |
     lgen(0, X_000104, sin(X_000103))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_173, plain,
    (X_000103 < 0 | ~ lgen(0, X_000104, X_000103 - X_000103 ^ 3 / 6) |
     X_000104 <= sin(X_000103)),
    inference(resolve, [], [refute_0_172, refute_0_171])).

cnf(refute_0_174, plain,
    (X_000103 * (1 + X_000103 * (X_000103 * -1/6)) < X_000104 |
     X_000103 < 0 | X_000104 <= sin(X_000103)),
    inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (skoX3C1 < 0 |
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) <= sin(skoX3C1)),
    inference(subst, [], [refute_0_174])).

cnf(refute_0_176, plain,
    (skoX3C1 < 0 |
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1)) | skoX1C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_175, refute_0_29])).

cnf(refute_0_177, plain,
    (skoX3C1 < 0 |
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1))),
    inference(resolve, [], [refute_0_176, refute_0_148])).

cnf(refute_0_178, plain,
    (skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
      cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1))) /
     (skoX1C1 * (skoX1C1 * -1))),
    inference(resolve, [], [refute_0_150, refute_0_177])).

cnf(refute_0_179, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) +
     cos(skoX3C1) * (skoX2C1 * (skoX1C1 * -1)) <
     skoX3C1 * (1 + skoX3C1 * (skoX3C1 * -1/6)) *
     (skoX1C1 * (skoX1C1 * -1)) | 0 <= skoX1C1 * (skoX1C1 * -1)),
    inference(resolve, [], [refute_0_170, refute_0_178])).

cnf(refute_0_180, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1) | skoX1C1 * skoX1C1 <= 0),
    inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     cos(skoX3C1) * (skoX2C1 * skoX1C1)),
    inference(resolve, [], [refute_0_180, refute_0_147])).

cnf(refute_0_182, plain,
    (skoX2C1 * skoX1C1 < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) < 0 |
     0 < skoX2C1 * skoX1C1), inference(split, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1),
    inference(arithmetic, [], [refute_0_182])).

cnf(refute_0_184, plain,
    (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) < 0 |
     0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(split, [], [refute_0_183])).

cnf(refute_0_185, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6))))),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_187, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6))))),
    inference(resolve, [], [refute_0_186, refute_0_185])).

cnf(refute_0_188, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_189, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6))))),
    inference(resolve, [], [refute_0_188, refute_0_187])).

cnf(refute_0_190, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(decision, [],
              [refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_191, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_169, refute_0_190])).

cnf(refute_0_192, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0 | skoX3C1 <= 0),
    inference(resolve, [], [refute_0_161, refute_0_191])).

cnf(refute_0_193, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_192, refute_0_139])).

cnf(refute_0_194, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_160, refute_0_193])).

cnf(refute_0_195, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * 1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800)))))))))))) |
     skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 = 0), inference(arithmetic, [], [refute_0_194])).

cnf(refute_0_196, plain,
    (skoX2C1 * skoX1C1 < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) < 0 |
     0 < skoX2C1 * skoX1C1), inference(split, [], [refute_0_154])).

cnf(refute_0_197, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1),
    inference(arithmetic, [], [refute_0_196])).

cnf(refute_0_198, plain,
    (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) < 0 |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     0 | 0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(split, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800))))))))))))),
    inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     (-1 + skoX3C1) / skoX3C1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_201, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     ln(skoX3C1) | skoX3C1 <= 0), inference(subst, [], [refute_0_73])).

cnf(refute_0_202, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * -1/2)) * (skoX2C1 * skoX1C1) |
     0 <= skoX2C1 * skoX1C1 |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <= 1 + skoX3C1 * (skoX3C1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_203, plain,
    (1 + skoX3C1 * (skoX3C1 * -1/2) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) |
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <= cos(skoX3C1)),
    inference(subst, [], [refute_0_109])).

cnf(refute_0_204, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | 0 <= skoX2C1 * skoX1C1),
    inference(split, [], [refute_0_181])).

cnf(refute_0_205, plain,
    (cos(skoX3C1) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (1 + skoX3C1 * (skoX3C1 * -1/2) <
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_203, refute_0_205])).

cnf(refute_0_207, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * -1/2)) * (skoX2C1 * skoX1C1) |
     0 <= skoX2C1 * skoX1C1 | skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_202, refute_0_206])).

cnf(refute_0_208, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1) <= 0), inference(split, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24))))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     (-1 + skoX3C1) / skoX3C1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_212, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) <=
     ln(skoX3C1) | skoX3C1 <= 0), inference(subst, [], [refute_0_73])).

cnf(refute_0_213, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | skoX2C1 * skoX1C1 <= 0 |
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_214, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     cos(skoX3C1) <=
     (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1)), inference(subst, [], [refute_0_86])).

cnf(refute_0_215, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) < cos(skoX3C1) | skoX2C1 * skoX1C1 <= 0),
    inference(split, [], [refute_0_181])).

cnf(refute_0_216, plain,
    ((skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
      ln(skoX3C1) *
      (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))))) /
     (skoX2C1 * skoX1C1) <
     1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 * (skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) +
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))) <
     (1 + skoX3C1 * (skoX3C1 * (-1/2 + skoX3C1 * (skoX3C1 * 1/24)))) *
     (skoX2C1 * skoX1C1) | skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_213, refute_0_216])).

cnf(refute_0_218, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0), inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0), inference(split, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(arithmetic, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_220, refute_0_190])).

cnf(refute_0_222, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0 | skoX3C1 <= 0),
    inference(resolve, [], [refute_0_212, refute_0_221])).

cnf(refute_0_223, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24)))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_222, refute_0_139])).

cnf(refute_0_224, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * -1/24))))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_211, refute_0_223])).

cnf(refute_0_225, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) |
     skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 = 0), inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) |
     skoX2C1 * skoX1C1 <= 0 | skoX3C1 = 0),
    inference(decision, [],
              [refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_227, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) |
     skoX2C1 * skoX1C1 <= 0 | skoX3C1 = 0),
    inference(resolve, [], [refute_0_225, refute_0_226])).

cnf(refute_0_228, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) | skoX3C1 != 0 |
     skoX2C1 * skoX1C1 <= 0),
    inference(decision, [],
              [refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_229, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 * (skoX2C1 * (skoX1C1 * 1/24))))) |
     skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_227, refute_0_228])).

cnf(refute_0_230, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(decision, [],
              [refute_0_229, refute_0_189, refute_0_147, refute_0_141,
               refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_231, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_210, refute_0_230])).

cnf(refute_0_232, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0 | skoX3C1 <= 0),
    inference(resolve, [], [refute_0_201, refute_0_231])).

cnf(refute_0_233, plain,
    ((-1 + skoX3C1) / skoX3C1 <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_232, refute_0_133])).

cnf(refute_0_234, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * -1/6)))) <
     (-1 + skoX3C1) / skoX3C1 *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_200, refute_0_233])).

cnf(refute_0_235, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 = 0), inference(arithmetic, [], [refute_0_234])).

cnf(refute_0_236, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0 | skoX3C1 = 0),
    inference(decision, [],
              [refute_0_229, refute_0_189, refute_0_147, refute_0_141,
               refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_237, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0 | skoX3C1 = 0),
    inference(resolve, [], [refute_0_235, refute_0_236])).

cnf(refute_0_238, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) | skoX3C1 != 0 |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(decision, [],
              [refute_0_229, refute_0_189, refute_0_147, refute_0_141,
               refute_0_133, refute_0_127, refute_0_77, refute_0_61,
               refute_0_39, refute_0_40])).

cnf(refute_0_239, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 * (skoX1C1 * (skoX1C1 * 1/6)))) |
     skoX2C1 * (skoX1C1 * -1) <= 0),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) +
      skoX3C1 *
      (skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 *
        (skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800)))))))))))) <=
     skoX2C1 * (skoX2C1 * (skoX2C1 * skoX2C1))),
    inference(decision, [],
              [refute_0_239, refute_0_229, refute_0_189, refute_0_147,
               refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_241, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(resolve, [], [refute_0_240, refute_0_199])).

cnf(refute_0_242, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_239, refute_0_229, refute_0_189, refute_0_147,
               refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_243, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(resolve, [], [refute_0_242, refute_0_241])).

cnf(refute_0_244, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(decision, [],
              [refute_0_239, refute_0_229, refute_0_189, refute_0_147,
               refute_0_141, refute_0_133, refute_0_127, refute_0_77,
               refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_245, plain,
    (0 < skoX2C1 * (skoX1C1 * -1) | 0 < skoX2C1 * skoX1C1),
    inference(resolve, [], [refute_0_244, refute_0_243])).

cnf(refute_0_246, plain,
    (skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -1) + skoX2C1 * (skoX2C1 * 868589/2000000) +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/2000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * 1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/120) +
          skoX3C1 *
          (skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/5040) +
            skoX3C1 *
            (skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/362880) +
              skoX3C1 *
              (skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800)))))))))))) <=
     skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) |
     skoX3C1 = 0),
    inference(decision, [],
              [refute_0_245, refute_0_239, refute_0_229, refute_0_189,
               refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_247, plain,
    (skoX2C1 * skoX1C1 <= 0 |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 = 0), inference(resolve, [], [refute_0_246, refute_0_195])).

cnf(refute_0_248, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX2C1 * skoX1C1 <= 0 | skoX3C1 = 0),
    inference(decision, [],
              [refute_0_245, refute_0_239, refute_0_229, refute_0_189,
               refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_249, plain, (skoX2C1 * skoX1C1 <= 0 | skoX3C1 = 0),
    inference(resolve, [], [refute_0_247, refute_0_248])).

cnf(refute_0_250, plain, (skoX3C1 != 0 | skoX2C1 * skoX1C1 <= 0),
    inference(decision, [],
              [refute_0_245, refute_0_239, refute_0_229, refute_0_189,
               refute_0_147, refute_0_141, refute_0_133, refute_0_127,
               refute_0_77, refute_0_61, refute_0_39, refute_0_40])).

cnf(refute_0_251, plain, (skoX2C1 * skoX1C1 <= 0),
    inference(resolve, [], [refute_0_249, refute_0_250])).

cnf(refute_0_252, plain, (0 < skoX2C1 * (skoX1C1 * -1)),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_253, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     ln(skoX3C1) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))))),
    inference(resolve, [], [refute_0_159, refute_0_252])).

cnf(refute_0_254, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(split, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0),
    inference(arithmetic, [], [refute_0_254])).

cnf(refute_0_256, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000)))),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_257, plain,
    (ln(skoX3C1) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))))),
    inference(resolve, [], [refute_0_255, refute_0_256])).

cnf(refute_0_258, plain,
    ((-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1)) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     skoX3C1 <= 0), inference(resolve, [], [refute_0_6, refute_0_257])).

cnf(refute_0_259, plain,
    ((-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1)) <
     (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) /
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000))))),
    inference(resolve, [], [refute_0_258, refute_0_141])).

cnf(refute_0_260, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     (-1/2 + skoX3C1 * (-2 + skoX3C1 * 5/2)) / (skoX3C1 * (2 + skoX3C1)) *
     (skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))) |
     0 <= skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/2000000)))),
    inference(resolve, [], [refute_0_1, refute_0_259])).

cnf(refute_0_261, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) / (2 + skoX3C1) |
     skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) <= 0 |
     skoX3C1 = 0), inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (0 < skoX3C1 * (skoX3C1 * (skoX2C1 * (skoX2C1 * 868589/2000000))) |
     skoX3C1 = 0),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_263, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) / (2 + skoX3C1) |
     skoX3C1 = 0), inference(resolve, [], [refute_0_261, refute_0_262])).

cnf(refute_0_264, plain, (skoX3C1 != 0),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_265, plain,
    (skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
     skoX3C1 *
     (skoX1C1 * skoX1C1 +
      skoX3C1 *
      (skoX2C1 * (skoX1C1 * 1/2) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * -1/6) +
        skoX3C1 *
        (skoX2C1 * (skoX1C1 * -1/24) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * 1/120) +
          skoX3C1 *
          (skoX2C1 * (skoX1C1 * 1/720) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * -1/5040) +
            skoX3C1 *
            (skoX2C1 * (skoX1C1 * -1/40320) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * 1/362880) +
              skoX3C1 *
              (skoX2C1 * (skoX1C1 * 1/3628800) +
               skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800)))))))))))) <
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) / (2 + skoX3C1)),
    inference(resolve, [], [refute_0_263, refute_0_264])).

cnf(refute_0_266, plain,
    ((skoX2C1 * (skoX1C1 * -1 + skoX2C1 * (skoX2C1 * skoX2C1)) +
      skoX3C1 *
      (skoX1C1 * skoX1C1 +
       skoX3C1 *
       (skoX2C1 * (skoX1C1 * 1/2) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * -1/6) +
         skoX3C1 *
         (skoX2C1 * (skoX1C1 * -1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * 1/120) +
           skoX3C1 *
           (skoX2C1 * (skoX1C1 * 1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * -1/5040) +
             skoX3C1 *
             (skoX2C1 * (skoX1C1 * -1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * 1/362880) +
               skoX3C1 *
               (skoX2C1 * (skoX1C1 * 1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * -1/39916800))))))))))))) *
     (2 + skoX3C1) <
     skoX3C1 *
     (skoX2C1 * (skoX2C1 * 868589/4000000) +
      skoX3C1 *
      (skoX2C1 * (skoX2C1 * 868589/1000000) +
       skoX3C1 * (skoX2C1 * (skoX2C1 * -868589/800000)))) |
     2 + skoX3C1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_265])).

cnf(refute_0_267, plain,
    (skoX2C1 * (skoX1C1 * -2 + skoX2C1 * (skoX2C1 * (skoX2C1 * 2))) <
     skoX3C1 *
     (skoX1C1 * (skoX1C1 * -2) +
      skoX2C1 *
      (skoX1C1 + skoX2C1 * (868589/4000000 + skoX2C1 * (skoX2C1 * -1))) +
      skoX3C1 *
      (skoX1C1 * (skoX1C1 * -1) +
       skoX2C1 * (skoX1C1 * -1 + skoX2C1 * 868589/1000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/3) +
        skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/800000) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * 1/6) + skoX2C1 * (skoX1C1 * 1/12) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/60) + skoX2C1 * (skoX1C1 * 1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * -1/120) + skoX2C1 * (skoX1C1 * -1/360) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/2520) + skoX2C1 * (skoX1C1 * -1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * 1/5040) + skoX2C1 * (skoX1C1 * 1/20160) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/181440) +
              skoX2C1 * (skoX1C1 * 1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * -1/362880) +
               skoX2C1 * (skoX1C1 * -1/1814400) +
               skoX3C1 *
               (skoX1C1 * (skoX1C1 * 1/19958400) +
                skoX2C1 * (skoX1C1 * -1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800))))))))))))) |
     skoX3C1 <= -2), inference(arithmetic, [], [refute_0_266])).

cnf(refute_0_268, plain,
    (skoX3C1 *
     (skoX1C1 * (skoX1C1 * -2) +
      skoX2C1 *
      (skoX1C1 + skoX2C1 * (868589/4000000 + skoX2C1 * (skoX2C1 * -1))) +
      skoX3C1 *
      (skoX1C1 * (skoX1C1 * -1) +
       skoX2C1 * (skoX1C1 * -1 + skoX2C1 * 868589/1000000) +
       skoX3C1 *
       (skoX1C1 * (skoX1C1 * 1/3) +
        skoX2C1 * (skoX1C1 * -1/2 + skoX2C1 * -868589/800000) +
        skoX3C1 *
        (skoX1C1 * (skoX1C1 * 1/6) + skoX2C1 * (skoX1C1 * 1/12) +
         skoX3C1 *
         (skoX1C1 * (skoX1C1 * -1/60) + skoX2C1 * (skoX1C1 * 1/24) +
          skoX3C1 *
          (skoX1C1 * (skoX1C1 * -1/120) + skoX2C1 * (skoX1C1 * -1/360) +
           skoX3C1 *
           (skoX1C1 * (skoX1C1 * 1/2520) + skoX2C1 * (skoX1C1 * -1/720) +
            skoX3C1 *
            (skoX1C1 * (skoX1C1 * 1/5040) + skoX2C1 * (skoX1C1 * 1/20160) +
             skoX3C1 *
             (skoX1C1 * (skoX1C1 * -1/181440) +
              skoX2C1 * (skoX1C1 * 1/40320) +
              skoX3C1 *
              (skoX1C1 * (skoX1C1 * -1/362880) +
               skoX2C1 * (skoX1C1 * -1/1814400) +
               skoX3C1 *
               (skoX1C1 * (skoX1C1 * 1/19958400) +
                skoX2C1 * (skoX1C1 * -1/3628800) +
                skoX3C1 * (skoX1C1 * (skoX1C1 * 1/39916800))))))))))))) <=
     skoX2C1 * (skoX1C1 * -2 + skoX2C1 * (skoX2C1 * (skoX2C1 * 2))) |
     skoX3C1 <= -2),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_269, plain, (skoX3C1 <= -2),
    inference(resolve, [], [refute_0_268, refute_0_267])).

cnf(refute_0_270, plain, (-2 < skoX3C1),
    inference(decision, [],
              [refute_0_251, refute_0_245, refute_0_239, refute_0_229,
               refute_0_189, refute_0_147, refute_0_141, refute_0_133,
               refute_0_127, refute_0_77, refute_0_61, refute_0_39,
               refute_0_40])).

cnf(refute_0_271, plain, ($false),
    inference(resolve, [], [refute_0_269, refute_0_270])).
% SZS output end CNFRefutation for Mastellone2c.tptp
