% SZS output start CNFRefutation for MUTANT-3-h.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf1, axiom,
    (X < 0 | 3/2 < X | ~ lgen(R, -(X + 2) / (X - 2), Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

cnf(sinh_upper_bound_pos1, axiom,
    (X < 0 | 3/2 < X | ~ lgen(R, -4 * X / ((X - 2) * (X + 2)), Y) |
     lgen(R, sinh(X), Y))).

cnf(cosh_upper_bound_pos1, axiom,
    (X < 0 | 3/2 < X |
     ~ lgen(R, -(8 * (X ^ 2 + 3)) / ((X - 2) * (X ^ 2 + 6 * X + 12)), Y) |
     lgen(R, cosh(X), Y))).

fof('MUTANT_3_h', conjecture,
    (! [X] :
       ((0 <= X & X <= 10) =>
        -2650 * exp(-139/500 + 10 ^ -3 * X) +
        19/12500000 * exp(-103/20000 * X) *
        (5020000000 * cosh(16/3125 * X) +
         24000000000 * sinh(16/3125 * X)) < 25000))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 10) =>
        -2650 * exp(-139/500 + 10 ^ -3 * X) +
        19/12500000 * exp(-103/20000 * X) *
        (5020000000 * cosh(16/3125 * X) +
         24000000000 * sinh(16/3125 * X)) < 25000)),
    inference(strip, [], ['MUTANT_3_h'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 10) =>
          -2650 * exp(-139/500 + 10 ^ -3 * X) +
          19/12500000 * exp(-103/20000 * X) *
          (5020000000 * cosh(16/3125 * X) +
           24000000000 * sinh(16/3125 * X)) < 25000)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 10 & 0 <= X &
        25000 <=
        -2650 * exp(-139/500 + 10 ^ -3 * X) +
        19/12500000 * exp(-103/20000 * X) *
        (5020000000 * cosh(16/3125 * X) +
         24000000000 * sinh(16/3125 * X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     25000 <=
     -2650 * exp(-139/500 + 10 ^ -3 * skoXC1) +
     19/12500000 * exp(-103/20000 * skoXC1) *
     (5020000000 * cosh(16/3125 * skoXC1) +
      24000000000 * sinh(16/3125 * skoXC1)) & skoXC1 <= 10),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (25000 <=
     -2650 * exp(-139/500 + 10 ^ -3 * skoXC1) +
     19/12500000 * exp(-103/20000 * skoXC1) *
     (5020000000 * cosh(16/3125 * skoXC1) +
      24000000000 * sinh(16/3125 * skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 10),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-21975552/6625 +
     skoXC1 *
     (35366904/4140625 +
      skoXC1 *
      (-7545068283/517578125000 +
       skoXC1 *
       (754493952/40435791015625 +
        skoXC1 *
        (109896950456/631809234619140625 +
         skoXC1 *
         (-193150451712/394880771636962890625 +
          skoXC1 * 103617169408/246800482273101806640625))))) <
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) *
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 <=
     -24 + skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125)) |
     (-21975552/6625 +
      skoXC1 *
      (35366904/4140625 +
       skoXC1 *
       (-7545068283/517578125000 +
        skoXC1 *
        (754493952/40435791015625 +
         skoXC1 *
         (109896950456/631809234619140625 +
          skoXC1 *
          (-193150451712/394880771636962890625 +
           skoXC1 * 103617169408/246800482273101806640625)))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) <=
     24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) *
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) |
     915648/6625 +
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 +
        skoXC1 * -50594321/25272369384765625))) <= 0 |
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) <=
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_2, plain,
    ((-12 + skoXC1 * (309/10000 + skoXC1 * -10609/400000000)) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     -24000/53 +
     skoXC1 *
     (-309/265 +
      skoXC1 *
      (521207/265000000 +
       skoXC1 * (19776/2587890625 + skoXC1 * 10609/1617431640625))) |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <= 0 |
     (-24000/53 +
      skoXC1 *
      (-309/265 +
       skoXC1 *
       (521207/265000000 +
        skoXC1 * (19776/2587890625 + skoXC1 * 10609/1617431640625)))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <=
     -12 + skoXC1 * (309/10000 + skoXC1 * -10609/400000000)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_3, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) *
     (12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000)) <
     12 + skoXC1 * (-309/10000 + skoXC1 * 10609/400000000) |
     12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000) <= 0 |
     (12 + skoXC1 * (-309/10000 + skoXC1 * 10609/400000000)) /
     (12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000)) <=
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_4, plain,
    (~ lgen(0, exp(X_000071), X_000072) | exp(X_000071) <= X_000072),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_5, plain,
    (~ lgen(0,
            (X_000071 ^ 2 + 6 * X_000071 + 12) /
            (X_000071 ^ 2 - 6 * X_000071 + 12), X_000072) | 0 <= X_000071 |
     lgen(0, exp(X_000071), X_000072)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_0_6, plain,
    (~ lgen(0,
            (X_000071 ^ 2 + 6 * X_000071 + 12) /
            (X_000071 ^ 2 - 6 * X_000071 + 12), X_000072) | 0 <= X_000071 |
     exp(X_000071) <= X_000072),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (X_000072 <
     (12 + X_000071 * (6 + X_000071)) / (12 + X_000071 * (-6 + X_000071)) |
     0 <= X_000071 | exp(X_000071) <= X_000072),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (12 + skoXC1 * -103/20000 * (6 + skoXC1 * -103/20000)) /
     (12 + skoXC1 * -103/20000 * (-6 + skoXC1 * -103/20000)) |
     0 <= skoXC1 * -103/20000 |
     exp(skoXC1 * -103/20000) <=
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)))),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (exp(skoXC1 * -103/20000) * (skoXC1 * -233472/828125) <
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) *
     (-4 + skoXC1 * (skoXC1 * 256/9765625)) |
     0 <= -4 + skoXC1 * (skoXC1 * 256/9765625) |
     exp(skoXC1 * -103/20000) * (skoXC1 * -233472/828125) /
     (-4 + skoXC1 * (skoXC1 * 256/9765625)) <=
     500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_10, plain,
    (500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625) <
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625)) *
     (exp(skoXC1 * -103/20000) * 3648/265) |
     exp(skoXC1 * -103/20000) * 3648/265 <= 0 |
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625)) <=
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_11, plain,
    (25000 <=
     -2650 * exp(-139/500 + 10 ^ -3 * skoXC1) +
     19/12500000 * exp(-103/20000 * skoXC1) *
     (5020000000 * cosh(16/3125 * skoXC1) +
      24000000000 * sinh(16/3125 * skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_12, plain,
    (exp(-139/500 + skoXC1 * 1/1000) <=
     -500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * 19076/6625) +
     sinh(skoXC1 * 16/3125) * (exp(skoXC1 * -103/20000) * 3648/265)),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_15, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (0 <
     -500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * 19076/6625) +
     sinh(skoXC1 * 16/3125) * (exp(skoXC1 * -103/20000) * 3648/265) |
     -500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * 19076/6625) +
     sinh(skoXC1 * 16/3125) * (exp(skoXC1 * -103/20000) * 3648/265) <
     exp(-139/500 + skoXC1 * 1/1000)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (0 <
     -500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * 19076/6625) +
     sinh(skoXC1 * 16/3125) * (exp(skoXC1 * -103/20000) * 3648/265)),
    inference(resolve, [], [refute_0_12, refute_0_16])).

cnf(refute_0_18, plain,
    (500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625) <
     sinh(skoXC1 * 16/3125) * (exp(skoXC1 * -103/20000) * 3648/265)),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    ((500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) < sinh(skoXC1 * 16/3125) |
     exp(skoXC1 * -103/20000) * 3648/265 <= 0),
    inference(split, [], [refute_0_18])).

cnf(refute_0_20, plain,
    ((500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) < sinh(skoXC1 * 16/3125)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (~ lgen(0, sinh(X_000110), X_000111) | sinh(X_000110) <= X_000111),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_22, plain,
    (3/2 < X_000110 | X_000110 < 0 |
     ~ lgen(0, -4 * X_000110 / ((X_000110 - 2) * (X_000110 + 2)),
            X_000111) | lgen(0, sinh(X_000110), X_000111)),
    inference(subst, [], [sinh_upper_bound_pos1])).

cnf(refute_0_23, plain,
    (3/2 < X_000110 | X_000110 < 0 |
     ~ lgen(0, -4 * X_000110 / ((X_000110 - 2) * (X_000110 + 2)),
            X_000111) | sinh(X_000110) <= X_000111),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (3/2 < X_000110 | X_000110 < 0 |
     X_000111 < X_000110 * -4 / (-4 + X_000110 * X_000110) |
     sinh(X_000110) <= X_000111),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) <
     skoXC1 * 16/3125 * -4 / (-4 + skoXC1 * 16/3125 * (skoXC1 * 16/3125)) |
     sinh(skoXC1 * 16/3125) <=
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265)),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) <
     skoXC1 * 16/3125 * -4 / (-4 + skoXC1 * 16/3125 * (skoXC1 * 16/3125))),
    inference(resolve, [], [refute_0_25, refute_0_20])).

cnf(refute_0_27, plain,
    (skoXC1 < 0 |
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) <
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625)) |
     9375/32 < skoXC1), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain, (skoXC1 <= 10),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_29, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_30, plain, (9375/32 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_31, plain,
    ((500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) <
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625)) |
     9375/32 < skoXC1),
    inference(resolve, [], [refute_0_30, refute_0_27])).

cnf(refute_0_32, plain, (skoXC1 <= 9375/32),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_33, plain,
    ((500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) /
     (exp(skoXC1 * -103/20000) * 3648/265) <
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625))),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625) <
     skoXC1 * -64/3125 / (-4 + skoXC1 * (skoXC1 * 256/9765625)) *
     (exp(skoXC1 * -103/20000) * 3648/265) |
     exp(skoXC1 * -103/20000) * 3648/265 <= 0),
    inference(resolve, [], [refute_0_10, refute_0_33])).

cnf(refute_0_35, plain,
    (500/53 +
     exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625) <
     exp(skoXC1 * -103/20000) * (skoXC1 * -233472/828125) /
     (-4 + skoXC1 * (skoXC1 * 256/9765625))),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (exp(skoXC1 * -103/20000) * (skoXC1 * -233472/828125) <
     (500/53 +
      exp(skoXC1 * -103/20000) * (cosh(skoXC1 * 16/3125) * -19076/6625)) *
     (-4 + skoXC1 * (skoXC1 * 256/9765625)) |
     0 <= -4 + skoXC1 * (skoXC1 * 256/9765625)),
    inference(resolve, [], [refute_0_9, refute_0_35])).

cnf(refute_0_37, plain,
    (2000/53 + skoXC1 * (skoXC1 * -1024/4140625) <
     exp(skoXC1 * -103/20000) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     skoXC1 * (skoXC1 * -256/9765625) <= -4),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain, (-4 < skoXC1 * (skoXC1 * -256/9765625)),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_39, plain,
    (2000/53 + skoXC1 * (skoXC1 * -1024/4140625) <
     exp(skoXC1 * -103/20000) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)))),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     exp(skoXC1 * -103/20000) |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <= 0),
    inference(split, [], [refute_0_39])).

cnf(refute_0_41, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     exp(skoXC1 * -103/20000) |
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <=
     skoXC1 * -233472/828125), inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (2000/53 + skoXC1 * (skoXC1 * -1024/4140625) < 0 |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) < 0 |
     0 <
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))),
    inference(split, [], [refute_0_39])).

cnf(refute_0_43, plain,
    (skoXC1 * 233472/828125 <
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) |
     skoXC1 * -233472/828125 <
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) |
     2000/53 < skoXC1 * (skoXC1 * 1024/4140625)),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain, (skoXC1 * (skoXC1 * 1024/4140625) <= 2000/53),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_45, plain,
    (skoXC1 * 233472/828125 <
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) |
     skoXC1 * -233472/828125 <
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (2000/53 + skoXC1 * (skoXC1 * -1024/4140625) <
     0 *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     0 <=
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) |
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_47, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_48, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_49, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_48, refute_0_47])).

cnf(refute_0_50, plain,
    (0 <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <=
     exp(skoXC1 * -103/20000)), inference(subst, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (exp(skoXC1 * -103/20000) <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     0 <=
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))),
    inference(split, [], [refute_0_39])).

cnf(refute_0_52, plain,
    (exp(skoXC1 * -103/20000) <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) <=
     skoXC1 * 233472/828125), inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (0 <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) <=
     skoXC1 * 233472/828125),
    inference(resolve, [], [refute_0_50, refute_0_52])).

cnf(refute_0_54, plain,
    (2000/53 + skoXC1 * (skoXC1 * -1024/4140625) <
     0 *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     0 <=
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) |
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) <=
     skoXC1 * 233472/828125),
    inference(resolve, [], [refute_0_46, refute_0_53])).

cnf(refute_0_55, plain,
    (2000/53 < skoXC1 * (skoXC1 * 1024/4140625) |
     cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) <=
     skoXC1 * 233472/828125), inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (cosh(skoXC1 * 16/3125) *
     (-76304/6625 + skoXC1 * (skoXC1 * 4883456/64697265625)) <=
     skoXC1 * 233472/828125),
    inference(resolve, [], [refute_0_44, refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 * -233472/828125 <
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))),
    inference(resolve, [], [refute_0_56, refute_0_45])).

cnf(refute_0_58, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     exp(skoXC1 * -103/20000)),
    inference(resolve, [], [refute_0_41, refute_0_57])).

cnf(refute_0_59, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (12 + skoXC1 * -103/20000 * (6 + skoXC1 * -103/20000)) /
     (12 + skoXC1 * -103/20000 * (-6 + skoXC1 * -103/20000)) |
     0 <= skoXC1 * -103/20000),
    inference(resolve, [], [refute_0_8, refute_0_58])).

cnf(refute_0_60, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (12 + skoXC1 * (-309/10000 + skoXC1 * 10609/400000000)) /
     (12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) *
     (12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000)) <
     12 + skoXC1 * (-309/10000 + skoXC1 * 10609/400000000) |
     12 + skoXC1 * (309/10000 + skoXC1 * 10609/400000000) <= 0 |
     skoXC1 <= 0), inference(resolve, [], [refute_0_3, refute_0_60])).

cnf(refute_0_62, plain,
    (-12 + skoXC1 * (309/10000 + skoXC1 * -10609/400000000) <
     (-24000/53 +
      skoXC1 *
      (-309/265 +
       skoXC1 *
       (521207/265000000 +
        skoXC1 * (19776/2587890625 + skoXC1 * 10609/1617431640625)))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     skoXC1 * (309/10000 + skoXC1 * 10609/400000000) <= -12 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (-12 < skoXC1 * (309/10000 + skoXC1 * 10609/400000000) | skoXC1 <= 0),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_64, plain,
    (-12 + skoXC1 * (309/10000 + skoXC1 * -10609/400000000) <
     (-24000/53 +
      skoXC1 *
      (-309/265 +
       skoXC1 *
       (521207/265000000 +
        skoXC1 * (19776/2587890625 + skoXC1 * 10609/1617431640625)))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain,
    ((-12 + skoXC1 * (309/10000 + skoXC1 * -10609/400000000)) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     -24000/53 +
     skoXC1 *
     (-309/265 +
      skoXC1 *
      (521207/265000000 +
       skoXC1 * (19776/2587890625 + skoXC1 * 10609/1617431640625))) |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <= 0 |
     skoXC1 <= 0), inference(resolve, [], [refute_0_2, refute_0_64])).

cnf(refute_0_66, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     cosh(skoXC1 * 16/3125) *
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) |
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <=
     skoXC1 * -233472/828125 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     cosh(skoXC1 * 16/3125) *
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_66, refute_0_57])).

cnf(refute_0_68, plain,
    ((24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     cosh(skoXC1 * 16/3125) |
     915648/6625 +
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 +
        skoXC1 * -50594321/25272369384765625))) <= 0 | skoXC1 <= 0),
    inference(split, [], [refute_0_67])).

cnf(refute_0_69, plain,
    ((24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     cosh(skoXC1 * 16/3125) |
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 +
        skoXC1 * -50594321/25272369384765625))) <= -915648/6625 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (-915648/6625 <
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 + skoXC1 * -50594321/25272369384765625))) |
     skoXC1 <= 0), inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_71, plain,
    ((24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     cosh(skoXC1 * 16/3125) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (~ lgen(0, cosh(X_000117), X_000118) | cosh(X_000117) <= X_000118),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_73, plain,
    (3/2 < X_000117 | X_000117 < 0 |
     ~ lgen(0,
            -(8 * (X_000117 ^ 2 + 3)) /
            ((X_000117 - 2) * (X_000117 ^ 2 + 6 * X_000117 + 12)),
            X_000118) | lgen(0, cosh(X_000117), X_000118)),
    inference(subst, [], [cosh_upper_bound_pos1])).

cnf(refute_0_74, plain,
    (3/2 < X_000117 | X_000117 < 0 |
     ~ lgen(0,
            -(8 * (X_000117 ^ 2 + 3)) /
            ((X_000117 - 2) * (X_000117 ^ 2 + 6 * X_000117 + 12)),
            X_000118) | cosh(X_000117) <= X_000118),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    (3/2 < X_000117 | X_000117 < 0 |
     X_000118 <
     (-24 + X_000117 * (X_000117 * -8)) /
     (-24 + X_000117 * (X_000117 * (4 + X_000117))) |
     cosh(X_000117) <= X_000118),
    inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     (-24 + skoXC1 * 16/3125 * (skoXC1 * 16/3125 * -8)) /
     (-24 +
      skoXC1 * 16/3125 * (skoXC1 * 16/3125 * (4 + skoXC1 * 16/3125))) |
     cosh(skoXC1 * 16/3125) <=
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625))))),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     (-24 + skoXC1 * 16/3125 * (skoXC1 * 16/3125 * -8)) /
     (-24 +
      skoXC1 * 16/3125 * (skoXC1 * 16/3125 * (4 + skoXC1 * 16/3125))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_76, refute_0_71])).

cnf(refute_0_78, plain,
    (skoXC1 < 0 |
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     9375/32 < skoXC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain, (9375/32 < skoXC1 | 0 <= skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_80, plain,
    ((24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     9375/32 < skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_79, refute_0_78])).

cnf(refute_0_81, plain, (skoXC1 <= 0 | skoXC1 <= 9375/32),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_82, plain,
    ((24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) /
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_81, refute_0_80])).

cnf(refute_0_83, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) *
     (915648/6625 +
      skoXC1 *
      (-1473621/4140625 +
       skoXC1 *
       (-2485636183/4140625000000 +
        skoXC1 *
        (94311744/40435791015625 +
         skoXC1 * -50594321/25272369384765625)))) |
     915648/6625 +
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 +
        skoXC1 * -50594321/25272369384765625))) <= 0 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_1, refute_0_82])).

cnf(refute_0_84, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     (-21975552/6625 +
      skoXC1 *
      (35366904/4140625 +
       skoXC1 *
       (-7545068283/517578125000 +
        skoXC1 *
        (754493952/40435791015625 +
         skoXC1 *
         (109896950456/631809234619140625 +
          skoXC1 *
          (-193150451712/394880771636962890625 +
           skoXC1 * 103617169408/246800482273101806640625)))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     skoXC1 *
     (-1473621/4140625 +
      skoXC1 *
      (-2485636183/4140625000000 +
       skoXC1 *
       (94311744/40435791015625 +
        skoXC1 * -50594321/25272369384765625))) <= -915648/6625 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     (-21975552/6625 +
      skoXC1 *
      (35366904/4140625 +
       skoXC1 *
       (-7545068283/517578125000 +
        skoXC1 *
        (754493952/40435791015625 +
         skoXC1 *
         (109896950456/631809234619140625 +
          skoXC1 *
          (-193150451712/394880771636962890625 +
           skoXC1 * 103617169408/246800482273101806640625)))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_84, refute_0_70])).

cnf(refute_0_86, plain,
    (-3662592/6625 +
     skoXC1 *
     (5894484/4140625 +
      skoXC1 *
      (-78135296/64697265625 +
       skoXC1 *
       (125748992/40435791015625 +
        skoXC1 *
        (20002635776/631809234619140625 +
         skoXC1 * -32191741952/394880771636962890625)))) <
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) *
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 <=
     -24 + skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125)) |
     (-3662592/6625 +
      skoXC1 *
      (5894484/4140625 +
       skoXC1 *
       (-78135296/64697265625 +
        skoXC1 *
        (125748992/40435791015625 +
         skoXC1 *
         (20002635776/631809234619140625 +
          skoXC1 * -32191741952/394880771636962890625))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) <=
     4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_87, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) *
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) |
     152608/6625 +
     skoXC1 *
     (-491207/8281250 +
      skoXC1 *
      (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) <= 0 |
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) <=
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_88, plain,
    ((-2 + skoXC1 * 103/20000) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     -4000/53 +
     skoXC1 *
     (-103/530 + skoXC1 * (2048/4140625 + skoXC1 * 3296/2587890625)) |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <= 0 |
     (-4000/53 +
      skoXC1 *
      (-103/530 + skoXC1 * (2048/4140625 + skoXC1 * 3296/2587890625))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <=
     -2 + skoXC1 * 103/20000),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_89, plain,
    (-2 + skoXC1 * 103/20000 <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) *
     (-2 + skoXC1 * -103/20000) | 0 <= -2 + skoXC1 * -103/20000 |
     (-2 + skoXC1 * 103/20000) / (-2 + skoXC1 * -103/20000) <=
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_90, plain,
    (~ lgen(0, exp(X_000103), X_000104) | exp(X_000103) <= X_000104),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_91, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, -(X_000103 + 2) / (X_000103 - 2), X_000104) |
     lgen(0, exp(X_000103), X_000104)),
    inference(subst, [], [exp_upper_bound_cf1])).

cnf(refute_0_92, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, -(X_000103 + 2) / (X_000103 - 2), X_000104) |
     exp(X_000103) <= X_000104),
    inference(resolve, [], [refute_0_91, refute_0_90])).

cnf(refute_0_93, plain,
    (3/2 < X_000103 | X_000103 < 0 |
     X_000104 < (-2 + X_000103 * -1) / (-2 + X_000103) |
     exp(X_000103) <= X_000104), inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (skoXC1 * -103/20000 < 0 | 3/2 < skoXC1 * -103/20000 |
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (-2 + skoXC1 * -103/20000 * -1) / (-2 + skoXC1 * -103/20000) |
     exp(skoXC1 * -103/20000) <=
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)))),
    inference(subst, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 * -103/20000 < 0 | 3/2 < skoXC1 * -103/20000 |
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (-2 + skoXC1 * -103/20000 * -1) / (-2 + skoXC1 * -103/20000)),
    inference(resolve, [], [refute_0_94, refute_0_58])).

cnf(refute_0_96, plain,
    (skoXC1 < -30000/103 |
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (-2 + skoXC1 * 103/20000) / (-2 + skoXC1 * -103/20000) | 0 < skoXC1),
    inference(arithmetic, [], [refute_0_95])).

cnf(refute_0_97, plain, (0 < skoXC1 | -30000/103 <= skoXC1),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_98, plain,
    ((2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     (-2 + skoXC1 * 103/20000) / (-2 + skoXC1 * -103/20000) | 0 < skoXC1),
    inference(resolve, [], [refute_0_97, refute_0_96])).

cnf(refute_0_99, plain,
    (-2 + skoXC1 * 103/20000 <
     (2000/53 + skoXC1 * (skoXC1 * -1024/4140625)) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) *
     (-2 + skoXC1 * -103/20000) | 0 < skoXC1 |
     0 <= -2 + skoXC1 * -103/20000),
    inference(resolve, [], [refute_0_89, refute_0_98])).

cnf(refute_0_100, plain,
    (-2 + skoXC1 * 103/20000 <
     (-4000/53 +
      skoXC1 *
      (-103/530 + skoXC1 * (2048/4140625 + skoXC1 * 3296/2587890625))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     0 < skoXC1 | skoXC1 <= -40000/103),
    inference(arithmetic, [], [refute_0_99])).

cnf(refute_0_101, plain, (-40000/103 < skoXC1 | 0 < skoXC1),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_102, plain,
    (-2 + skoXC1 * 103/20000 <
     (-4000/53 +
      skoXC1 *
      (-103/530 + skoXC1 * (2048/4140625 + skoXC1 * 3296/2587890625))) /
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) |
     0 < skoXC1), inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    ((-2 + skoXC1 * 103/20000) *
     (skoXC1 * 233472/828125 +
      cosh(skoXC1 * 16/3125) *
      (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625))) <
     -4000/53 +
     skoXC1 *
     (-103/530 + skoXC1 * (2048/4140625 + skoXC1 * 3296/2587890625)) |
     0 < skoXC1 |
     skoXC1 * 233472/828125 +
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <= 0),
    inference(resolve, [], [refute_0_88, refute_0_102])).

cnf(refute_0_104, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     cosh(skoXC1 * 16/3125) *
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) |
     0 < skoXC1 |
     cosh(skoXC1 * 16/3125) *
     (76304/6625 + skoXC1 * (skoXC1 * -4883456/64697265625)) <=
     skoXC1 * -233472/828125), inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     cosh(skoXC1 * 16/3125) *
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) |
     0 < skoXC1), inference(resolve, [], [refute_0_104, refute_0_57])).

cnf(refute_0_106, plain,
    ((4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     cosh(skoXC1 * 16/3125) | 0 < skoXC1 |
     152608/6625 +
     skoXC1 *
     (-491207/8281250 +
      skoXC1 *
      (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) <= 0),
    inference(split, [], [refute_0_105])).

cnf(refute_0_107, plain,
    ((4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     cosh(skoXC1 * 16/3125) | 0 < skoXC1 |
     skoXC1 *
     (-491207/8281250 +
      skoXC1 *
      (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) <=
     -152608/6625), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (-152608/6625 <
     skoXC1 *
     (-491207/8281250 +
      skoXC1 * (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) |
     0 < skoXC1), inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_109, plain,
    ((4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     cosh(skoXC1 * 16/3125) | 0 < skoXC1),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     (-24 + skoXC1 * 16/3125 * (skoXC1 * 16/3125 * -8)) /
     (-24 +
      skoXC1 * 16/3125 * (skoXC1 * 16/3125 * (4 + skoXC1 * 16/3125))) |
     cosh(skoXC1 * 16/3125) <=
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)))),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_111, plain,
    (skoXC1 * 16/3125 < 0 | 3/2 < skoXC1 * 16/3125 |
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     (-24 + skoXC1 * 16/3125 * (skoXC1 * 16/3125 * -8)) /
     (-24 +
      skoXC1 * 16/3125 * (skoXC1 * 16/3125 * (4 + skoXC1 * 16/3125))) |
     0 < skoXC1), inference(resolve, [], [refute_0_110, refute_0_109])).

cnf(refute_0_112, plain,
    (skoXC1 < 0 |
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1 | 9375/32 < skoXC1),
    inference(arithmetic, [], [refute_0_111])).

cnf(refute_0_113, plain, (0 < skoXC1 | 9375/32 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_114, plain,
    ((4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1 | 9375/32 < skoXC1),
    inference(resolve, [], [refute_0_113, refute_0_112])).

cnf(refute_0_115, plain, (0 < skoXC1 | skoXC1 <= 9375/32),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_116, plain,
    ((4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) /
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1), inference(resolve, [], [refute_0_115, refute_0_114])).

cnf(refute_0_117, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     (-24 + skoXC1 * (skoXC1 * -2048/9765625)) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) *
     (152608/6625 +
      skoXC1 *
      (-491207/8281250 +
       skoXC1 *
       (-9766912/64697265625 + skoXC1 * 15718624/40435791015625))) |
     0 < skoXC1 |
     152608/6625 +
     skoXC1 *
     (-491207/8281250 +
      skoXC1 *
      (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) <= 0),
    inference(resolve, [], [refute_0_87, refute_0_116])).

cnf(refute_0_118, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     (-3662592/6625 +
      skoXC1 *
      (5894484/4140625 +
       skoXC1 *
       (-78135296/64697265625 +
        skoXC1 *
        (125748992/40435791015625 +
         skoXC1 *
         (20002635776/631809234619140625 +
          skoXC1 * -32191741952/394880771636962890625))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1 |
     skoXC1 *
     (-491207/8281250 +
      skoXC1 *
      (-9766912/64697265625 + skoXC1 * 15718624/40435791015625)) <=
     -152608/6625), inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (4000/53 +
     skoXC1 *
     (-612013/1656250 +
      skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625)) <
     (-3662592/6625 +
      skoXC1 *
      (5894484/4140625 +
       skoXC1 *
       (-78135296/64697265625 +
        skoXC1 *
        (125748992/40435791015625 +
         skoXC1 *
         (20002635776/631809234619140625 +
          skoXC1 * -32191741952/394880771636962890625))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1), inference(resolve, [], [refute_0_118, refute_0_108])).

cnf(refute_0_120, plain,
    (-3662592/6625 +
     skoXC1 *
     (5894484/4140625 +
      skoXC1 *
      (-78135296/64697265625 +
       skoXC1 *
       (125748992/40435791015625 +
        skoXC1 *
        (20002635776/631809234619140625 +
         skoXC1 * -32191741952/394880771636962890625)))) <
     (4000/53 +
      skoXC1 *
      (-612013/1656250 +
       skoXC1 * (495488/517578125 + skoXC1 * -3296/2587890625))) *
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 < skoXC1 |
     0 <=
     -24 + skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))),
    inference(resolve, [], [refute_0_86, refute_0_119])).

cnf(refute_0_121, plain,
    (8337408/6625 <
     skoXC1 *
     (30826296/4140625 +
      skoXC1 *
      (-896328704/64697265625 +
       skoXC1 *
       (-46902272/40435791015625 +
        skoXC1 *
        (12084762624/631809234619140625 +
         skoXC1 *
         (30193713152/394880771636962890625 +
          skoXC1 * -13500416/78976154327392578125))))) | 0 < skoXC1 |
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24), inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (0 < skoXC1 |
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24 |
     skoXC1 *
     (30826296/4140625 +
      skoXC1 *
      (-896328704/64697265625 +
       skoXC1 *
       (-46902272/40435791015625 +
        skoXC1 *
        (12084762624/631809234619140625 +
         skoXC1 *
         (30193713152/394880771636962890625 +
          skoXC1 * -13500416/78976154327392578125))))) <= 8337408/6625),
    inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_123, plain,
    (0 < skoXC1 |
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24), inference(resolve, [], [refute_0_122, refute_0_121])).

cnf(refute_0_124, plain,
    (-24 <
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) |
     0 < skoXC1), inference(decision, [], [refute_0_29, refute_0_28])).

cnf(refute_0_125, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (24000/53 +
     skoXC1 *
     (-1836039/828125 +
      skoXC1 *
      (223420517/33125000000 +
       skoXC1 * (-4890852/323486328125 + skoXC1 * -10609/1617431640625))) <
     (-21975552/6625 +
      skoXC1 *
      (35366904/4140625 +
       skoXC1 *
       (-7545068283/517578125000 +
        skoXC1 *
        (754493952/40435791015625 +
         skoXC1 *
         (109896950456/631809234619140625 +
          skoXC1 *
          (-193150451712/394880771636962890625 +
           skoXC1 * 103617169408/246800482273101806640625)))))) /
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125)))),
    inference(resolve, [], [refute_0_85, refute_0_125])).

cnf(refute_0_127, plain,
    (-21975552/6625 +
     skoXC1 *
     (35366904/4140625 +
      skoXC1 *
      (-7545068283/517578125000 +
       skoXC1 *
       (754493952/40435791015625 +
        skoXC1 *
        (109896950456/631809234619140625 +
         skoXC1 *
         (-193150451712/394880771636962890625 +
          skoXC1 * 103617169408/246800482273101806640625))))) <
     (24000/53 +
      skoXC1 *
      (-1836039/828125 +
       skoXC1 *
       (223420517/33125000000 +
        skoXC1 *
        (-4890852/323486328125 + skoXC1 * -10609/1617431640625)))) *
     (-24 +
      skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))) |
     0 <=
     -24 + skoXC1 * (skoXC1 * (1024/9765625 + skoXC1 * 4096/30517578125))),
    inference(resolve, [], [refute_0_0, refute_0_126])).

cnf(refute_0_128, plain,
    (50024448/6625 <
     skoXC1 *
     (184957776/4140625 +
      skoXC1 *
      (-6457703199/64697265625 +
       skoXC1 *
       (6975142368/40435791015625 +
        skoXC1 *
        (248393064944/631809234619140625 +
         skoXC1 *
         (-75405777088/394880771636962890625 +
          skoXC1 *
          (-774184414208/246800482273101806640625 +
           skoXC1 * -43454464/49360096454620361328125)))))) |
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24), inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24 |
     skoXC1 *
     (184957776/4140625 +
      skoXC1 *
      (-6457703199/64697265625 +
       skoXC1 *
       (6975142368/40435791015625 +
        skoXC1 *
        (248393064944/631809234619140625 +
         skoXC1 *
         (-75405777088/394880771636962890625 +
          skoXC1 *
          (-774184414208/246800482273101806640625 +
           skoXC1 * -43454464/49360096454620361328125)))))) <=
     50024448/6625),
    inference(decision, [], [refute_0_125, refute_0_29, refute_0_28])).

cnf(refute_0_130, plain,
    (skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125)) <=
     -24), inference(resolve, [], [refute_0_129, refute_0_128])).

cnf(refute_0_131, plain,
    (-24 <
     skoXC1 * (skoXC1 * (-1024/9765625 + skoXC1 * -4096/30517578125))),
    inference(decision, [], [refute_0_125, refute_0_29, refute_0_28])).

cnf(refute_0_132, plain, ($false),
    inference(resolve, [], [refute_0_130, refute_0_131])).
% SZS output end CNFRefutation for MUTANT-3-h.tptp
