% SZS output start CNFRefutation for MUTANT-2.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

fof('MUTANT', conjecture,
    (! [X] :
       (0 <= X =>
        40000 * exp(-139/500000 * X) + 211/10000 * exp(-103/10000 * X) <
        49964))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        40000 * exp(-139/500000 * X) + 211/10000 * exp(-103/10000 * X) <
        49964)), inference(strip, [], ['MUTANT'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          40000 * exp(-139/500000 * X) + 211/10000 * exp(-103/10000 * X) <
          49964)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        49964 <=
        40000 * exp(-139/500000 * X) + 211/10000 * exp(-103/10000 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     49964 <=
     40000 * exp(-139/500000 * skoXC1) +
     211/10000 * exp(-103/10000 * skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (49964 <=
     40000 * exp(-139/500000 * skoXC1) +
     211/10000 * exp(-103/10000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (2304 <
     (-826454016/625 +
      skoXC1 *
      (-151008872832/9765625 +
       skoXC1 *
       (-13639144631382/152587890625 +
        skoXC1 *
        (-13018295417563233/38146972656250000 +
         skoXC1 *
         (-544528567848190080159/610351562500000000000000 +
          skoXC1 *
          (-489119267414870397839493/305175781250000000000000000000 +
           skoXC1 *
           (-1164447569999682163026047259/
            610351562500000000000000000000000000 +
            skoXC1 *
            (-6836135912988367705582648977/
             6103515625000000000000000000000000000000 +
             skoXC1 *
             (-36066465222743293467342643737/
              244140625000000000000000000000000000000000000 +
              skoXC1 *
              (-64456610154457681374671706849/
               4882812500000000000000000000000000000000000000000 +
               skoXC1 *
               (-306044602821710899768313001651/
                390625000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-119220833812171406481648593829/
                 3906250000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -107574505431956767290462148979/
                 156250000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) |
     0 <=
     -1899/1562500 +
     skoXC1 *
     (-263961/781250000000 +
      skoXC1 *
      (-36690579/781250000000000000 +
       skoXC1 *
       (-1699996827/390625000000000000000000 +
        skoXC1 *
        (-1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (-32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          -1521847926176971/
          6250000000000000000000000000000000000000000))))) |
     2304 *
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     -826454016/625 +
     skoXC1 *
     (-151008872832/9765625 +
      skoXC1 *
      (-13639144631382/152587890625 +
       skoXC1 *
       (-13018295417563233/38146972656250000 +
        skoXC1 *
        (-544528567848190080159/610351562500000000000000 +
         skoXC1 *
         (-489119267414870397839493/305175781250000000000000000000 +
          skoXC1 *
          (-1164447569999682163026047259/
           610351562500000000000000000000000000 +
           skoXC1 *
           (-6836135912988367705582648977/
            6103515625000000000000000000000000000000 +
            skoXC1 *
            (-36066465222743293467342643737/
             244140625000000000000000000000000000000000000 +
             skoXC1 *
             (-64456610154457681374671706849/
              4882812500000000000000000000000000000000000000000 +
              skoXC1 *
              (-306044602821710899768313001651/
               390625000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-119220833812171406481648593829/
                3906250000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -107574505431956767290462148979/
                156250000000000000000000000000000000000000000000000000000000000)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (14832/625 +
       skoXC1 *
       (95481/781250 +
        skoXC1 *
        (3278181/7812500000 +
         skoXC1 *
         (2363568501/2500000000000000 +
          skoXC1 *
          (34778222229/25000000000000000000 +
           skoXC1 * 1194052296529/1000000000000000000000000)))))) <
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) *
     (2304 +
      skoXC1 *
      (14832/625 +
       skoXC1 *
       (95481/781250 +
        skoXC1 *
        (3278181/7812500000 +
         skoXC1 *
         (2363568501/2500000000000000 +
          skoXC1 *
          (34778222229/25000000000000000000 +
           skoXC1 * 1194052296529/1000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (14832/625 +
      skoXC1 *
      (95481/781250 +
       skoXC1 *
       (3278181/7812500000 +
        skoXC1 *
        (2363568501/2500000000000000 +
         skoXC1 *
         (34778222229/25000000000000000000 +
          skoXC1 * 1194052296529/1000000000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (exp(X_000033) < X_000034 | ~ lgen(1, exp(X_000033), X_000034)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (0 < X_000033 |
     ~ lgen(1,
            2304 /
            (-(X_000033 ^ 3) + 6 * X_000033 ^ 2 - 24 * X_000033 + 48) ^ 2,
            X_000034) | lgen(1, exp(X_000033), X_000034)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_4, plain,
    (0 < X_000033 | exp(X_000033) < X_000034 |
     ~ lgen(1,
            2304 /
            (-(X_000033 ^ 3) + 6 * X_000033 ^ 2 - 24 * X_000033 + 48) ^ 2,
            X_000034)), inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (0 < X_000033 | exp(X_000033) < X_000034 |
     X_000034 <=
     2304 /
     (2304 +
      X_000033 *
      (-2304 +
       X_000033 *
       (1152 +
        X_000033 *
        (-384 + X_000033 * (84 + X_000033 * (-12 + X_000033))))))),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (0 < skoXC1 * -103/10000 |
     exp(skoXC1 * -103/10000) <
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -103/10000 *
      (-2304 +
       skoXC1 * -103/10000 *
       (1152 +
        skoXC1 * -103/10000 *
        (-384 +
         skoXC1 * -103/10000 *
         (84 + skoXC1 * -103/10000 * (-12 + skoXC1 * -103/10000))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <
     12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 |
     (12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000) *
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_8, plain,
    (49964 <=
     40000 * exp(-139/500000 * skoXC1) +
     211/10000 * exp(-103/10000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    (12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 <=
     exp(skoXC1 * -139/500000)), inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (0 < skoXC1 * -139/500000 |
     exp(skoXC1 * -139/500000) <
     12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 |
     12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 <=
     2304 /
     (2304 +
      skoXC1 * -139/500000 *
      (-2304 +
       skoXC1 * -139/500000 *
       (1152 +
        skoXC1 * -139/500000 *
        (-384 +
         skoXC1 * -139/500000 *
         (84 + skoXC1 * -139/500000 * (-12 + skoXC1 * -139/500000))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_11, plain,
    (0 < skoXC1 * -139/500000 |
     12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 <=
     2304 /
     (2304 +
      skoXC1 * -139/500000 *
      (-2304 +
       skoXC1 * -139/500000 *
       (1152 +
        skoXC1 * -139/500000 *
        (-384 +
         skoXC1 * -139/500000 *
         (84 + skoXC1 * -139/500000 * (-12 + skoXC1 * -139/500000))))))),
    inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (skoXC1 < 0 |
     12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 <=
     2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain,
    (12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000 <=
     2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    ((12491/10000 + exp(skoXC1 * -103/10000) * -211/400000000) *
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_14, refute_0_7])).

cnf(refute_0_16, plain,
    (exp(skoXC1 * -103/10000) *
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     -358704/625 +
     skoXC1 *
     (-15626241/19531250 +
      skoXC1 *
      (-2172047499/19531250000000 +
       skoXC1 *
       (-100638200787/9765625000000000000 +
        skoXC1 *
        (-97920969365751/156250000000000000000000000 +
         skoXC1 *
         (-1944430677405627/78125000000000000000000000000000 +
          skoXC1 *
          -90091954719794051/156250000000000000000000000000000000000))))) |
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (-2304 <
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 * 7212549413161/15625000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -103/10000) *
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     -358704/625 +
     skoXC1 *
     (-15626241/19531250 +
      skoXC1 *
      (-2172047499/19531250000000 +
       skoXC1 *
       (-100638200787/9765625000000000000 +
        skoXC1 *
        (-97920969365751/156250000000000000000000000 +
         skoXC1 *
         (-1944430677405627/78125000000000000000000000000000 +
          skoXC1 *
          -90091954719794051/156250000000000000000000000000000000000)))))),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (0 <=
     -1899/1562500 +
     skoXC1 *
     (-263961/781250000000 +
      skoXC1 *
      (-36690579/781250000000000000 +
       skoXC1 *
       (-1699996827/390625000000000000000000 +
        skoXC1 *
        (-1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (-32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          -1521847926176971/
          6250000000000000000000000000000000000000000))))) |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     exp(skoXC1 * -103/10000)), inference(split, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (skoXC1 *
     (263961/781250000000 +
      skoXC1 *
      (36690579/781250000000000000 +
       skoXC1 *
       (1699996827/390625000000000000000000 +
        skoXC1 *
        (1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          1521847926176971/
          6250000000000000000000000000000000000000000))))) <=
     -1899/1562500 |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     exp(skoXC1 * -103/10000)), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (-1899/1562500 <
     skoXC1 *
     (263961/781250000000 +
      skoXC1 *
      (36690579/781250000000000000 +
       skoXC1 *
       (1699996827/390625000000000000000000 +
        skoXC1 *
        (1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          1521847926176971/
          6250000000000000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_22, plain,
    ((-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     exp(skoXC1 * -103/10000)),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (0 < skoXC1 * -103/10000 |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 * -103/10000 *
      (-2304 +
       skoXC1 * -103/10000 *
       (1152 +
        skoXC1 * -103/10000 *
        (-384 +
         skoXC1 * -103/10000 *
         (84 + skoXC1 * -103/10000 * (-12 + skoXC1 * -103/10000))))))),
    inference(resolve, [], [refute_0_22, refute_0_6])).

cnf(refute_0_24, plain,
    (skoXC1 < 0 |
     (-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (14832/625 +
       skoXC1 *
       (95481/781250 +
        skoXC1 *
        (3278181/7812500000 +
         skoXC1 *
         (2363568501/2500000000000000 +
          skoXC1 *
          (34778222229/25000000000000000000 +
           skoXC1 * 1194052296529/1000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    ((-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (14832/625 +
       skoXC1 *
       (95481/781250 +
        skoXC1 *
        (3278181/7812500000 +
         skoXC1 *
         (2363568501/2500000000000000 +
          skoXC1 *
          (34778222229/25000000000000000000 +
           skoXC1 * 1194052296529/1000000000000000000000000))))))),
    inference(resolve, [], [refute_0_13, refute_0_24])).

cnf(refute_0_26, plain,
    ((-358704/625 +
      skoXC1 *
      (-15626241/19531250 +
       skoXC1 *
       (-2172047499/19531250000000 +
        skoXC1 *
        (-100638200787/9765625000000000000 +
         skoXC1 *
         (-97920969365751/156250000000000000000000000 +
          skoXC1 *
          (-1944430677405627/78125000000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           156250000000000000000000000000000000000)))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) *
     (2304 +
      skoXC1 *
      (14832/625 +
       skoXC1 *
       (95481/781250 +
        skoXC1 *
        (3278181/7812500000 +
         skoXC1 *
         (2363568501/2500000000000000 +
          skoXC1 *
          (34778222229/25000000000000000000 +
           skoXC1 * 1194052296529/1000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (14832/625 +
      skoXC1 *
      (95481/781250 +
       skoXC1 *
       (3278181/7812500000 +
        skoXC1 *
        (2363568501/2500000000000000 +
         skoXC1 *
         (34778222229/25000000000000000000 +
          skoXC1 * 1194052296529/1000000000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_25, refute_0_1])).

cnf(refute_0_27, plain,
    (skoXC1 *
     (14832/625 +
      skoXC1 *
      (95481/781250 +
       skoXC1 *
       (3278181/7812500000 +
        skoXC1 *
        (2363568501/2500000000000000 +
         skoXC1 *
         (34778222229/25000000000000000000 +
          skoXC1 * 1194052296529/1000000000000000000000000))))) <= -2304 |
     (-826454016/625 +
      skoXC1 *
      (-151008872832/9765625 +
       skoXC1 *
       (-13639144631382/152587890625 +
        skoXC1 *
        (-13018295417563233/38146972656250000 +
         skoXC1 *
         (-544528567848190080159/610351562500000000000000 +
          skoXC1 *
          (-489119267414870397839493/305175781250000000000000000000 +
           skoXC1 *
           (-1164447569999682163026047259/
            610351562500000000000000000000000000 +
            skoXC1 *
            (-6836135912988367705582648977/
             6103515625000000000000000000000000000000 +
             skoXC1 *
             (-36066465222743293467342643737/
              244140625000000000000000000000000000000000000 +
              skoXC1 *
              (-64456610154457681374671706849/
               4882812500000000000000000000000000000000000000000 +
               skoXC1 *
               (-306044602821710899768313001651/
                390625000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-119220833812171406481648593829/
                 3906250000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -107574505431956767290462148979/
                 156250000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <= 2304),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (-2304 <
     skoXC1 *
     (14832/625 +
      skoXC1 *
      (95481/781250 +
       skoXC1 *
       (3278181/7812500000 +
        skoXC1 *
        (2363568501/2500000000000000 +
         skoXC1 *
         (34778222229/25000000000000000000 +
          skoXC1 * 1194052296529/1000000000000000000000000)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_29, plain,
    ((-826454016/625 +
      skoXC1 *
      (-151008872832/9765625 +
       skoXC1 *
       (-13639144631382/152587890625 +
        skoXC1 *
        (-13018295417563233/38146972656250000 +
         skoXC1 *
         (-544528567848190080159/610351562500000000000000 +
          skoXC1 *
          (-489119267414870397839493/305175781250000000000000000000 +
           skoXC1 *
           (-1164447569999682163026047259/
            610351562500000000000000000000000000 +
            skoXC1 *
            (-6836135912988367705582648977/
             6103515625000000000000000000000000000000 +
             skoXC1 *
             (-36066465222743293467342643737/
              244140625000000000000000000000000000000000000 +
              skoXC1 *
              (-64456610154457681374671706849/
               4882812500000000000000000000000000000000000000000 +
               skoXC1 *
               (-306044602821710899768313001651/
                390625000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-119220833812171406481648593829/
                 3906250000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -107574505431956767290462148979/
                 156250000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <= 2304),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (0 <=
     -1899/1562500 +
     skoXC1 *
     (-263961/781250000000 +
      skoXC1 *
      (-36690579/781250000000000000 +
       skoXC1 *
       (-1699996827/390625000000000000000000 +
        skoXC1 *
        (-1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (-32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          -1521847926176971/
          6250000000000000000000000000000000000000000))))) |
     2304 *
     (-1899/1562500 +
      skoXC1 *
      (-263961/781250000000 +
       skoXC1 *
       (-36690579/781250000000000000 +
        skoXC1 *
        (-1699996827/390625000000000000000000 +
         skoXC1 *
         (-1654096912671/6250000000000000000000000000000 +
          skoXC1 *
          (-32845638694467/3125000000000000000000000000000000000 +
           skoXC1 *
           -1521847926176971/
           6250000000000000000000000000000000000000000)))))) <=
     -826454016/625 +
     skoXC1 *
     (-151008872832/9765625 +
      skoXC1 *
      (-13639144631382/152587890625 +
       skoXC1 *
       (-13018295417563233/38146972656250000 +
        skoXC1 *
        (-544528567848190080159/610351562500000000000000 +
         skoXC1 *
         (-489119267414870397839493/305175781250000000000000000000 +
          skoXC1 *
          (-1164447569999682163026047259/
           610351562500000000000000000000000000 +
           skoXC1 *
           (-6836135912988367705582648977/
            6103515625000000000000000000000000000000 +
            skoXC1 *
            (-36066465222743293467342643737/
             244140625000000000000000000000000000000000000 +
             skoXC1 *
             (-64456610154457681374671706849/
              4882812500000000000000000000000000000000000000000 +
              skoXC1 *
              (-306044602821710899768313001651/
               390625000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-119220833812171406481648593829/
                3906250000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -107574505431956767290462148979/
                156250000000000000000000000000000000000000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_29, refute_0_0])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (263961/781250000000 +
      skoXC1 *
      (36690579/781250000000000000 +
       skoXC1 *
       (1699996827/390625000000000000000000 +
        skoXC1 *
        (1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          1521847926176971/
          6250000000000000000000000000000000000000000))))) <=
     -1899/1562500 |
     skoXC1 *
     (94380540768702/6103515625 +
      skoXC1 *
      (272782892297424789/3051757812500000 +
       skoXC1 *
       (520731816687229348557/1525878906250000000000 +
        skoXC1 *
        (21781142713912716334145961/24414062500000000000000000000 +
         skoXC1 *
         (19564770696594520302831469797/
          12207031250000000000000000000000000 +
          skoXC1 *
          (46577902799987272824410554767261/
           24414062500000000000000000000000000000000 +
           skoXC1 *
           (6836135912988367705582648977/
            6103515625000000000000000000000000000000 +
            skoXC1 *
            (36066465222743293467342643737/
             244140625000000000000000000000000000000000000 +
             skoXC1 *
             (64456610154457681374671706849/
              4882812500000000000000000000000000000000000000000 +
              skoXC1 *
              (306044602821710899768313001651/
               390625000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (119220833812171406481648593829/
                3906250000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                107574505431956767290462148979/
                156250000000000000000000000000000000000000000000000000000000000))))))))))) <=
     -516532666176/390625), inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (-516532666176/390625 <
     skoXC1 *
     (94380540768702/6103515625 +
      skoXC1 *
      (272782892297424789/3051757812500000 +
       skoXC1 *
       (520731816687229348557/1525878906250000000000 +
        skoXC1 *
        (21781142713912716334145961/24414062500000000000000000000 +
         skoXC1 *
         (19564770696594520302831469797/
          12207031250000000000000000000000000 +
          skoXC1 *
          (46577902799987272824410554767261/
           24414062500000000000000000000000000000000 +
           skoXC1 *
           (6836135912988367705582648977/
            6103515625000000000000000000000000000000 +
            skoXC1 *
            (36066465222743293467342643737/
             244140625000000000000000000000000000000000000 +
             skoXC1 *
             (64456610154457681374671706849/
              4882812500000000000000000000000000000000000000000 +
              skoXC1 *
              (306044602821710899768313001651/
               390625000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (119220833812171406481648593829/
                3906250000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                107574505431956767290462148979/
                156250000000000000000000000000000000000000000000000000000000000))))))))))) |
     skoXC1 *
     (263961/781250000000 +
      skoXC1 *
      (36690579/781250000000000000 +
       skoXC1 *
       (1699996827/390625000000000000000000 +
        skoXC1 *
        (1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          1521847926176971/
          6250000000000000000000000000000000000000000))))) <=
     -1899/1562500), inference(decision, [], [refute_0_13])).

cnf(refute_0_33, plain,
    (skoXC1 *
     (263961/781250000000 +
      skoXC1 *
      (36690579/781250000000000000 +
       skoXC1 *
       (1699996827/390625000000000000000000 +
        skoXC1 *
        (1654096912671/6250000000000000000000000000000 +
         skoXC1 *
         (32845638694467/3125000000000000000000000000000000000 +
          skoXC1 *
          1521847926176971/
          6250000000000000000000000000000000000000000))))) <=
     -1899/1562500), inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain, ($false),
    inference(resolve, [], [refute_0_33, refute_0_21])).
% SZS output end CNFRefutation for MUTANT-2.tptp
