% SZS output start CNFRefutation for MUTANT-3.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

fof('MUTANT', conjecture,
    (! [X] :
       (0 <= X =>
        19448 * exp(-139/500000 * X) - 14448 * exp(-103/10000 * X) <
        49964))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        19448 * exp(-139/500000 * X) - 14448 * exp(-103/10000 * X) <
        49964)), inference(strip, [], ['MUTANT'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          19448 * exp(-139/500000 * X) - 14448 * exp(-103/10000 * X) <
          49964)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        49964 <=
        19448 * exp(-139/500000 * X) - 14448 * exp(-103/10000 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     49964 <=
     19448 * exp(-139/500000 * skoXC1) - 14448 * exp(-103/10000 * skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (49964 <=
     19448 * exp(-139/500000 * skoXC1) - 14448 * exp(-103/10000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <
     -8788608/2431 +
     skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000))))) |
     4161024/2431 +
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     0 |
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) |
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <
     exp(skoXC1 * -103/10000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <
     12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 |
     (12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431) *
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (49964 <=
     19448 * exp(-139/500000 * skoXC1) - 14448 * exp(-103/10000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 <=
     exp(skoXC1 * -139/500000)), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -139/500000 |
     exp(skoXC1 * -139/500000) <
     12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 |
     12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 <=
     2304 /
     (2304 +
      skoXC1 * -139/500000 *
      (-2304 +
       skoXC1 * -139/500000 *
       (1152 +
        skoXC1 * -139/500000 *
        (-384 +
         skoXC1 * -139/500000 *
         (84 + skoXC1 * -139/500000 * (-12 + skoXC1 * -139/500000))))))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -139/500000 |
     12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 <=
     2304 /
     (2304 +
      skoXC1 * -139/500000 *
      (-2304 +
       skoXC1 * -139/500000 *
       (1152 +
        skoXC1 * -139/500000 *
        (-384 +
         skoXC1 * -139/500000 *
         (84 + skoXC1 * -139/500000 * (-12 + skoXC1 * -139/500000))))))),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 <=
     2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431 <=
     2304 /
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((12491/4862 + exp(skoXC1 * -103/10000) * 1806/2431) *
     (2304 +
      skoXC1 *
      (10008/15625 +
       skoXC1 *
       (173889/1953125000 +
        skoXC1 *
        (8056857/976562500000000 +
         skoXC1 *
         (7839321861/15625000000000000000000 +
          skoXC1 *
          (155666534097/7812500000000000000000000000 +
           skoXC1 *
           7212549413161/15625000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_5])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -103/10000) *
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <=
     -8788608/2431 +
     skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000))))) |
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 *
          7212549413161/15625000000000000000000000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (-2304 <
     skoXC1 *
     (10008/15625 +
      skoXC1 *
      (173889/1953125000 +
       skoXC1 *
       (8056857/976562500000000 +
        skoXC1 *
        (7839321861/15625000000000000000000 +
         skoXC1 *
         (155666534097/7812500000000000000000000000 +
          skoXC1 * 7212549413161/15625000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_20, plain,
    (exp(skoXC1 * -103/10000) *
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <=
     -8788608/2431 +
     skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000)))))),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (4161024/2431 +
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     0 |
     exp(skoXC1 * -103/10000) <=
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000))))))),
    inference(split, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     -4161024/2431 |
     exp(skoXC1 * -103/10000) <=
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-4161024/2431 <
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000)))))),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_24, plain,
    (exp(skoXC1 * -103/10000) <=
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (0 <
     (-8788608/2431 +
      skoXC1 *
      (-62504964/37984375 +
       skoXC1 *
       (-2172047499/9496093750000 +
        skoXC1 *
        (-100638200787/4748046875000000000 +
         skoXC1 *
         (-97920969365751/75968750000000000000000000 +
          skoXC1 *
          (-1944430677405627/37984375000000000000000000000000 +
           skoXC1 *
           -90091954719794051/
           75968750000000000000000000000000000000)))))) /
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_24, refute_0_4])).

cnf(refute_0_26, plain,
    (0 *
     (4161024/2431 +
      skoXC1 *
      (18074448/37984375 +
       skoXC1 *
       (157021767/2374023437500 +
        skoXC1 *
        (7275341871/1187011718750000000 +
         skoXC1 *
         (7078907640483/18992187500000000000000000 +
          skoXC1 *
          (140566880289591/9496093750000000000000000000000 +
           skoXC1 *
           6512932120084383/18992187500000000000000000000000000000)))))) <
     -8788608/2431 +
     skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000))))) |
     4161024/2431 +
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_25])).

cnf(refute_0_27, plain,
    (8788608/2431 <
     skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000))))) |
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     -4161024/2431), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 *
     (-62504964/37984375 +
      skoXC1 *
      (-2172047499/9496093750000 +
       skoXC1 *
       (-100638200787/4748046875000000000 +
        skoXC1 *
        (-97920969365751/75968750000000000000000000 +
         skoXC1 *
         (-1944430677405627/37984375000000000000000000000000 +
          skoXC1 *
          -90091954719794051/75968750000000000000000000000000000000))))) <=
     8788608/2431 |
     skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     -4161024/2431), inference(decision, [], [refute_0_15])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (18074448/37984375 +
      skoXC1 *
      (157021767/2374023437500 +
       skoXC1 *
       (7275341871/1187011718750000000 +
        skoXC1 *
        (7078907640483/18992187500000000000000000 +
         skoXC1 *
         (140566880289591/9496093750000000000000000000000 +
          skoXC1 *
          6512932120084383/18992187500000000000000000000000000000))))) <=
     -4161024/2431), inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain, ($false),
    inference(resolve, [], [refute_0_29, refute_0_23])).
% SZS output end CNFRefutation for MUTANT-3.tptp
