% SZS output start CNFRefutation for NAVIGATION-2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_lower_taylor_5_cubed, axiom,
    (~ lgen(R, Y,
            (1 + X / 3 + 1/2 * (X / 3) ^ 2 + 1/6 * (X / 3) ^ 3 +
             1/24 * (X / 3) ^ 4 + 1/120 * (X / 3) ^ 5) ^ 3) |
     lgen(R, Y, exp(X)))).

fof('NAVIGATION', conjecture,
    (! [X] :
       (0 <= X =>
        -173/250 * exp(-13/10 * X) + 319/100 - X - exp(-11/10 * X) < 2))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        -173/250 * exp(-13/10 * X) + 319/100 - X - exp(-11/10 * X) < 2)),
    inference(strip, [], ['NAVIGATION'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          -173/250 * exp(-13/10 * X) + 319/100 - X - exp(-11/10 * X) < 2)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        2 <= -173/250 * exp(-13/10 * X) + 319/100 - X - exp(-11/10 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     2 <=
     -173/250 * exp(-13/10 * skoXC1) + 319/100 - skoXC1 -
     exp(-11/10 * skoXC1)), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (2 <=
     -173/250 * exp(-13/10 * skoXC1) + 319/100 - skoXC1 -
     exp(-11/10 * skoXC1)), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (X_000017 < exp(X_000016) | ~ lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3) |
     lgen(1, X_000017, exp(X_000016))),
    inference(subst, [], [exp_lower_taylor_5_cubed])).

cnf(refute_0_2, plain,
    (X_000017 < exp(X_000016) |
     ~ lgen(1, X_000017,
            (1 + X_000016 / 3 + 1/2 * (X_000016 / 3) ^ 2 +
             1/6 * (X_000016 / 3) ^ 3 + 1/24 * (X_000016 / 3) ^ 4 +
             1/120 * (X_000016 / 3) ^ 5) ^ 3)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000017 < exp(X_000016) |
     1 +
     X_000016 *
     (1 +
      X_000016 *
      (1/2 +
       X_000016 *
       (1/6 +
        X_000016 *
        (1/24 +
         X_000016 *
         (1/120 +
          X_000016 *
          (121/87480 +
           X_000016 *
           (17/87480 +
            X_000016 *
            (49/2099520 +
             X_000016 *
             (409/170061120 +
              X_000016 *
              (361/1700611200 +
               X_000016 *
               (1/62985600 +
                X_000016 *
                (181/183666009600 +
                 X_000016 *
                 (1/20407334400 +
                  X_000016 *
                  (1/550998028800 +
                   X_000016 * 1/24794911296000)))))))))))))) <= X_000017),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (95/346 +
     skoXC1 *
     (25/173 +
      skoXC1 *
      (-605/692 +
       skoXC1 *
       (1331/4152 +
        skoXC1 *
        (-14641/166080 +
         skoXC1 *
         (161051/8304000 +
          skoXC1 *
          (-214358881/60536160000 +
           skoXC1 *
           (331281907/605361600000 +
            skoXC1 *
            (-10503585169/145286784000000 +
             skoXC1 *
             (964400605619/117682295040000000 +
              skoXC1 *
              (-9363410280961/11768229504000000000 +
               skoXC1 *
               (285311670611/4358603520000000000 +
                skoXC1 *
                (-568055536186501/127096878643200000000000 +
                 skoXC1 *
                 (34522712143931/141218754048000000000000 +
                  skoXC1 *
                  (-379749833583241/38129063592960000000000000 +
                   skoXC1 *
                   4177248169415651/
                   17158078616832000000000000000)))))))))))))) <
     exp(skoXC1 * -13/10) |
     1 +
     skoXC1 * -13/10 *
     (1 +
      skoXC1 * -13/10 *
      (1/2 +
       skoXC1 * -13/10 *
       (1/6 +
        skoXC1 * -13/10 *
        (1/24 +
         skoXC1 * -13/10 *
         (1/120 +
          skoXC1 * -13/10 *
          (121/87480 +
           skoXC1 * -13/10 *
           (17/87480 +
            skoXC1 * -13/10 *
            (49/2099520 +
             skoXC1 * -13/10 *
             (409/170061120 +
              skoXC1 * -13/10 *
              (361/1700611200 +
               skoXC1 * -13/10 *
               (1/62985600 +
                skoXC1 * -13/10 *
                (181/183666009600 +
                 skoXC1 * -13/10 *
                 (1/20407334400 +
                  skoXC1 * -13/10 *
                  (1/550998028800 +
                   skoXC1 * -13/10 * 1/24794911296000)))))))))))))) <=
     95/346 +
     skoXC1 *
     (25/173 +
      skoXC1 *
      (-605/692 +
       skoXC1 *
       (1331/4152 +
        skoXC1 *
        (-14641/166080 +
         skoXC1 *
         (161051/8304000 +
          skoXC1 *
          (-214358881/60536160000 +
           skoXC1 *
           (331281907/605361600000 +
            skoXC1 *
            (-10503585169/145286784000000 +
             skoXC1 *
             (964400605619/117682295040000000 +
              skoXC1 *
              (-9363410280961/11768229504000000000 +
               skoXC1 *
               (285311670611/4358603520000000000 +
                skoXC1 *
                (-568055536186501/127096878643200000000000 +
                 skoXC1 *
                 (34522712143931/141218754048000000000000 +
                  skoXC1 *
                  (-379749833583241/38129063592960000000000000 +
                   skoXC1 *
                   4177248169415651/
                   17158078616832000000000000000))))))))))))))),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2 <=
     -173/250 * exp(-13/10 * skoXC1) + 319/100 - skoXC1 -
     exp(-11/10 * skoXC1)), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_6, plain,
    (exp(skoXC1 * -11/10) <=
     119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250),
    inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250 <
     exp(skoXC1 * -11/10) |
     1 +
     skoXC1 * -11/10 *
     (1 +
      skoXC1 * -11/10 *
      (1/2 +
       skoXC1 * -11/10 *
       (1/6 +
        skoXC1 * -11/10 *
        (1/24 +
         skoXC1 * -11/10 *
         (1/120 +
          skoXC1 * -11/10 *
          (121/87480 +
           skoXC1 * -11/10 *
           (17/87480 +
            skoXC1 * -11/10 *
            (49/2099520 +
             skoXC1 * -11/10 *
             (409/170061120 +
              skoXC1 * -11/10 *
              (361/1700611200 +
               skoXC1 * -11/10 *
               (1/62985600 +
                skoXC1 * -11/10 *
                (181/183666009600 +
                 skoXC1 * -11/10 *
                 (1/20407334400 +
                  skoXC1 * -11/10 *
                  (1/550998028800 +
                   skoXC1 * -11/10 * 1/24794911296000)))))))))))))) <=
     119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_8, plain,
    (1 +
     skoXC1 * -11/10 *
     (1 +
      skoXC1 * -11/10 *
      (1/2 +
       skoXC1 * -11/10 *
       (1/6 +
        skoXC1 * -11/10 *
        (1/24 +
         skoXC1 * -11/10 *
         (1/120 +
          skoXC1 * -11/10 *
          (121/87480 +
           skoXC1 * -11/10 *
           (17/87480 +
            skoXC1 * -11/10 *
            (49/2099520 +
             skoXC1 * -11/10 *
             (409/170061120 +
              skoXC1 * -11/10 *
              (361/1700611200 +
               skoXC1 * -11/10 *
               (1/62985600 +
                skoXC1 * -11/10 *
                (181/183666009600 +
                 skoXC1 * -11/10 *
                 (1/20407334400 +
                  skoXC1 * -11/10 *
                  (1/550998028800 +
                   skoXC1 * -11/10 * 1/24794911296000)))))))))))))) <=
     119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250),
    inference(resolve, [], [refute_0_6, refute_0_7])).

cnf(refute_0_9, plain,
    (exp(skoXC1 * -13/10) <=
     95/346 +
     skoXC1 *
     (25/173 +
      skoXC1 *
      (-605/692 +
       skoXC1 *
       (1331/4152 +
        skoXC1 *
        (-14641/166080 +
         skoXC1 *
         (161051/8304000 +
          skoXC1 *
          (-214358881/60536160000 +
           skoXC1 *
           (331281907/605361600000 +
            skoXC1 *
            (-10503585169/145286784000000 +
             skoXC1 *
             (964400605619/117682295040000000 +
              skoXC1 *
              (-9363410280961/11768229504000000000 +
               skoXC1 *
               (285311670611/4358603520000000000 +
                skoXC1 *
                (-568055536186501/127096878643200000000000 +
                 skoXC1 *
                 (34522712143931/141218754048000000000000 +
                  skoXC1 *
                  (-379749833583241/38129063592960000000000000 +
                   skoXC1 *
                   4177248169415651/
                   17158078616832000000000000000))))))))))))))),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (1 +
     skoXC1 * -13/10 *
     (1 +
      skoXC1 * -13/10 *
      (1/2 +
       skoXC1 * -13/10 *
       (1/6 +
        skoXC1 * -13/10 *
        (1/24 +
         skoXC1 * -13/10 *
         (1/120 +
          skoXC1 * -13/10 *
          (121/87480 +
           skoXC1 * -13/10 *
           (17/87480 +
            skoXC1 * -13/10 *
            (49/2099520 +
             skoXC1 * -13/10 *
             (409/170061120 +
              skoXC1 * -13/10 *
              (361/1700611200 +
               skoXC1 * -13/10 *
               (1/62985600 +
                skoXC1 * -13/10 *
                (181/183666009600 +
                 skoXC1 * -13/10 *
                 (1/20407334400 +
                  skoXC1 * -13/10 *
                  (1/550998028800 +
                   skoXC1 * -13/10 * 1/24794911296000)))))))))))))) <=
     95/346 +
     skoXC1 *
     (25/173 +
      skoXC1 *
      (-605/692 +
       skoXC1 *
       (1331/4152 +
        skoXC1 *
        (-14641/166080 +
         skoXC1 *
         (161051/8304000 +
          skoXC1 *
          (-214358881/60536160000 +
           skoXC1 *
           (331281907/605361600000 +
            skoXC1 *
            (-10503585169/145286784000000 +
             skoXC1 *
             (964400605619/117682295040000000 +
              skoXC1 *
              (-9363410280961/11768229504000000000 +
               skoXC1 *
               (285311670611/4358603520000000000 +
                skoXC1 *
                (-568055536186501/127096878643200000000000 +
                 skoXC1 *
                 (34522712143931/141218754048000000000000 +
                  skoXC1 *
                  (-379749833583241/38129063592960000000000000 +
                   skoXC1 *
                   4177248169415651/
                   17158078616832000000000000000))))))))))))))),
    inference(resolve, [], [refute_0_9, refute_0_4])).

cnf(refute_0_11, plain,
    (skoXC1 *
     (-2499/1730 +
      skoXC1 *
      (59487/34600 +
       skoXC1 *
       (-712831/1038000 +
        skoXC1 *
        (2867101/13840000 +
         skoXC1 *
         (-104496439/2076000000 +
          skoXC1 *
          (17181034783/1681560000000 +
           skoXC1 *
           (-267363865247/151340400000000 +
            skoXC1 *
            (3180281871389/12107232000000000 +
             skoXC1 *
             (-991442713540111/29420573760000000000 +
              skoXC1 *
              (3650176320561949/980685792000000000000 +
               skoXC1 *
               (-381371665821151/1089650880000000000000 +
                skoXC1 *
                (32279511980995289/1176822950400000000000000 +
                 skoXC1 *
                 (-61028071476442519/35304688512000000000000000 +
                  skoXC1 *
                  (258701191040595749/3177421966080000000000000000 +
                   skoXC1 *
                   -9899471533791613711/
                   4289519654208000000000000000000)))))))))))))) <=
     -251/346), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_14, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (0 < 595/346 + skoXC1 * -250/173 |
     595/346 + skoXC1 * -250/173 <= exp(skoXC1 * -13/10)),
    inference(subst, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_18, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (0 < 119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250 |
     119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250 <
     exp(skoXC1 * -11/10)), inference(subst, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (0 < 119/100 + skoXC1 * -1 + exp(skoXC1 * -13/10) * -173/250),
    inference(resolve, [], [refute_0_6, refute_0_19])).

cnf(refute_0_21, plain,
    (exp(skoXC1 * -13/10) < 595/346 + skoXC1 * -250/173),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain, (0 < 595/346 + skoXC1 * -250/173),
    inference(resolve, [], [refute_0_15, refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 < 119/100),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_25, plain,
    (-251/346 <
     skoXC1 *
     (-2499/1730 +
      skoXC1 *
      (59487/34600 +
       skoXC1 *
       (-712831/1038000 +
        skoXC1 *
        (2867101/13840000 +
         skoXC1 *
         (-104496439/2076000000 +
          skoXC1 *
          (17181034783/1681560000000 +
           skoXC1 *
           (-267363865247/151340400000000 +
            skoXC1 *
            (3180281871389/12107232000000000 +
             skoXC1 *
             (-991442713540111/29420573760000000000 +
              skoXC1 *
              (3650176320561949/980685792000000000000 +
               skoXC1 *
               (-381371665821151/1089650880000000000000 +
                skoXC1 *
                (32279511980995289/1176822950400000000000000 +
                 skoXC1 *
                 (-61028071476442519/35304688512000000000000000 +
                  skoXC1 *
                  (258701191040595749/3177421966080000000000000000 +
                   skoXC1 *
                   -9899471533791613711/
                   4289519654208000000000000000000))))))))))))))),
    inference(decision, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain, ($false),
    inference(resolve, [], [refute_0_11, refute_0_25])).
% SZS output end CNFRefutation for NAVIGATION-2.tptp
