% SZS output start CNFRefutation for nfm-middle-right-verify-1.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(atan_lower_bound_case_14, axiom,
    (X <= 1 | ~ lgen(R, Y, 157079/100000 - 1 / X) |
     lgen(R, Y, arctan(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof(aeronautical, conjecture,
    (! [X] :
       ((36442/51063 < X & X < 9717/12946) =>
        -3 * pi / 4 <
        -pi +
        arctan(1/5 *
               (4900000000000000000000000000 * X ^ 8 +
                60000000000000000000000000000 * X ^ 6 -
                100000000000000000000000000000 * X ^ 4 +
                239000000000000000000000000000 * X ^ 2 - 161000000000 +
                5300000000000 * X ^ 10) /
               (X *
                (370000000000000000000000000 * X ^ 8 +
                 4500000000000000000000000000 * X ^ 6 -
                 8900000000000000000000000000 * X ^ 4 +
                 3690000000000000000000000000 * X ^ 2 -
                 96000000000000000000000000)))))).

fof(subgoal_0, plain,
    (! [X] :
       ((36442/51063 < X & X < 9717/12946) =>
        -3 * pi / 4 <
        -pi +
        arctan(1/5 *
               (4900000000000000000000000000 * X ^ 8 +
                60000000000000000000000000000 * X ^ 6 -
                100000000000000000000000000000 * X ^ 4 +
                239000000000000000000000000000 * X ^ 2 - 161000000000 +
                5300000000000 * X ^ 10) /
               (X *
                (370000000000000000000000000 * X ^ 8 +
                 4500000000000000000000000000 * X ^ 6 -
                 8900000000000000000000000000 * X ^ 4 +
                 3690000000000000000000000000 * X ^ 2 -
                 96000000000000000000000000))))),
    inference(strip, [], [aeronautical])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((36442/51063 < X & X < 9717/12946) =>
          -3 * pi / 4 <
          -pi +
          arctan(1/5 *
                 (4900000000000000000000000000 * X ^ 8 +
                  60000000000000000000000000000 * X ^ 6 -
                  100000000000000000000000000000 * X ^ 4 +
                  239000000000000000000000000000 * X ^ 2 - 161000000000 +
                  5300000000000 * X ^ 10) /
                 (X *
                  (370000000000000000000000000 * X ^ 8 +
                   4500000000000000000000000000 * X ^ 6 -
                   8900000000000000000000000000 * X ^ 4 +
                   3690000000000000000000000000 * X ^ 2 -
                   96000000000000000000000000))))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (36442/51063 < X & X < 9717/12946 &
        -pi +
        arctan(1/5 *
               (4900000000000000000000000000 * X ^ 8 +
                60000000000000000000000000000 * X ^ 6 -
                100000000000000000000000000000 * X ^ 4 +
                239000000000000000000000000000 * X ^ 2 - 161000000000 +
                5300000000000 * X ^ 10) /
               (X *
                (370000000000000000000000000 * X ^ 8 +
                 4500000000000000000000000000 * X ^ 6 -
                 8900000000000000000000000000 * X ^ 4 +
                 3690000000000000000000000000 * X ^ 2 -
                 96000000000000000000000000))) <= -3 * pi / 4)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 9717/12946 & 36442/51063 < skoXC1 &
     -pi +
     arctan(1/5 *
            (4900000000000000000000000000 * skoXC1 ^ 8 +
             60000000000000000000000000000 * skoXC1 ^ 6 -
             100000000000000000000000000000 * skoXC1 ^ 4 +
             239000000000000000000000000000 * skoXC1 ^ 2 - 161000000000 +
             5300000000000 * skoXC1 ^ 10) /
            (skoXC1 *
             (370000000000000000000000000 * skoXC1 ^ 8 +
              4500000000000000000000000000 * skoXC1 ^ 6 -
              8900000000000000000000000000 * skoXC1 ^ 4 +
              3690000000000000000000000000 * skoXC1 ^ 2 -
              96000000000000000000000000))) <= -3 * pi / 4),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-pi +
     arctan(1/5 *
            (4900000000000000000000000000 * skoXC1 ^ 8 +
             60000000000000000000000000000 * skoXC1 ^ 6 -
             100000000000000000000000000000 * skoXC1 ^ 4 +
             239000000000000000000000000000 * skoXC1 ^ 2 - 161000000000 +
             5300000000000 * skoXC1 ^ 10) /
            (skoXC1 *
             (370000000000000000000000000 * skoXC1 ^ 8 +
              4500000000000000000000000000 * skoXC1 ^ 6 -
              8900000000000000000000000000 * skoXC1 ^ 4 +
              3690000000000000000000000000 * skoXC1 ^ 2 -
              96000000000000000000000000))) <= -3 * pi / 4),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_4, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_5, plain, (36442/51063 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoXC1 < 9717/12946),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 <
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) |
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) <= 0 |
     -32200000000 +
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <=
     1 *
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000)))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    (skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) /
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) <
     157079/100000 + pi * -1/4 |
     (157079/100000 + pi * -1/4) *
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) <=
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) |
     -32200000000 +
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (-pi +
     arctan(1/5 *
            (4900000000000000000000000000 * skoXC1 ^ 8 +
             60000000000000000000000000000 * skoXC1 ^ 6 -
             100000000000000000000000000000 * skoXC1 ^ 4 +
             239000000000000000000000000000 * skoXC1 ^ 2 - 161000000000 +
             5300000000000 * skoXC1 ^ 10) /
            (skoXC1 *
             (370000000000000000000000000 * skoXC1 ^ 8 +
              4500000000000000000000000000 * skoXC1 ^ 6 -
              8900000000000000000000000000 * skoXC1 ^ 4 +
              3690000000000000000000000000 * skoXC1 ^ 2 -
              96000000000000000000000000))) <= -3 * pi / 4),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (arctan((-32200000000 +
             skoXC1 *
             (skoXC1 *
              (47800000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-20000000000000000000000000000 +
                 skoXC1 *
                 (skoXC1 *
                  (12000000000000000000000000000 +
                   skoXC1 *
                   (skoXC1 *
                    (980000000000000000000000000 +
                     skoXC1 * (skoXC1 * 1060000000000)))))))))) /
            (skoXC1 *
             (-96000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (3690000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 (-8900000000000000000000000000 +
                  skoXC1 *
                  (skoXC1 *
                   (4500000000000000000000000000 +
                    skoXC1 *
                    (skoXC1 * 370000000000000000000000000)))))))))) <=
     pi * 1/4), inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (X_000048 < arctan(X_000047) | ~ lgen(1, X_000048, arctan(X_000047))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_5, plain,
    (~ lgen(1, X_000048, 157079/100000 - 1 / X_000047) | X_000047 <= 1 |
     lgen(1, X_000048, arctan(X_000047))),
    inference(subst, [], [atan_lower_bound_case_14])).

cnf(refute_0_6, plain,
    (X_000048 < arctan(X_000047) |
     ~ lgen(1, X_000048, 157079/100000 - 1 / X_000047) | X_000047 <= 1),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (X_000048 < arctan(X_000047) | X_000047 <= 1 |
     157079/100000 + 1 / X_000047 * -1 <= X_000048),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (pi * 1/4 <
     arctan((-32200000000 +
             skoXC1 *
             (skoXC1 *
              (47800000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (-20000000000000000000000000000 +
                 skoXC1 *
                 (skoXC1 *
                  (12000000000000000000000000000 +
                   skoXC1 *
                   (skoXC1 *
                    (980000000000000000000000000 +
                     skoXC1 * (skoXC1 * 1060000000000)))))))))) /
            (skoXC1 *
             (-96000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               (3690000000000000000000000000 +
                skoXC1 *
                (skoXC1 *
                 (-8900000000000000000000000000 +
                  skoXC1 *
                  (skoXC1 *
                   (4500000000000000000000000000 +
                    skoXC1 *
                    (skoXC1 * 370000000000000000000000000)))))))))) |
     157079/100000 +
     1 /
     ((-32200000000 +
       skoXC1 *
       (skoXC1 *
        (47800000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-20000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (12000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (980000000000000000000000000 +
               skoXC1 * (skoXC1 * 1060000000000)))))))))) /
      (skoXC1 *
       (-96000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (3690000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-8900000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (4500000000000000000000000000 +
              skoXC1 * (skoXC1 * 370000000000000000000000000)))))))))) *
     -1 <= pi * 1/4 |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (157079/100000 +
     1 /
     ((-32200000000 +
       skoXC1 *
       (skoXC1 *
        (47800000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-20000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (12000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (980000000000000000000000000 +
               skoXC1 * (skoXC1 * 1060000000000)))))))))) /
      (skoXC1 *
       (-96000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (3690000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-8900000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (4500000000000000000000000000 +
              skoXC1 * (skoXC1 * 370000000000000000000000000)))))))))) *
     -1 <= pi * 1/4 |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(resolve, [], [refute_0_3, refute_0_8])).

cnf(refute_0_10, plain,
    (157079/100000 + pi * -1/4 <=
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) /
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    ((157079/100000 + pi * -1/4) *
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) <=
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) |
     -32200000000 +
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <= 0 |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(resolve, [], [refute_0_10, refute_0_1])).

cnf(refute_0_12, plain,
    (skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <= 32200000000 |
     skoXC1 *
     (96000000000000000000000000 +
      skoXC1 *
      (75083762000000000000000000000 +
       pi * -11950000000000000000000000000 +
       skoXC1 *
       (-3690000000000000000000000000 +
        skoXC1 *
        (-31415800000000000000000000000 +
         pi * 5000000000000000000000000000 +
         skoXC1 *
         (8900000000000000000000000000 +
          skoXC1 *
          (18849480000000000000000000000 +
           pi * -3000000000000000000000000000 +
           skoXC1 *
           (-4500000000000000000000000000 +
            skoXC1 *
            (1539374200000000000000000000 +
             pi * -245000000000000000000000000 +
             skoXC1 *
             (-370000000000000000000000000 +
              skoXC1 * (1665037400000 + pi * -265000000000)))))))))) <=
     50579438000 + pi * -8050000000 |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_15, plain, (36442/51063 < skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_16, plain, (skoXC1 < 9717/12946),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_17, plain,
    (50579438000 + pi * -8050000000 <
     skoXC1 *
     (96000000000000000000000000 +
      skoXC1 *
      (75083762000000000000000000000 +
       pi * -11950000000000000000000000000 +
       skoXC1 *
       (-3690000000000000000000000000 +
        skoXC1 *
        (-31415800000000000000000000000 +
         pi * 5000000000000000000000000000 +
         skoXC1 *
         (8900000000000000000000000000 +
          skoXC1 *
          (18849480000000000000000000000 +
           pi * -3000000000000000000000000000 +
           skoXC1 *
           (-4500000000000000000000000000 +
            skoXC1 *
            (1539374200000000000000000000 +
             pi * -245000000000000000000000000 +
             skoXC1 *
             (-370000000000000000000000000 +
              skoXC1 * (1665037400000 + pi * -265000000000)))))))))) |
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <= 32200000000),
    inference(decision, [],
              [refute_0_13, refute_0_14, refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <= 32200000000 |
     (-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(resolve, [], [refute_0_12, refute_0_17])).

cnf(refute_0_19, plain,
    (32200000000 <
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000)))))))))),
    inference(decision, [],
              [refute_0_13, refute_0_14, refute_0_15, refute_0_16])).

cnf(refute_0_20, plain,
    ((-32200000000 +
      skoXC1 *
      (skoXC1 *
       (47800000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-20000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (12000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (980000000000000000000000000 +
              skoXC1 * (skoXC1 * 1060000000000)))))))))) /
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000))))))))) <= 1),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) <= 0 |
     -32200000000 +
     skoXC1 *
     (skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (980000000000000000000000000 +
             skoXC1 * (skoXC1 * 1060000000000))))))))) <=
     1 *
     (skoXC1 *
      (-96000000000000000000000000 +
       skoXC1 *
       (skoXC1 *
        (3690000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-8900000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (4500000000000000000000000000 +
             skoXC1 * (skoXC1 * 370000000000000000000000000)))))))))),
    inference(resolve, [], [refute_0_20, refute_0_0])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) <= 0 |
     skoXC1 *
     (96000000000000000000000000 +
      skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (-3690000000000000000000000000 +
        skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (8900000000000000000000000000 +
          skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (-4500000000000000000000000000 +
            skoXC1 *
            (980000000000000000000000000 +
             skoXC1 *
             (-370000000000000000000000000 +
              skoXC1 * 1060000000000))))))))) <= 32200000000),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (32200000000 <
     skoXC1 *
     (96000000000000000000000000 +
      skoXC1 *
      (47800000000000000000000000000 +
       skoXC1 *
       (-3690000000000000000000000000 +
        skoXC1 *
        (-20000000000000000000000000000 +
         skoXC1 *
         (8900000000000000000000000000 +
          skoXC1 *
          (12000000000000000000000000000 +
           skoXC1 *
           (-4500000000000000000000000000 +
            skoXC1 *
            (980000000000000000000000000 +
             skoXC1 *
             (-370000000000000000000000000 +
              skoXC1 * 1060000000000))))))))) |
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) <= 0),
    inference(decision, [],
              [refute_0_13, refute_0_14, refute_0_15, refute_0_16])).

cnf(refute_0_24, plain,
    (skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000)))))))) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (0 <
     skoXC1 *
     (-96000000000000000000000000 +
      skoXC1 *
      (skoXC1 *
       (3690000000000000000000000000 +
        skoXC1 *
        (skoXC1 *
         (-8900000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (4500000000000000000000000000 +
            skoXC1 * (skoXC1 * 370000000000000000000000000))))))))),
    inference(decision, [],
              [refute_0_13, refute_0_14, refute_0_15, refute_0_16])).

cnf(refute_0_26, plain, ($false),
    inference(resolve, [], [refute_0_24, refute_0_25])).
% SZS output end CNFRefutation for nfm-middle-right-verify-1.tptp
