% SZS output start CNFRefutation for Nichols-Plot-Simple-Example-3.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(atan_lower_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, Y, -3927/2500 - 3 * X / (1 + 3 * X ^ 2)) |
     lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

fof(nichols_plot_3, conjecture,
    (! [X] :
       ((4131/10000 <= X & X <= 1941/2500) =>
        36 <
        55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55))))).

fof(subgoal_0, plain,
    (! [X] :
       ((4131/10000 <= X & X <= 1941/2500) =>
        36 <
        55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55)))), inference(strip, [], [nichols_plot_3])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((4131/10000 <= X & X <= 1941/2500) =>
          36 <
          55 +
          8 *
          arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
                 (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 -
                  1597 * X ^ 2 + 55)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 1941/2500 & 4131/10000 <= X &
        55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55)) <= 36)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (4131/10000 <= skoXC1 &
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <= 36 & skoXC1 <= 1941/2500),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <= 36),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 1941/2500),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (4131/10000 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < -1 |
     -1 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <
     4021/5000 |
     4021/5000 *
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-87492693 +
           skoXC1 *
           (skoXC1 *
            (519476652 +
             skoXC1 *
             (skoXC1 *
              (237863352 +
               skoXC1 *
               (skoXC1 *
                (46000773 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     30607933 +
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <= 36),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))))) <=
     -19/8), inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (X_000076 < arctan(X_000075) | ~ lgen(1, X_000076, arctan(X_000075))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_5, plain,
    (~ lgen(1, X_000076,
            -3927/2500 - 3 * X_000075 / (1 + 3 * X_000075 ^ 2)) |
     -1 <= X_000075 | lgen(1, X_000076, arctan(X_000075))),
    inference(subst, [], [atan_lower_bound_case_11])).

cnf(refute_0_6, plain,
    (X_000076 < arctan(X_000075) |
     ~ lgen(1, X_000076,
            -3927/2500 - 3 * X_000075 / (1 + 3 * X_000075 ^ 2)) |
     -1 <= X_000075), inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (X_000076 < arctan(X_000075) | -1 <= X_000075 |
     -3927/2500 + X_000075 * 3 / (1 + X_000075 * (X_000075 * 3)) * -1 <=
     X_000076), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (-19/8 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))))) |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     -3927/2500 +
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (1 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
       3)) * -1 <= -19/8), inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     -3927/2500 +
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (1 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
       3)) * -1 <= -19/8),
    inference(resolve, [], [refute_0_3, refute_0_8])).

cnf(refute_0_10, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     4021/5000 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-37972146 +
           skoXC1 *
           (skoXC1 *
            (130046134 +
             skoXC1 *
             (skoXC1 *
              (113235524 +
               skoXC1 *
               (skoXC1 *
                (42355209 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) =
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) =
     -55), inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_12, plain,
    (-19/8 <
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <=
     -19/8 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_13, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (~ lgen(1, X_000034, X_000033) | 0 <= X_000033 |
     lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_15, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (X_000034 < arctan(X_000033) | X_000033 <= X_000034 | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (-19/8 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))))) |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -19/8), inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -19/8), inference(resolve, [], [refute_0_3, refute_0_17])).

cnf(refute_0_19, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <=
     -19/8 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_18, refute_0_12])).

cnf(refute_0_20, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= -55),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain, (skoXC1 <= 1941/2500),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_22, plain, (4131/10000 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_23, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8), inference(decision, [], [refute_0_21, refute_0_22])).

cnf(refute_0_24, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8), inference(resolve, [], [refute_0_20, refute_0_23])).

cnf(refute_0_25, plain,
    (0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8 |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_24, refute_0_11])).

cnf(refute_0_26, plain,
    (skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(decision, [], [refute_0_21, refute_0_22])).

cnf(refute_0_28, plain,
    (skoXC1 *
     (skoXC1 *
      (-24303/8 +
       skoXC1 *
       (skoXC1 *
        (172041/4 +
         skoXC1 * (skoXC1 * (97633/8 + skoXC1 * (skoXC1 * 1900))))))) <=
     24507/8 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(resolve, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-37972146 +
           skoXC1 *
           (skoXC1 *
            (130046134 +
             skoXC1 *
             (skoXC1 *
              (113235524 +
               skoXC1 *
               (skoXC1 *
                (42355209 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) !=
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) =
     -55),
    inference(decision, [], [refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_30, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     4021/5000 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) =
     -55), inference(resolve, [], [refute_0_10, refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) != -55),
    inference(decision, [], [refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_32, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     4021/5000 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))))),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     4021/5000 *
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (11904606 +
            skoXC1 *
            (skoXC1 *
             (754182316 +
              skoXC1 *
              (skoXC1 *
               (182010248 +
                skoXC1 *
                (skoXC1 *
                 (44260836 +
                  skoXC1 *
                  (skoXC1 *
                   (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-87492693 +
           skoXC1 *
           (skoXC1 *
            (519476652 +
             skoXC1 *
             (skoXC1 *
              (237863352 +
               skoXC1 *
               (skoXC1 *
                (46000773 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     30607933 +
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     0), inference(resolve, [], [refute_0_32, refute_0_1])).

cnf(refute_0_34, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-13602063929/500 +
       skoXC1 *
       (skoXC1 *
        (-62757939463/625 +
         skoXC1 *
         (skoXC1 *
          (242665942863/2500 +
           skoXC1 *
           (skoXC1 *
            (108795958159/1250 +
             skoXC1 *
             (skoXC1 *
              (-57181694099/625 +
               skoXC1 *
               (skoXC1 *
                (-13007760861/1250 +
                 skoXC1 *
                 (skoXC1 *
                  (106682904/25 +
                   skoXC1 * (skoXC1 * 514688))))))))))))))) <=
     -125709548593/5000 |
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     -30607933), inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (-3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_36, plain,
    (-19/8 <
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 + skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-87492693 +
           skoXC1 *
           (skoXC1 *
            (519476652 +
             skoXC1 *
             (skoXC1 *
              (237863352 +
               skoXC1 *
               (skoXC1 *
                (46000773 + skoXC1 * (skoXC1 * 1912800))))))))))))) <=
     -19/8 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 + skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-97290854 +
           skoXC1 *
           (skoXC1 *
            (598183796 +
             skoXC1 *
             (skoXC1 *
              (362631480 +
               skoXC1 *
               (skoXC1 *
                (127700836 +
                 skoXC1 *
                 (skoXC1 *
                  (23054400 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_37, plain,
    (X_000062 < arctan(X_000061) | ~ lgen(1, X_000062, arctan(X_000061))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_38, plain,
    (X_000061 < 0 |
     ~ lgen(1, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3)) |
     lgen(1, X_000062, arctan(X_000061))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_39, plain,
    (X_000061 < 0 | X_000062 < arctan(X_000061) |
     ~ lgen(1, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3))),
    inference(resolve, [], [refute_0_38, refute_0_37])).

cnf(refute_0_40, plain,
    (X_000061 < 0 | X_000062 < arctan(X_000061) |
     X_000061 * 3 / (3 + X_000061 * X_000061) <= X_000062),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     -19/8 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (3 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))))) <=
     -19/8), inference(subst, [], [refute_0_40])).

cnf(refute_0_42, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (3 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))))) <=
     -19/8), inference(resolve, [], [refute_0_3, refute_0_41])).

cnf(refute_0_43, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -19/8 |
     skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-37972146 +
           skoXC1 *
           (skoXC1 *
            (130046134 +
             skoXC1 *
             (skoXC1 *
              (113235524 +
               skoXC1 *
               (skoXC1 *
                (42355209 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) =
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) =
     -55), inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -19/8 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) =
     -55), inference(resolve, [], [refute_0_43, refute_0_29])).

cnf(refute_0_45, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-87492693 +
            skoXC1 *
            (skoXC1 *
             (519476652 +
              skoXC1 *
              (skoXC1 *
               (237863352 +
                skoXC1 *
                (skoXC1 *
                 (46000773 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -19/8),
    inference(resolve, [], [refute_0_44, refute_0_31])).

cnf(refute_0_46, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-87492693 +
           skoXC1 *
           (skoXC1 *
            (519476652 +
             skoXC1 *
             (skoXC1 *
              (237863352 +
               skoXC1 *
               (skoXC1 *
                (46000773 + skoXC1 * (skoXC1 * 1912800))))))))))))) <=
     -19/8 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-97290854 +
            skoXC1 *
            (skoXC1 *
             (598183796 +
              skoXC1 *
              (skoXC1 *
               (362631480 +
                skoXC1 *
                (skoXC1 *
                 (127700836 +
                  skoXC1 *
                  (skoXC1 *
                   (23054400 + skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-97290854 +
           skoXC1 *
           (skoXC1 *
            (598183796 +
             skoXC1 *
             (skoXC1 *
              (362631480 +
               skoXC1 *
               (skoXC1 *
                (127700836 +
                 skoXC1 *
                 (skoXC1 *
                  (23054400 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(resolve, [], [refute_0_45, refute_0_36])).

cnf(refute_0_47, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-97290854 +
           skoXC1 *
           (skoXC1 *
            (598183796 +
             skoXC1 *
             (skoXC1 *
              (362631480 +
               skoXC1 *
               (skoXC1 *
                (127700836 +
                 skoXC1 *
                 (skoXC1 *
                  (23054400 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     -10210711 |
     skoXC1 *
     (skoXC1 *
      (10883591/4 +
       skoXC1 *
       (skoXC1 *
        (-305969320 +
         skoXC1 *
         (skoXC1 *
          (-1274233885/4 +
           skoXC1 *
           (skoXC1 *
            (3880326335/2 +
             skoXC1 *
             (skoXC1 *
              (1099113117 +
               skoXC1 *
               (skoXC1 *
                (698580517/2 +
                 skoXC1 *
                 (skoXC1 *
                  (56667000 + skoXC1 * (skoXC1 * 4560000))))))))))))))) <=
     -189787429/8), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (-10210711 <
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-97290854 +
           skoXC1 *
           (skoXC1 *
            (598183796 +
             skoXC1 *
             (skoXC1 *
              (362631480 +
               skoXC1 *
               (skoXC1 *
                (127700836 +
                 skoXC1 *
                 (skoXC1 *
                  (23054400 + skoXC1 * (skoXC1 * 1920000))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (10883591/4 +
       skoXC1 *
       (skoXC1 *
        (-305969320 +
         skoXC1 *
         (skoXC1 *
          (-1274233885/4 +
           skoXC1 *
           (skoXC1 *
            (3880326335/2 +
             skoXC1 *
             (skoXC1 *
              (1099113117 +
               skoXC1 *
               (skoXC1 *
                (698580517/2 +
                 skoXC1 *
                 (skoXC1 *
                  (56667000 + skoXC1 * (skoXC1 * 4560000))))))))))))))) <=
     -189787429/8),
    inference(decision, [], [refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_49, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     skoXC1 *
     (skoXC1 *
      (10883591/4 +
       skoXC1 *
       (skoXC1 *
        (-305969320 +
         skoXC1 *
         (skoXC1 *
          (-1274233885/4 +
           skoXC1 *
           (skoXC1 *
            (3880326335/2 +
             skoXC1 *
             (skoXC1 *
              (1099113117 +
               skoXC1 *
               (skoXC1 *
                (698580517/2 +
                 skoXC1 *
                 (skoXC1 *
                  (56667000 + skoXC1 * (skoXC1 * 4560000))))))))))))))) <=
     -189787429/8), inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (-189787429/8 <
     skoXC1 *
     (skoXC1 *
      (10883591/4 +
       skoXC1 *
       (skoXC1 *
        (-305969320 +
         skoXC1 *
         (skoXC1 *
          (-1274233885/4 +
           skoXC1 *
           (skoXC1 *
            (3880326335/2 +
             skoXC1 *
             (skoXC1 *
              (1099113117 +
               skoXC1 *
               (skoXC1 *
                (698580517/2 +
                 skoXC1 *
                 (skoXC1 *
                  (56667000 + skoXC1 * (skoXC1 * 4560000)))))))))))))))),
    inference(decision, [], [refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_51, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) < 0),
    inference(resolve, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain,
    (-3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_35, refute_0_51])).

cnf(refute_0_53, plain,
    (-3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= -55),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) |
     -3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797)))))),
    inference(decision, [], [refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_55, plain,
    (-3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797)))))),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (-125709548593/5000 <
     skoXC1 *
     (skoXC1 *
      (-13602063929/500 +
       skoXC1 *
       (skoXC1 *
        (-62757939463/625 +
         skoXC1 *
         (skoXC1 *
          (242665942863/2500 +
           skoXC1 *
           (skoXC1 *
            (108795958159/1250 +
             skoXC1 *
             (skoXC1 *
              (-57181694099/625 +
               skoXC1 *
               (skoXC1 *
                (-13007760861/1250 +
                 skoXC1 *
                 (skoXC1 *
                  (106682904/25 +
                   skoXC1 * (skoXC1 * 514688))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     -30607933),
    inference(decision, [],
              [refute_0_55, refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_57, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     -30607933), inference(resolve, [], [refute_0_34, refute_0_56])).

cnf(refute_0_58, plain,
    (-30607933 <
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (11904606 +
           skoXC1 *
           (skoXC1 *
            (754182316 +
             skoXC1 *
             (skoXC1 *
              (182010248 +
               skoXC1 *
               (skoXC1 *
                (44260836 +
                 skoXC1 *
                 (skoXC1 *
                  (7684800 + skoXC1 * (skoXC1 * 640000)))))))))))))))),
    inference(decision, [],
              [refute_0_55, refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_59, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))))),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (-1 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_59, refute_0_0])).

cnf(refute_0_61, plain,
    (skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-5600 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 + skoXC1 * (skoXC1 * (4803 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-5600 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139),
    inference(decision, [],
              [refute_0_55, refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_63, plain,
    (skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-5600 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (-3139 <
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-5600 + skoXC1 * (skoXC1 * -800)))))))),
    inference(decision, [],
              [refute_0_55, refute_0_28, refute_0_21, refute_0_22])).

cnf(refute_0_65, plain, ($false),
    inference(resolve, [], [refute_0_63, refute_0_64])).
% SZS output end CNFRefutation for Nichols-Plot-Simple-Example-3.tptp
