% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-7.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, 3 * X / (X ^ 2 + 3), Y) | lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_1_7', conjecture,
    (! [X] :
       (X = 15839/128 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            4000000000000000 * X * (X + 10 ^ 30) /
            (X ^ 2 + 6000000000000000000000000000000 * X + 10 ^ 60)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 15839/128 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_7'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 15839/128 =>
          20 / ln(10) *
          ln(36400000000000 /
             sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2)) <
          -12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18 <=
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) & X = 15839/128)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) &
     skoXC1 = 15839/128), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 15839/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (arctan(X_000082) < X_000083 | ~ lgen(1, arctan(X_000082), X_000083)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, 3 * X_000082 / (X_000082 ^ 2 + 3), X_000083) |
     0 <= X_000082 | lgen(1, arctan(X_000082), X_000083)),
    inference(subst, [], [atan_upper_bound_case_12])).

cnf(refute_0_2, plain,
    (arctan(X_000082) < X_000083 |
     ~ lgen(1, 3 * X_000082 / (X_000082 ^ 2 + 3), X_000083) |
     0 <= X_000082), inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (arctan(X_000082) < X_000083 |
     X_000083 <= X_000082 * 3 / (3 + X_000082 * X_000082) | 0 <= X_000082),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (arctan(-5533264095502723/3391227738231320302080) < pi * 1/2 |
     0 <= -5533264095502723/3391227738231320302080 |
     pi * 1/2 <=
     -5533264095502723/3391227738231320302080 * 3 /
     (3 +
      -5533264095502723/3391227738231320302080 *
      -5533264095502723/3391227738231320302080)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_6, plain,
    (~ lgen(1, X_000034, X_000033) | 0 <= X_000033 |
     lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_7, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_6, refute_0_5])).

cnf(refute_0_8, plain,
    (X_000034 < arctan(X_000033) | X_000033 <= X_000034 | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (pi *
     -1443791129577942626020214494512436471137565286734885336349911054926890895292420542819147154282180846998929/
     226074717212884075979569153819244221548092729357025975783507366975131871345827575566604866252331945434705000 <
     arctan(-5533264095502723/3391227738231320302080) |
     -5533264095502723/3391227738231320302080 <=
     pi *
     -1443791129577942626020214494512436471137565286734885336349911054926890895292420542819147154282180846998929/
     226074717212884075979569153819244221548092729357025975783507366975131871345827575566604866252331945434705000 |
     0 <= -5533264095502723/3391227738231320302080),
    inference(subst, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <
     46889570977487067610577305977324727432197010533309091273616342780027351093949423080481009296779959053124000000000000000/
     30140414370304268248774434496453531157237322725547856658651235248082979853487355770120252324194989763281 |
     46889570977487067610577305977324727432197010533309091273616342780027351093949423080481009296779959053124000000000000000/
     30140414370304268248774434496453531157237322725547856658651235248082979853487355770120252324194989763281 *
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <=
     pi * 728000000000000 |
     pi * 5/7 + arctan(-5533264095502723/3391227738231320302080) * 270/7 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_11, plain,
    (45/14 <
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) |
     0 <=
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12) |
     45/14 *
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     pi * 728000000000000 +
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) * (pi * -20)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_12, plain,
    ((36400000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * -1) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (36400000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * -1) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_13, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain, (skoXC1 = 15839/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_15, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 15839/128 |
     skoXC1 ^ 2 = 15839/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_17, plain, (skoXC1 != 15839/128 | skoXC1 ^ 2 = 15839/128 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain, (skoXC1 ^ 2 = 15839/128 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain, (29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_21])).

cnf(refute_0_23, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_24, plain,
    (29 * skoXC1 ^ 2 != 29 * 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (29 * skoXC1 ^ 2 != 29 * 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 15839/128 ^ 2 - 90000000 |
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_29, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 15839/128 ^ 2 - 90000000 |
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 15839/128 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 != 15839/128 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_14, refute_0_33])).

cnf(refute_0_35, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_36, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_39, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_34, refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_30, refute_0_41])).

cnf(refute_0_43, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 15839/128 |
     skoXC1 ^ 4 = 15839/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_49, plain, (skoXC1 != 15839/128 | skoXC1 ^ 4 = 15839/128 ^ 4),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (skoXC1 ^ 4 = 15839/128 ^ 4),
    inference(resolve, [], [refute_0_14, refute_0_49])).

cnf(refute_0_51, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 15839/128 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (skoXC1 ^ 4 != 15839/128 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_56, plain,
    (3 * skoXC1 ^ 4 != 3 * 15839/128 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_57, plain,
    (3 * skoXC1 ^ 4 != 3 * 15839/128 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_54, refute_0_57])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_61])).

cnf(refute_0_63, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_68, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_58, refute_0_68])).

cnf(refute_0_70, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_71, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_72, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_75, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_76, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_73, refute_0_76])).

cnf(refute_0_78, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_79, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_80, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_77, refute_0_80])).

cnf(refute_0_82, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_83, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_46, refute_0_83])).

cnf(refute_0_85, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_94, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_95, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_92, refute_0_95])).

cnf(refute_0_97, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_102, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 | skoXC1 ^ 2 != 15839/128 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_105])).

cnf(refute_0_107, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_106, refute_0_109])).

cnf(refute_0_111, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 15839/128 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoXC1 ^ 4 != 15839/128 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_113])).

cnf(refute_0_115, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 15839/128 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 15839/128 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 15839/128 |
     skoXC1 ^ 6 = 15839/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 6 = 15839/128 ^ 6),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain, (skoXC1 ^ 6 = 15839/128 ^ 6),
    inference(resolve, [], [refute_0_14, refute_0_121])).

cnf(refute_0_123, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 15839/128 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (skoXC1 ^ 6 != 15839/128 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_128, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 15839/128 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_129, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 15839/128 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_126, refute_0_129])).

cnf(refute_0_131, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 15839/128 |
     skoXC1 ^ 8 = 15839/128 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 8 = 15839/128 ^ 8),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 ^ 8 = 15839/128 ^ 8),
    inference(resolve, [], [refute_0_14, refute_0_133])).

cnf(refute_0_135, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 15839/128 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (skoXC1 ^ 8 != 15839/128 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 15839/128 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 15839/128 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_138, refute_0_141])).

cnf(refute_0_143, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_144, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 15839/128 |
     skoXC1 ^ 10 = 15839/128 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_145, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 10 = 15839/128 ^ 10),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain, (skoXC1 ^ 10 = 15839/128 ^ 10),
    inference(resolve, [], [refute_0_14, refute_0_145])).

cnf(refute_0_147, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_149, plain,
    (skoXC1 ^ 10 != 15839/128 ^ 10 | 9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    inference(resolve, [], [refute_0_147, refute_0_148])).

cnf(refute_0_150, plain, (9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    inference(resolve, [], [refute_0_146, refute_0_149])).

cnf(refute_0_151, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 != 9 * 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 != 9 * 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_150, refute_0_153])).

cnf(refute_0_155, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_156, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_142, refute_0_156])).

cnf(refute_0_158, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_163, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_130, refute_0_163])).

cnf(refute_0_165, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_164, refute_0_167])).

cnf(refute_0_169, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_170, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_118, refute_0_170])).

cnf(refute_0_172, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_173, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_174, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_171, refute_0_174])).

cnf(refute_0_176, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_177, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_110, refute_0_177])).

cnf(refute_0_179, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_180, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_181, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_178, refute_0_181])).

cnf(refute_0_183, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_184, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_185, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_183, refute_0_184])).

cnf(refute_0_186, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_182, refute_0_185])).

cnf(refute_0_187, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_188, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
          10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
          4000000000000000000000000 * 15839/128 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_186, refute_0_189])).

cnf(refute_0_191, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_192, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_190, refute_0_193])).

cnf(refute_0_195, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2)) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_196, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_102, refute_0_196])).

cnf(refute_0_198, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
             10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
             4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_13, refute_0_197])).

cnf(refute_0_199, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * -20 / ln(10) <=
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     pi | pi = 0), inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    ((pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     pi <
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * -20 / ln(10) |
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 |
     pi <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_201, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 |
     pi <= 0 | pi = 0),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 |
     pi <= 0 | pi = 0), inference(arithmetic, [], [refute_0_201])).

cnf(refute_0_203, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_204, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_205, plain, (0 < pi | pi = 0),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_206, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 |
     pi = 0), inference(resolve, [], [refute_0_202, refute_0_205])).

cnf(refute_0_207, plain, (pi != 0),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_208, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 <
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) / ln(10) |
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) <=
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) *
     ln(10) | ln(10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_210, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) <=
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) *
     ln(10) | ln(10) <= 0),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12)),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (0 <= pi * -20 |
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424))),
    inference(split, [], [refute_0_211])).

cnf(refute_0_213, plain,
    (ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) | pi <= 0),
    inference(arithmetic, [], [refute_0_212])).

cnf(refute_0_214, plain, (0 < pi),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_215, plain,
    (ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424))),
    inference(resolve, [], [refute_0_213, refute_0_214])).

cnf(refute_0_216, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_217, plain,
    (~ lgen(1, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_218, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, X_000047 - 1, X_000048) |
     X_000047 <= 0), inference(resolve, [], [refute_0_217, refute_0_216])).

cnf(refute_0_219, plain,
    (ln(X_000047) < X_000048 | X_000047 <= 0 | X_000048 <= -1 + X_000047),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (ln(36400000000000 /
        sqrt(2884015903937927215862476800992050452535136618222409/
             1180591620717411303424)) <
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) |
     36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424)), inference(subst, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424)),
    inference(resolve, [], [refute_0_215, refute_0_220])).

cnf(refute_0_222, plain,
    (36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * -1) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424)),
    inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (0 <
     36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) |
     36400000000000 <=
     0 *
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_224, plain,
    (0 <
     36400000000000 /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424)),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * -1) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424)),
    inference(resolve, [], [refute_0_222, refute_0_224])).

cnf(refute_0_226, plain,
    (0 <= pi * -20 |
     (36400000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * -1) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12)),
    inference(resolve, [], [refute_0_225, refute_0_12])).

cnf(refute_0_227, plain,
    ((pi * -728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * 20)) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) |
     pi <= 0), inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain,
    ((pi * -728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * 20)) /
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     ln(10) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12)),
    inference(resolve, [], [refute_0_227, refute_0_214])).

cnf(refute_0_229, plain,
    (0 <= pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12 |
     -((pi * -728000000000000 +
       sqrt(2884015903937927215862476800992050452535136618222409/
            1180591620717411303424) * (pi * 20)) /
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424)) /
     (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12) <=
     ln(10)), inference(split, [], [refute_0_228])).

cnf(refute_0_230, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     ln(10)), inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (ln(X_000103) < X_000104 | ~ lgen(1, ln(X_000103), X_000104)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_232, plain,
    (~ lgen(1, 1/2 * (X_000103 + 5) * (X_000103 - 1) / (2 * X_000103 + 1),
            X_000104) | X_000103 <= 0 | lgen(1, ln(X_000103), X_000104)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_233, plain,
    (ln(X_000103) < X_000104 |
     ~ lgen(1, 1/2 * (X_000103 + 5) * (X_000103 - 1) / (2 * X_000103 + 1),
            X_000104) | X_000103 <= 0),
    inference(resolve, [], [refute_0_232, refute_0_231])).

cnf(refute_0_234, plain,
    (ln(X_000103) < X_000104 | X_000103 <= 0 |
     X_000104 <=
     (-5/2 + X_000103 * (2 + X_000103 * 1/2)) / (1 + X_000103 * 2)),
    inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (ln(10) <
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) |
     10 <= 0 |
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(subst, [], [refute_0_234])).

cnf(refute_0_236, plain,
    (10 <= 0 |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(resolve, [], [refute_0_230, refute_0_235])).

cnf(refute_0_237, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     (pi * 728000000000000 +
      sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) * (pi * -20)) /
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     45/14), inference(arithmetic, [], [refute_0_236])).

cnf(refute_0_238, plain,
    (0 <=
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12) |
     45/14 *
     (sqrt(2884015903937927215862476800992050452535136618222409/
           1180591620717411303424) *
      (pi * -6 + arctan(-5533264095502723/3391227738231320302080) * 12)) <=
     pi * 728000000000000 +
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) * (pi * -20) |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_237, refute_0_11])).

cnf(refute_0_239, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <=
     pi * 728000000000000 |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) <=
     0), inference(arithmetic, [], [refute_0_238])).

cnf(refute_0_240, plain,
    (0 <
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_241, plain,
    (0 <
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) |
     pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12 <= 0 |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     0 /
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12)),
    inference(split, [], [refute_0_240])).

cnf(refute_0_242, plain,
    (0 <
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 6 + arctan(-5533264095502723/3391227738231320302080) * -12) |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) *
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <=
     pi * 728000000000000),
    inference(resolve, [], [refute_0_239, refute_0_242])).

cnf(refute_0_244, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     pi * 5/7 + arctan(-5533264095502723/3391227738231320302080) * 270/7 <=
     0 |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7)),
    inference(split, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     arctan(-5533264095502723/3391227738231320302080) <= pi * -1/54 |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7)),
    inference(arithmetic, [], [refute_0_244])).

cnf(refute_0_246, plain,
    (pi * -1/54 < arctan(-5533264095502723/3391227738231320302080) |
     -5533264095502723/3391227738231320302080 <= pi * -1/54 |
     0 <= -5533264095502723/3391227738231320302080),
    inference(subst, [], [refute_0_8])).

cnf(refute_0_247, plain,
    (pi * -1/54 < arctan(-5533264095502723/3391227738231320302080) |
     pi <= 49799376859524507/565204623038553383680),
    inference(arithmetic, [], [refute_0_246])).

cnf(refute_0_248, plain, (49799376859524507/565204623038553383680 < pi),
    inference(decision, [], [refute_0_204, refute_0_203, refute_0_14])).

cnf(refute_0_249, plain,
    (pi * -1/54 < arctan(-5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_247, refute_0_248])).

cnf(refute_0_250, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7)),
    inference(resolve, [], [refute_0_245, refute_0_249])).

cnf(refute_0_251, plain,
    (X_000118 < sqrt(X_000117) | ~ lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_252, plain,
    (X_000117 < 0 |
     ~ lgen(1, X_000118,
            4000000000000000 * X_000117 * (X_000117 + 10 ^ 30) /
            (X_000117 ^ 2 + 6000000000000000000000000000000 * X_000117 +
             10 ^ 60)) | lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_253, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     ~ lgen(1, X_000118,
            4000000000000000 * X_000117 * (X_000117 + 10 ^ 30) /
            (X_000117 ^ 2 + 6000000000000000000000000000000 * X_000117 +
             10 ^ 60))),
    inference(resolve, [], [refute_0_252, refute_0_251])).

cnf(refute_0_254, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     X_000117 *
     (4000000000000000000000000000000000000000000000 +
      X_000117 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      X_000117 * (6000000000000000000000000000000 + X_000117)) <=
     X_000118), inference(arithmetic, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (2884015903937927215862476800992050452535136618222409/
     1180591620717411303424 < 0 |
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <
     sqrt(2884015903937927215862476800992050452535136618222409/
          1180591620717411303424) |
     2884015903937927215862476800992050452535136618222409/
     1180591620717411303424 *
     (4000000000000000000000000000000000000000000000 +
      2884015903937927215862476800992050452535136618222409/
      1180591620717411303424 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      2884015903937927215862476800992050452535136618222409/
      1180591620717411303424 *
      (6000000000000000000000000000000 +
       2884015903937927215862476800992050452535136618222409/
       1180591620717411303424)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7)),
    inference(subst, [], [refute_0_254])).

cnf(refute_0_256, plain,
    (2884015903937927215862476800992050452535136618222409/
     1180591620717411303424 < 0 |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     2884015903937927215862476800992050452535136618222409/
     1180591620717411303424 *
     (4000000000000000000000000000000000000000000000 +
      2884015903937927215862476800992050452535136618222409/
      1180591620717411303424 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      2884015903937927215862476800992050452535136618222409/
      1180591620717411303424 *
      (6000000000000000000000000000000 +
       2884015903937927215862476800992050452535136618222409/
       1180591620717411303424)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7)),
    inference(resolve, [], [refute_0_250, refute_0_255])).

cnf(refute_0_257, plain,
    (46889570977487067610577305977324727432197010533309091273616342780027351093949423080481009296779959053124000000000000000/
     30140414370304268248774434496453531157237322725547856658651235248082979853487355770120252324194989763281 <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (46889570977487067610577305977324727432197010533309091273616342780027351093949423080481009296779959053124000000000000000/
     30140414370304268248774434496453531157237322725547856658651235248082979853487355770120252324194989763281 *
     (pi * 5/7 +
      arctan(-5533264095502723/3391227738231320302080) * 270/7) <=
     pi * 728000000000000 |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     pi * 5/7 + arctan(-5533264095502723/3391227738231320302080) * 270/7 <=
     0), inference(resolve, [], [refute_0_257, refute_0_10])).

cnf(refute_0_259, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     arctan(-5533264095502723/3391227738231320302080) <= pi * -1/54 |
     arctan(-5533264095502723/3391227738231320302080) <=
     pi *
     -1443791129577942626020214494512436471137565286734885336349911054926890895292420542819147154282180846998929/
     226074717212884075979569153819244221548092729357025975783507366975131871345827575566604866252331945434705000),
    inference(arithmetic, [], [refute_0_258])).

cnf(refute_0_260, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     arctan(-5533264095502723/3391227738231320302080) <=
     pi *
     -1443791129577942626020214494512436471137565286734885336349911054926890895292420542819147154282180846998929/
     226074717212884075979569153819244221548092729357025975783507366975131871345827575566604866252331945434705000),
    inference(resolve, [], [refute_0_259, refute_0_249])).

cnf(refute_0_261, plain,
    (-5533264095502723/3391227738231320302080 <=
     pi *
     -1443791129577942626020214494512436471137565286734885336349911054926890895292420542819147154282180846998929/
     226074717212884075979569153819244221548092729357025975783507366975131871345827575566604866252331945434705000 |
     0 <= -5533264095502723/3391227738231320302080 |
     pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_260, refute_0_9])).

cnf(refute_0_262, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080) |
     pi <=
     658149249560674542309286504685814084876864452654369009494377263124407115167132958573869596017149579030832868248072375/
     2576038326723619686936625867403312410264515395006848816340737807615978384589794200067109096220731455064307784677896604224),
    inference(arithmetic, [], [refute_0_261])).

cnf(refute_0_263, plain,
    (658149249560674542309286504685814084876864452654369009494377263124407115167132958573869596017149579030832868248072375/
     2576038326723619686936625867403312410264515395006848816340737807615978384589794200067109096220731455064307784677896604224 <
     pi),
    inference(decision, [], [refute_0_204, refute_0_203, refute_0_14])).

cnf(refute_0_264, plain,
    (pi * 1/2 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (0 <= -5533264095502723/3391227738231320302080 |
     pi * 1/2 <=
     -5533264095502723/3391227738231320302080 * 3 /
     (3 +
      -5533264095502723/3391227738231320302080 *
      -5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_264, refute_0_4])).

cnf(refute_0_266, plain,
    (pi <=
     -112587352101769629685144091122935383040/
     34501276717679165892368517383979661457393929),
    inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (-112587352101769629685144091122935383040/
     34501276717679165892368517383979661457393929 < pi),
    inference(decision, [], [refute_0_204, refute_0_203, refute_0_14])).

cnf(refute_0_268, plain, ($false),
    inference(resolve, [], [refute_0_266, refute_0_267])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-7.tptp
