% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-6-sqrt.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Fails_1_6', conjecture,
    (! [X] :
       (X = 265/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 *
                    sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                         21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                         15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900)) + 3/2)))).

cnf(sqrt_upper_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R,
            1/8000000 *
            (X ^ 2 + 24000000000000 * X + 16000000000000000000000000) /
            (X + 4000000000000), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            8000000 * X * (X + 4000000000000) /
            (X ^ 2 + 24000000000000 * X + 16000000000000000000000000)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 265/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 *
                    sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                         21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                         15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900)) + 3/2))),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Fails_1_6'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 265/128 =>
          interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 *
                      sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                           21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                           15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900)) + 3/2))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (~ interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 *
                      sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                           21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                           15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900)) + 3/2) &
        X = 265/128)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) &
     skoXC1 = 265/128), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 265/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (~ lgen(0, arctan(X_000040), X_000041) | arctan(X_000040) <= X_000041),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (X_000040 < 0 | ~ lgen(0, X_000040, X_000041) |
     lgen(0, arctan(X_000040), X_000041)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_2, plain,
    (X_000040 < 0 | ~ lgen(0, X_000040, X_000041) |
     arctan(X_000040) <= X_000041),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000041 < X_000040 | X_000040 < 0 | arctan(X_000040) <= X_000041),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (836177642653679/2278501083899584 < 0 |
     pi *
     15833834480630041832220569266204020314032293580577068249197350480063/
     21614713530174939011217907414420533002588781281411275945246445273088 <
     836177642653679/2278501083899584 |
     arctan(836177642653679/2278501083899584) <=
     pi *
     15833834480630041832220569266204020314032293580577068249197350480063/
     21614713530174939011217907414420533002588781281411275945246445273088),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 <
     231993/101125 *
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12 |
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) <=
     231993/101125), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000223, ln(X_000222)) | X_000223 <= ln(X_000222)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (~ lgen(0, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0 |
     lgen(0, X_000223, ln(X_000222))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0 | X_000223 <= ln(X_000222)),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    ((-1/5 +
      X_000222 *
      (-95/6 +
       X_000222 *
       (-140/3 +
        X_000222 * (20 + X_000222 * (115/3 + X_000222 * 131/30))))) /
     (X_000222 *
      (5 +
       X_000222 * (40 + X_000222 * (60 + X_000222 * (20 + X_000222))))) <
     X_000223 | X_000222 <= 0 | X_000223 <= ln(X_000222)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     10 <= 0 |
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) <=
     ln(10)), inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (pi * 20 +
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     20062757577026320309797141508079595540476649613348567212816874129/
     9467448208080222069846420188471815866695581538881372160000 *
     (pi * 14953358137753600/3119868895908289175433) |
     20062757577026320309797141508079595540476649613348567212816874129/
     9467448208080222069846420188471815866695581538881372160000 <=
     (pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) |
     pi * 14953358137753600/3119868895908289175433 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain,
    (~ lgen(0, sqrt(X_000103), X_000104) | sqrt(X_000103) <= X_000104),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (X_000103 < 0 |
     ~ lgen(0,
            1/8000000 *
            (X_000103 ^ 2 + 24000000000000 * X_000103 +
             16000000000000000000000000) / (X_000103 + 4000000000000),
            X_000104) | lgen(0, sqrt(X_000103), X_000104)),
    inference(subst, [], [sqrt_upper_bound_1_scaled])).

cnf(refute_0_14, plain,
    (X_000103 < 0 |
     ~ lgen(0,
            1/8000000 *
            (X_000103 ^ 2 + 24000000000000 * X_000103 +
             16000000000000000000000000) / (X_000103 + 4000000000000),
            X_000104) | sqrt(X_000103) <= X_000104),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (X_000103 < 0 |
     X_000104 <
     (2000000000000000000 + X_000103 * (3000000 + X_000103 * 1/8000000)) /
     (4000000000000 + X_000103) | sqrt(X_000103) <= X_000104),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     (pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     (2000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (3000000 +
       1325421053866224634595698711821825/295147905179352825856 *
       1/8000000)) /
     (4000000000000 +
      1325421053866224634595698711821825/295147905179352825856) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433) * (pi * 20) | pi * 20 <= 0 |
     -1 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 <=
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_18, plain,
    (~ lgen(0, ln(X_000047), X_000048) | ln(X_000047) <= X_000048),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(0, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     ln(X_000047) <= X_000048),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (X_000048 < -1 + X_000047 | X_000047 <= 0 | ln(X_000047) <= X_000048),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     -1 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 <= 0 |
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) <=
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20)), inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * 3/2) +
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * 20) |
     (ln(10) * (pi * 3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * 20)) / ln(10) <=
     pi * 6 + arctan(836177642653679/2278501083899584) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_24, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) * pi |
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) <=
     (pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi |
     pi <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_25, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_26, plain, (skoXC1 = 265/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 265/128 |
     skoXC1 ^ 2 = 265/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_29, plain, (skoXC1 != 265/128 | skoXC1 ^ 2 = 265/128 ^ 2),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain, (skoXC1 ^ 2 = 265/128 ^ 2),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain, (406627 * skoXC1 ^ 2 = 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (406627 * skoXC1 ^ 2 != 406627 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 265/128 ^ 2 |
     406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 ^ 2 != 265/128 ^ 2 |
     406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain, (406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_33])).

cnf(refute_0_35, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_36, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 265/128 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 265/128 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_34, refute_0_37])).

cnf(refute_0_39, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_40, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 265/128 |
     skoXC1 ^ 4 = 265/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_41, plain, (skoXC1 != 265/128 | skoXC1 ^ 4 = 265/128 ^ 4),
    inference(resolve, [], [refute_0_39, refute_0_40])).

cnf(refute_0_42, plain, (skoXC1 ^ 4 = 265/128 ^ 4),
    inference(resolve, [], [refute_0_26, refute_0_41])).

cnf(refute_0_43, plain, (14325 * skoXC1 ^ 4 = 14325 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (14325 * skoXC1 ^ 4 != 14325 * skoXC1 ^ 4 | skoXC1 ^ 4 != 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (skoXC1 ^ 4 != 265/128 ^ 4 | 14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain, (14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_46, refute_0_49])).

cnf(refute_0_51, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_52, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_55, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_50, refute_0_56])).

cnf(refute_0_58, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_38, refute_0_57])).

cnf(refute_0_59, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660),
    introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_58, refute_0_61])).

cnf(refute_0_63, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660 |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660 |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain, (169932 * skoXC1 = 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_68, plain,
    (169932 * skoXC1 != 169932 * skoXC1 | skoXC1 != 265/128 |
     169932 * skoXC1 = 169932 * 265/128),
    introduced(tautology, [equality])).

cnf(refute_0_69, plain,
    (skoXC1 != 265/128 | 169932 * skoXC1 = 169932 * 265/128),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain, (169932 * skoXC1 = 169932 * 265/128),
    inference(resolve, [], [refute_0_26, refute_0_69])).

cnf(refute_0_71, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_72, plain,
    (169932 * skoXC1 != 169932 * 265/128 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    introduced(tautology, [equality])).

cnf(refute_0_73, plain,
    (169932 * skoXC1 != 169932 * 265/128 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_71, refute_0_72])).

cnf(refute_0_74, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_70, refute_0_73])).

cnf(refute_0_75, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_76, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 265/128 |
     skoXC1 ^ 5 = 265/128 ^ 5), introduced(tautology, [equality])).

cnf(refute_0_77, plain, (skoXC1 != 265/128 | skoXC1 ^ 5 = 265/128 ^ 5),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain, (skoXC1 ^ 5 = 265/128 ^ 5),
    inference(resolve, [], [refute_0_26, refute_0_77])).

cnf(refute_0_79, plain, (3890 * skoXC1 ^ 5 = 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_80, plain,
    (3890 * skoXC1 ^ 5 != 3890 * skoXC1 ^ 5 | skoXC1 ^ 5 != 265/128 ^ 5 |
     3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_81, plain,
    (skoXC1 ^ 5 != 265/128 ^ 5 | 3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain, (3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_78, refute_0_81])).

cnf(refute_0_83, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_84, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 265/128 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_85, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 265/128 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_82, refute_0_85])).

cnf(refute_0_87, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_88, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 265/128 |
     skoXC1 ^ 3 = 265/128 ^ 3), introduced(tautology, [equality])).

cnf(refute_0_89, plain, (skoXC1 != 265/128 | skoXC1 ^ 3 = 265/128 ^ 3),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain, (skoXC1 ^ 3 = 265/128 ^ 3),
    inference(resolve, [], [refute_0_26, refute_0_89])).

cnf(refute_0_91, plain, (105247 * skoXC1 ^ 3 = 105247 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_92, plain,
    (105247 * skoXC1 ^ 3 != 105247 * skoXC1 ^ 3 |
     skoXC1 ^ 3 != 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (skoXC1 ^ 3 != 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain, (105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    inference(resolve, [], [refute_0_90, refute_0_93])).

cnf(refute_0_95, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_96, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_97, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_94, refute_0_97])).

cnf(refute_0_99, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_100, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_86, refute_0_100])).

cnf(refute_0_102, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_103, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_104, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_101, refute_0_104])).

cnf(refute_0_106, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_107, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_74, refute_0_107])).

cnf(refute_0_109, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_110, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_111, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_108, refute_0_111])).

cnf(refute_0_113, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_114, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_112, refute_0_113])).

cnf(refute_0_115, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_66, refute_0_114])).

cnf(refute_0_116, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_117, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_118, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_116, refute_0_117])).

cnf(refute_0_119, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_115, refute_0_118])).

cnf(refute_0_120, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_121, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_122, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_119, refute_0_122])).

cnf(refute_0_124, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_125, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_126, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_123, refute_0_126])).

cnf(refute_0_128, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_129, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_130, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_127, refute_0_130])).

cnf(refute_0_132, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_123, refute_0_136])).

cnf(refute_0_138, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18),
    introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_143, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_145, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 265/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_146, plain,
    (skoXC1 ^ 2 != 265/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_146])).

cnf(refute_0_148, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_149, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 265/128 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_150, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 265/128 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_147, refute_0_150])).

cnf(refute_0_152, plain, (217230815 * skoXC1 ^ 4 = 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_153, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 265/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_154, plain,
    (skoXC1 ^ 4 != 265/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain,
    (217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_154])).

cnf(refute_0_156, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_157, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 265/128 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_158, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 265/128 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_155, refute_0_158])).

cnf(refute_0_160, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_161, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 265/128 |
     skoXC1 ^ 6 = 265/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_162, plain, (skoXC1 != 265/128 | skoXC1 ^ 6 = 265/128 ^ 6),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain, (skoXC1 ^ 6 = 265/128 ^ 6),
    inference(resolve, [], [refute_0_26, refute_0_162])).

cnf(refute_0_164, plain, (3783025 * skoXC1 ^ 6 = 3783025 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_165, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_166, plain,
    (skoXC1 ^ 6 != 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_164, refute_0_165])).

cnf(refute_0_167, plain, (3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_163, refute_0_166])).

cnf(refute_0_168, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_169, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_170, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_167, refute_0_170])).

cnf(refute_0_172, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_173, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_171, refute_0_172])).

cnf(refute_0_174, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_159, refute_0_173])).

cnf(refute_0_175, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_176, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_177, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_174, refute_0_177])).

cnf(refute_0_179, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2), inference(subst, [], [refute_0_55])).

cnf(refute_0_180, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_151, refute_0_180])).

cnf(refute_0_182, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900), introduced(tautology, [refl])).

cnf(refute_0_183, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    introduced(tautology, [equality])).

cnf(refute_0_184, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_181, refute_0_184])).

cnf(refute_0_186, plain,
    (4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)), introduced(tautology, [refl])).

cnf(refute_0_187, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900 |
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    introduced(tautology, [equality])).

cnf(refute_0_188, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900 |
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_185, refute_0_188])).

cnf(refute_0_190, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_191, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 265/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_192, plain,
    (skoXC1 ^ 2 != 265/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_190, refute_0_191])).

cnf(refute_0_193, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_192])).

cnf(refute_0_194, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_195, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 265/128 ^ 2 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_196, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 265/128 ^ 2 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_193, refute_0_196])).

cnf(refute_0_198, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_199, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 265/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_200, plain,
    (skoXC1 ^ 4 != 265/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_200])).

cnf(refute_0_202, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_203, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 265/128 ^ 4 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_204, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 265/128 ^ 4 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_201, refute_0_204])).

cnf(refute_0_206, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_207, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 265/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 265/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_208, plain,
    (skoXC1 ^ 6 != 265/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_163, refute_0_208])).

cnf(refute_0_210, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_211, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 265/128 ^ 6 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_212, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 265/128 ^ 6 |
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_209, refute_0_212])).

cnf(refute_0_214, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_215, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 265/128 |
     skoXC1 ^ 8 = 265/128 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_216, plain, (skoXC1 != 265/128 | skoXC1 ^ 8 = 265/128 ^ 8),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain, (skoXC1 ^ 8 = 265/128 ^ 8),
    inference(resolve, [], [refute_0_26, refute_0_216])).

cnf(refute_0_218, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_219, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 265/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 265/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_220, plain,
    (skoXC1 ^ 8 != 265/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_217, refute_0_220])).

cnf(refute_0_222, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_223, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 265/128 ^ 8 |
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_224, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 265/128 ^ 8 |
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_221, refute_0_224])).

cnf(refute_0_226, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_227, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 265/128 |
     skoXC1 ^ 10 = 265/128 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_228, plain, (skoXC1 != 265/128 | skoXC1 ^ 10 = 265/128 ^ 10),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain, (skoXC1 ^ 10 = 265/128 ^ 10),
    inference(resolve, [], [refute_0_26, refute_0_228])).

cnf(refute_0_230, plain, (15132100 * skoXC1 ^ 10 = 15132100 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * skoXC1 ^ 10 |
     skoXC1 ^ 10 != 265/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 265/128 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_232, plain,
    (skoXC1 ^ 10 != 265/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 265/128 ^ 10),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (15132100 * skoXC1 ^ 10 = 15132100 * 265/128 ^ 10),
    inference(resolve, [], [refute_0_229, refute_0_232])).

cnf(refute_0_234, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_235, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 265/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_236, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 265/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_233, refute_0_236])).

cnf(refute_0_238, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_239, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8),
    inference(resolve, [], [refute_0_225, refute_0_239])).

cnf(refute_0_241, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_242, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_243, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_241, refute_0_242])).

cnf(refute_0_244, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_240, refute_0_243])).

cnf(refute_0_245, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6), inference(subst, [], [refute_0_55])).

cnf(refute_0_246, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_213, refute_0_246])).

cnf(refute_0_248, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_249, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_250, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_247, refute_0_250])).

cnf(refute_0_252, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_253, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_251, refute_0_252])).

cnf(refute_0_254, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_205, refute_0_253])).

cnf(refute_0_255, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_256, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_257, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_255, refute_0_256])).

cnf(refute_0_258, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_254, refute_0_257])).

cnf(refute_0_259, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2), inference(subst, [], [refute_0_55])).

cnf(refute_0_260, plain,
    (15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_258, refute_0_259])).

cnf(refute_0_261, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_197, refute_0_260])).

cnf(refute_0_262, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600), introduced(tautology, [refl])).

cnf(refute_0_263, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 + 277555600),
    introduced(tautology, [equality])).

cnf(refute_0_264, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_261, refute_0_264])).

cnf(refute_0_266, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_267, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 + 277555600 |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_268, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
     21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
     15328072984 * 265/128 ^ 2 + 277555600 |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_266, refute_0_267])).

cnf(refute_0_269, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_265, refute_0_268])).

cnf(refute_0_270, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_271, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_272, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_270, refute_0_271])).

cnf(refute_0_273, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_269, refute_0_272])).

cnf(refute_0_274, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)), introduced(tautology, [refl])).

cnf(refute_0_275, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    introduced(tautology, [equality])).

cnf(refute_0_276, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_274, refute_0_275])).

cnf(refute_0_277, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_273, refute_0_276])).

cnf(refute_0_278, plain,
    (4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_279, plain,
    (4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_277, refute_0_278])).

cnf(refute_0_280, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_189, refute_0_279])).

cnf(refute_0_281, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_282, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_283, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
          21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
          15328072984 * 265/128 ^ 2 + 277555600) /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_281, refute_0_282])).

cnf(refute_0_284, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_280, refute_0_283])).

cnf(refute_0_285, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_286, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_287, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_285, refute_0_286])).

cnf(refute_0_288, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_284, refute_0_287])).

cnf(refute_0_289, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2),
    introduced(tautology, [refl])).

cnf(refute_0_290, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    introduced(tautology, [equality])).

cnf(refute_0_291, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    inference(resolve, [], [refute_0_289, refute_0_290])).

cnf(refute_0_292, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    inference(resolve, [], [refute_0_288, refute_0_291])).

cnf(refute_0_293, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 265/128 ^ 10 + 1024027285 * 265/128 ^ 8 +
             21404723599 * 265/128 ^ 6 + 129098541721 * 265/128 ^ 4 +
             15328072984 * 265/128 ^ 2 + 277555600) /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 265/128 ^ 10 +
                        1024027285 * 265/128 ^ 8 +
                        21404723599 * 265/128 ^ 6 +
                        129098541721 * 265/128 ^ 4 +
                        15328072984 * 265/128 ^ 2 + 277555600) /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_294, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 265/128 ^ 10 +
                        1024027285 * 265/128 ^ 8 +
                        21404723599 * 265/128 ^ 6 +
                        129098541721 * 265/128 ^ 4 +
                        15328072984 * 265/128 ^ 2 + 277555600) /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_292, refute_0_293])).

cnf(refute_0_295, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 265/128 ^ 10 +
                        1024027285 * 265/128 ^ 8 +
                        21404723599 * 265/128 ^ 6 +
                        129098541721 * 265/128 ^ 4 +
                        15328072984 * 265/128 ^ 2 + 277555600) /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_294, refute_0_143])).

cnf(refute_0_296, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 265/128 ^ 10 +
                        1024027285 * 265/128 ^ 8 +
                        21404723599 * 265/128 ^ 6 +
                        129098541721 * 265/128 ^ 4 +
                        15328072984 * 265/128 ^ 2 + 277555600) /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_295, refute_0_133])).

cnf(refute_0_297, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 265/128 ^ 10 +
                        1024027285 * 265/128 ^ 8 +
                        21404723599 * 265/128 ^ 6 +
                        129098541721 * 265/128 ^ 4 +
                        15328072984 * 265/128 ^ 2 + 277555600) /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_296, refute_0_25])).

cnf(refute_0_298, plain,
    (~ interval(0,
                (pi * -6 + arctan(836177642653679/2278501083899584) * 12) /
                pi, 0,
                (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
                pi,
                (ln(10) * 3/2 +
                 ln(sqrt(1325421053866224634595698711821825/
                         295147905179352825856) *
                    747667906887680/3119868895908289175433) * 20) /
                ln(10)) | pi = 0),
    inference(arithmetic, [], [refute_0_297])).

cnf(refute_0_299, plain,
    (~ lgen(0,
            (ln(10) * 3/2 +
             ln(sqrt(1325421053866224634595698711821825/
                     295147905179352825856) *
                747667906887680/3119868895908289175433) * 20) / ln(10),
            (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi,
            (ln(10) * 3/2 +
             ln(sqrt(1325421053866224634595698711821825/
                     295147905179352825856) *
                747667906887680/3119868895908289175433) * 20) / ln(10)) |
     interval(0,
              (pi * -6 + arctan(836177642653679/2278501083899584) * 12) /
              pi, 0,
              (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
              pi,
              (ln(10) * 3/2 +
               ln(sqrt(1325421053866224634595698711821825/
                       295147905179352825856) *
                  747667906887680/3119868895908289175433) * 20) / ln(10))),
    inference(subst, [], [interval_intro])).

cnf(refute_0_300, plain,
    (~ lgen(0,
            (ln(10) * 3/2 +
             ln(sqrt(1325421053866224634595698711821825/
                     295147905179352825856) *
                747667906887680/3119868895908289175433) * 20) / ln(10),
            (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi,
            (ln(10) * 3/2 +
             ln(sqrt(1325421053866224634595698711821825/
                     295147905179352825856) *
                747667906887680/3119868895908289175433) * 20) / ln(10)) |
     pi = 0), inference(resolve, [], [refute_0_299, refute_0_298])).

cnf(refute_0_301, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) |
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) <
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi |
     pi = 0), inference(arithmetic, [], [refute_0_300])).

cnf(refute_0_302, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_303, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_304, plain, (pi != 0),
    inference(decision, [], [refute_0_26, refute_0_302, refute_0_303])).

cnf(refute_0_305, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) |
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) <
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi),
    inference(resolve, [], [refute_0_301, refute_0_304])).

cnf(refute_0_306, plain,
    (arctan(X_000040) < X_000041 | ~ lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_307, plain,
    (X_000040 < 0 | ~ lgen(1, X_000040, X_000041) |
     lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_308, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 |
     ~ lgen(1, X_000040, X_000041)),
    inference(resolve, [], [refute_0_307, refute_0_306])).

cnf(refute_0_309, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 | X_000041 <= X_000040),
    inference(arithmetic, [], [refute_0_308])).

cnf(refute_0_310, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) < pi * 13592395/5567832 |
     pi * 13592395/5567832 <= 836177642653679/2278501083899584),
    inference(subst, [], [refute_0_309])).

cnf(refute_0_311, plain,
    ((pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 |
     pi *
     30392992063010096599082701152059392/
     7874467779473709920279975121876516072046628601 +
     arctan(836177642653679/2278501083899584) *
     -2305524983369630620753898890067968/
     1458234773976612948199995392940095568897523815 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_312, plain,
    (X_000118 < sqrt(X_000117) | ~ lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_313, plain,
    (X_000117 < 0 |
     ~ lgen(1, X_000118,
            8000000 * X_000117 * (X_000117 + 4000000000000) /
            (X_000117 ^ 2 + 24000000000000 * X_000117 +
             16000000000000000000000000)) |
     lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_314, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     ~ lgen(1, X_000118,
            8000000 * X_000117 * (X_000117 + 4000000000000) /
            (X_000117 ^ 2 + 24000000000000 * X_000117 +
             16000000000000000000000000))),
    inference(resolve, [], [refute_0_313, refute_0_312])).

cnf(refute_0_315, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     X_000117 * (32000000000000000000 + X_000117 * 8000000) /
     (16000000000000000000000000 +
      X_000117 * (24000000000000 + X_000117)) <= X_000118),
    inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(subst, [], [refute_0_315])).

cnf(refute_0_317, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825),
    introduced(tautology, [assume])).

cnf(refute_0_318, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi *
     30392992063010096599082701152059392/
     7874467779473709920279975121876516072046628601 +
     arctan(836177642653679/2278501083899584) *
     -2305524983369630620753898890067968/
     1458234773976612948199995392940095568897523815 <= 0 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(split, [], [refute_0_317])).

cnf(refute_0_319, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi * 13592395/5567832 <= arctan(836177642653679/2278501083899584) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(arithmetic, [], [refute_0_318])).

cnf(refute_0_320, plain,
    (arctan(836177642653679/2278501083899584) < pi * 13592395/5567832 |
     pi <= 581962079556464856741/3871285842536410757960),
    inference(arithmetic, [], [refute_0_310])).

cnf(refute_0_321, plain,
    (581962079556464856741/3871285842536410757960 < pi),
    inference(decision, [], [refute_0_303, refute_0_302, refute_0_26])).

cnf(refute_0_322, plain,
    (arctan(836177642653679/2278501083899584) < pi * 13592395/5567832),
    inference(resolve, [], [refute_0_320, refute_0_321])).

cnf(refute_0_323, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(resolve, [], [refute_0_319, refute_0_322])).

cnf(refute_0_324, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(resolve, [], [refute_0_323, refute_0_316])).

cnf(refute_0_325, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 <=
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825) /
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(arithmetic, [], [refute_0_324])).

cnf(refute_0_326, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 |
     pi *
     30392992063010096599082701152059392/
     7874467779473709920279975121876516072046628601 +
     arctan(836177642653679/2278501083899584) *
     -2305524983369630620753898890067968/
     1458234773976612948199995392940095568897523815 <= 0),
    inference(resolve, [], [refute_0_325, refute_0_311])).

cnf(refute_0_327, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi *
     57906967292335754501879804210975345143130949399924887346515987545331455/
     140070595474597260989405683839026020401891798426393059590131589435418328 <=
     arctan(836177642653679/2278501083899584) |
     pi * 13592395/5567832 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_326])).

cnf(refute_0_328, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) <
     pi *
     57906967292335754501879804210975345143130949399924887346515987545331455/
     140070595474597260989405683839026020401891798426393059590131589435418328 |
     pi *
     57906967292335754501879804210975345143130949399924887346515987545331455/
     140070595474597260989405683839026020401891798426393059590131589435418328 <=
     836177642653679/2278501083899584),
    inference(subst, [], [refute_0_309])).

cnf(refute_0_329, plain,
    (arctan(836177642653679/2278501083899584) <
     pi *
     57906967292335754501879804210975345143130949399924887346515987545331455/
     140070595474597260989405683839026020401891798426393059590131589435418328 |
     pi <=
     14640487541130726921627386832526000415313558292201111004600988465454729505147174153589/
     16492635967615596937107978351044290773281087887973170990333962969697558743930652076840),
    inference(arithmetic, [], [refute_0_328])).

cnf(refute_0_330, plain,
    (14640487541130726921627386832526000415313558292201111004600988465454729505147174153589/
     16492635967615596937107978351044290773281087887973170990333962969697558743930652076840 <
     pi),
    inference(decision, [], [refute_0_26, refute_0_302, refute_0_303])).

cnf(refute_0_331, plain,
    (arctan(836177642653679/2278501083899584) <
     pi *
     57906967292335754501879804210975345143130949399924887346515987545331455/
     140070595474597260989405683839026020401891798426393059590131589435418328),
    inference(resolve, [], [refute_0_329, refute_0_330])).

cnf(refute_0_332, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi * 13592395/5567832 <= arctan(836177642653679/2278501083899584)),
    inference(resolve, [], [refute_0_327, refute_0_331])).

cnf(refute_0_333, plain,
    (836177642653679/2278501083899584 < 0 |
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi * 13592395/5567832 <= 836177642653679/2278501083899584),
    inference(resolve, [], [refute_0_332, refute_0_310])).

cnf(refute_0_334, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi <= 581962079556464856741/3871285842536410757960),
    inference(arithmetic, [], [refute_0_333])).

cnf(refute_0_335, plain,
    (pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(resolve, [], [refute_0_334, refute_0_321])).

cnf(refute_0_336, plain,
    (pi * -10 <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     0 <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <=
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_337, plain,
    (~ lgen(0, X_000118, sqrt(X_000117)) | X_000118 <= sqrt(X_000117)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_338, plain,
    (X_000117 < 0 |
     ~ lgen(0, X_000118,
            8000000 * X_000117 * (X_000117 + 4000000000000) /
            (X_000117 ^ 2 + 24000000000000 * X_000117 +
             16000000000000000000000000)) |
     lgen(0, X_000118, sqrt(X_000117))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_339, plain,
    (X_000117 < 0 |
     ~ lgen(0, X_000118,
            8000000 * X_000117 * (X_000117 + 4000000000000) /
            (X_000117 ^ 2 + 24000000000000 * X_000117 +
             16000000000000000000000000)) | X_000118 <= sqrt(X_000117)),
    inference(resolve, [], [refute_0_338, refute_0_337])).

cnf(refute_0_340, plain,
    (X_000117 * (32000000000000000000 + X_000117 * 8000000) /
     (16000000000000000000000000 +
      X_000117 * (24000000000000 + X_000117)) < X_000118 | X_000117 < 0 |
     X_000118 <= sqrt(X_000117)),
    inference(arithmetic, [], [refute_0_339])).

cnf(refute_0_341, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <=
     sqrt(1325421053866224634595698711821825/295147905179352825856)),
    inference(subst, [], [refute_0_340])).

cnf(refute_0_342, plain,
    (pi * -10 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * -29906716275507200/3119868895908289175433 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       27950364949490226519206789120000/
       9733581927656007318022218939278759050737489)) <
     231993/101125 *
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     0 <=
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) <= 231993/101125),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_343, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) | 10 <= 0 |
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) <= ln(10)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_344, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (-1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) * (pi * -20) |
     0 <= pi * -20 |
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) <=
     (-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (-1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_345, plain,
    (~ lgen(0, X_000125, ln(X_000124)) | X_000125 <= ln(X_000124)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_346, plain,
    (~ lgen(0, X_000125,
            1/2 * (1 + 5 * X_000124) * (X_000124 - 1) /
            (X_000124 * (2 + X_000124))) | X_000124 <= 0 |
     lgen(0, X_000125, ln(X_000124))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_347, plain,
    (~ lgen(0, X_000125,
            1/2 * (1 + 5 * X_000124) * (X_000124 - 1) /
            (X_000124 * (2 + X_000124))) | X_000124 <= 0 |
     X_000125 <= ln(X_000124)),
    inference(resolve, [], [refute_0_346, refute_0_345])).

cnf(refute_0_348, plain,
    ((-1/2 + X_000124 * (-2 + X_000124 * 5/2)) /
     (X_000124 * (2 + X_000124)) < X_000125 | X_000124 <= 0 |
     X_000125 <= ln(X_000124)), inference(arithmetic, [], [refute_0_347])).

cnf(refute_0_349, plain,
    ((-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433 *
      (-2 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       747667906887680/3119868895908289175433 * 5/2)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433 *
      (2 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       747667906887680/3119868895908289175433)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 <= 0 |
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) <=
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433)),
    inference(subst, [], [refute_0_348])).

cnf(refute_0_350, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * -3/2) +
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * -20) |
     (ln(10) * (pi * -3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * -20)) / ln(10) <=
     pi * 6 + arctan(836177642653679/2278501083899584) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_351, plain,
    ((ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) * pi <
     pi * -6 + arctan(836177642653679/2278501083899584) * 12 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) | pi <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_352, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * -3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * -20)) / ln(10) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_351])).

cnf(refute_0_353, plain, (0 < pi),
    inference(decision, [], [refute_0_26, refute_0_302, refute_0_303])).

cnf(refute_0_354, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * -3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * -20)) / ln(10) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_352, refute_0_353])).

cnf(refute_0_355, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * -3/2) +
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) |
     ln(10) <= 0), inference(resolve, [], [refute_0_350, refute_0_354])).

cnf(refute_0_356, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_355])).

cnf(refute_0_357, plain,
    (ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_356])).

cnf(refute_0_358, plain,
    (ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_357])).

cnf(refute_0_359, plain,
    (ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_358, refute_0_353])).

cnf(refute_0_360, plain,
    ((-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433 *
      (-2 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       747667906887680/3119868895908289175433 * 5/2)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433 *
      (2 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       747667906887680/3119868895908289175433)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_349, refute_0_359])).

cnf(refute_0_361, plain,
    ((-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (-1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_360])).

cnf(refute_0_362, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1/2 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (-1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) * (pi * -20) |
     0 <= pi * -20 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_344, refute_0_361])).

cnf(refute_0_363, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (pi * 10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * 29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_362])).

cnf(refute_0_364, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (pi * 10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * 29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (1495335813775360/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_363, refute_0_353])).

cnf(refute_0_365, plain,
    (ln(10) <
     -((pi * 10 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi * 29906716275507200/3119868895908289175433 +
        sqrt(1325421053866224634595698711821825/295147905179352825856) *
        (pi *
         -27950364949490226519206789120000/
         9733581927656007318022218939278759050737489))) /
      (sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (1495335813775360/3119868895908289175433 +
        sqrt(1325421053866224634595698711821825/295147905179352825856) *
        559007298989804530384135782400/
        9733581927656007318022218939278759050737489))) /
     (pi * -15/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -15/2 + arctan(836177642653679/2278501083899584) * 12 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_364])).

cnf(refute_0_366, plain,
    (ln(10) <
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     pi * 5/8 <= arctan(836177642653679/2278501083899584) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_365])).

cnf(refute_0_367, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) < pi * 5/8 |
     pi * 5/8 <= 836177642653679/2278501083899584),
    inference(subst, [], [refute_0_309])).

cnf(refute_0_368, plain,
    (arctan(836177642653679/2278501083899584) < pi * 5/8 |
     pi <= 836177642653679/1424063177437240),
    inference(arithmetic, [], [refute_0_367])).

cnf(refute_0_369, plain, (836177642653679/1424063177437240 < pi),
    inference(decision, [], [refute_0_303, refute_0_302, refute_0_26])).

cnf(refute_0_370, plain,
    (arctan(836177642653679/2278501083899584) < pi * 5/8),
    inference(resolve, [], [refute_0_368, refute_0_369])).

cnf(refute_0_371, plain,
    (ln(10) <
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_366, refute_0_370])).

cnf(refute_0_372, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) | 10 <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_343, refute_0_371])).

cnf(refute_0_373, plain,
    (231993/101125 <
     (pi * -10 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -29906716275507200/3119868895908289175433 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_372])).

cnf(refute_0_374, plain,
    (pi * -10 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * -29906716275507200/3119868895908289175433 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       27950364949490226519206789120000/
       9733581927656007318022218939278759050737489)) <
     231993/101125 *
     (sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi * -3738339534438400/1039956298636096391811 +
       arctan(836177642653679/2278501083899584) *
       5981343255101440/1039956298636096391811 +
       sqrt(1325421053866224634595698711821825/295147905179352825856) *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163 +
        arctan(836177642653679/2278501083899584) *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     0 <=
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_342, refute_0_373])).

cnf(refute_0_375, plain,
    (pi * -10 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    ((pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     pi * -3738339534438400/1039956298636096391811 +
     arctan(836177642653679/2278501083899584) *
     5981343255101440/1039956298636096391811 |
     pi *
     1397518247474511325960339456000/
     3244527309218669106007406313092919683579163 +
     arctan(836177642653679/2278501083899584) *
     -2236029195959218121536543129600/
     3244527309218669106007406313092919683579163 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_377, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(subst, [], [refute_0_315])).

cnf(refute_0_378, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_379, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi * 3738339534438400/1039956298636096391811 +
     arctan(836177642653679/2278501083899584) *
     -5981343255101440/1039956298636096391811 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <= 0 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     0 /
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163))),
    inference(split, [], [refute_0_378])).

cnf(refute_0_380, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     pi * -3738339534438400/1039956298636096391811 +
     arctan(836177642653679/2278501083899584) *
     5981343255101440/1039956298636096391811),
    inference(arithmetic, [], [refute_0_379])).

cnf(refute_0_381, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi *
     1397518247474511325960339456000/
     3244527309218669106007406313092919683579163 +
     arctan(836177642653679/2278501083899584) *
     -2236029195959218121536543129600/
     3244527309218669106007406313092919683579163 <= 0 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(split, [], [refute_0_380])).

cnf(refute_0_382, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi * 5/8 <= arctan(836177642653679/2278501083899584) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(arithmetic, [], [refute_0_381])).

cnf(refute_0_383, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(resolve, [], [refute_0_382, refute_0_370])).

cnf(refute_0_384, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(resolve, [], [refute_0_383, refute_0_377])).

cnf(refute_0_385, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 <=
     (pi * -3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      5981343255101440/1039956298636096391811) /
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163)),
    inference(arithmetic, [], [refute_0_384])).

cnf(refute_0_386, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163 +
      arctan(836177642653679/2278501083899584) *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     pi * -3738339534438400/1039956298636096391811 +
     arctan(836177642653679/2278501083899584) *
     5981343255101440/1039956298636096391811 |
     pi *
     1397518247474511325960339456000/
     3244527309218669106007406313092919683579163 +
     arctan(836177642653679/2278501083899584) *
     -2236029195959218121536543129600/
     3244527309218669106007406313092919683579163 <= 0),
    inference(resolve, [], [refute_0_385, refute_0_376])).

cnf(refute_0_387, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi * 5/8 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_386])).

cnf(refute_0_388, plain,
    (836177642653679/2278501083899584 < 0 |
     0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi * 5/8 <= 836177642653679/2278501083899584),
    inference(resolve, [], [refute_0_387, refute_0_367])).

cnf(refute_0_389, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     pi <= 836177642653679/1424063177437240),
    inference(arithmetic, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (0 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 3738339534438400/1039956298636096391811 +
      arctan(836177642653679/2278501083899584) *
      -5981343255101440/1039956298636096391811 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       1397518247474511325960339456000/
       3244527309218669106007406313092919683579163 +
       arctan(836177642653679/2278501083899584) *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163))),
    inference(resolve, [], [refute_0_389, refute_0_369])).

cnf(refute_0_391, plain,
    (pi * -10 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_375, refute_0_390])).

cnf(refute_0_392, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     0 <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_392])).

cnf(refute_0_394, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_393, refute_0_335])).

cnf(refute_0_395, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_341, refute_0_394])).

cnf(refute_0_396, plain,
    (42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 <
     pi * -10 /
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_395])).

cnf(refute_0_397, plain,
    (pi * -10 <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi * 16900434900870496256/12619869683949029714626485 +
      arctan(836177642653679/2278501083899584) *
      92508651052049891328/7011038713305016508125825 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      (pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601 +
       arctan(836177642653679/2278501083899584) *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     0 <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_336, refute_0_396])).

cnf(refute_0_398, plain,
    (pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (836177642653679/2278501083899584 < 0 |
     pi *
     5420893806341993382637297203529332039166492333159280732126159724509478579931604856185112568291394355492815305643482432193293244053/
     42285618809853058977549136289423924650674202321052667037531674864681743287963298623375817152997086009900265236309732823100555264000 <
     836177642653679/2278501083899584 |
     arctan(836177642653679/2278501083899584) <=
     pi *
     5420893806341993382637297203529332039166492333159280732126159724509478579931604856185112568291394355492815305643482432193293244053/
     42285618809853058977549136289423924650674202321052667037531674864681743287963298623375817152997086009900265236309732823100555264000),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_400, plain,
    (pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361 <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     0 <=
     pi *
     -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
     50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
     arctan(836177642653679/2278501083899584) *
     18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
     5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257 |
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_401, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     sqrt(1325421053866224634595698711821825/295147905179352825856)),
    inference(subst, [], [refute_0_340])).

cnf(refute_0_402, plain,
    (pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361 <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    introduced(tautology, [assume])).

cnf(refute_0_403, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     0 <=
     pi *
     -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
     50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
     arctan(836177642653679/2278501083899584) *
     18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
     5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(split, [], [refute_0_402])).

cnf(refute_0_404, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     pi * 13592395/5567832 <= arctan(836177642653679/2278501083899584) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(arithmetic, [], [refute_0_403])).

cnf(refute_0_405, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_404, refute_0_322])).

cnf(refute_0_406, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     1325421053866224634595698711821825/295147905179352825856 *
     (32000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 * 8000000) /
     (16000000000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (24000000000000 +
       1325421053866224634595698711821825/295147905179352825856)) <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_401, refute_0_405])).

cnf(refute_0_407, plain,
    (42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 <
     (pi *
      -208369267434883597078022827337544726537100348452271858329677836503610/
      16230770879814293130625887480036392792245609537198990875168851170361 +
      arctan(836177642653679/2278501083899584) *
      -453832509263802398820275933468995036873097815302452478535343800320000/
      16230770879814293130625887480036392792245609537198990875168851170361) /
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(arithmetic, [], [refute_0_406])).

cnf(refute_0_408, plain,
    (pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361 <
     42515481089904115134363026767675203809038897329161094502534993032000000/
     20062757577026320309797141508079595540476649613348567212816874129 *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) |
     0 <=
     pi *
     -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
     50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
     arctan(836177642653679/2278501083899584) *
     18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
     5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_400, refute_0_407])).

cnf(refute_0_409, plain,
    (pi *
     5420893806341993382637297203529332039166492333159280732126159724509478579931604856185112568291394355492815305643482432193293244053/
     42285618809853058977549136289423924650674202321052667037531674864681743287963298623375817152997086009900265236309732823100555264000 <
     arctan(836177642653679/2278501083899584) |
     pi * 13592395/5567832 <= arctan(836177642653679/2278501083899584) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(arithmetic, [], [refute_0_408])).

cnf(refute_0_410, plain,
    (pi *
     5420893806341993382637297203529332039166492333159280732126159724509478579931604856185112568291394355492815305643482432193293244053/
     42285618809853058977549136289423924650674202321052667037531674864681743287963298623375817152997086009900265236309732823100555264000 <
     arctan(836177642653679/2278501083899584) |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_409, refute_0_322])).

cnf(refute_0_411, plain,
    (836177642653679/2278501083899584 < 0 |
     pi *
     5420893806341993382637297203529332039166492333159280732126159724509478579931604856185112568291394355492815305643482432193293244053/
     42285618809853058977549136289423924650674202321052667037531674864681743287963298623375817152997086009900265236309732823100555264000 <
     836177642653679/2278501083899584 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_399, refute_0_410])).

cnf(refute_0_412, plain,
    (pi <
     552473266477734347505311022828319823645637028887137624691839808204483076200762114002637467931967771771618353536525635204311488618964927381504000/
     192992381460230836321910482024388200641362216117752621854425722235231372777384351712422360285551278175067380820614054649310612846773676860268343 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(arithmetic, [], [refute_0_411])).

cnf(refute_0_413, plain,
    (552473266477734347505311022828319823645637028887137624691839808204483076200762114002637467931967771771618353536525635204311488618964927381504000/
     192992381460230836321910482024388200641362216117752621854425722235231372777384351712422360285551278175067380820614054649310612846773676860268343 <=
     pi),
    inference(decision, [], [refute_0_303, refute_0_302, refute_0_26])).

cnf(refute_0_414, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      -414175313909888084614911483816013105032812935472003596170312590870438916128768000000/
      50637877224546630013720350482519418838975367853522946219785078801740436685268089698941313 +
      arctan(836177642653679/2278501083899584) *
      18850889015491710163723593939096743098058852535103937221039525684532806366003200000/
      5626430802727403334857816720279935426552818650391438468865008755748937409474232188771257) <=
     pi *
     -208369267434883597078022827337544726537100348452271858329677836503610/
     16230770879814293130625887480036392792245609537198990875168851170361 +
     arctan(836177642653679/2278501083899584) *
     -453832509263802398820275933468995036873097815302452478535343800320000/
     16230770879814293130625887480036392792245609537198990875168851170361),
    inference(resolve, [], [refute_0_413, refute_0_412])).

cnf(refute_0_415, plain,
    (sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi *
      30392992063010096599082701152059392/
      7874467779473709920279975121876516072046628601 +
      arctan(836177642653679/2278501083899584) *
      -2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <=
     pi * 16900434900870496256/12619869683949029714626485 +
     arctan(836177642653679/2278501083899584) *
     92508651052049891328/7011038713305016508125825 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_414, refute_0_398])).

cnf(refute_0_416, plain,
    ((pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_415, refute_0_335])).

cnf(refute_0_417, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_416, refute_0_305])).

cnf(refute_0_418, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * 3/2 +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * 20) / ln(10) * pi |
     pi <= 0), inference(resolve, [], [refute_0_24, refute_0_417])).

cnf(refute_0_419, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * 3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * 20)) / ln(10) |
     pi <= 0), inference(arithmetic, [], [refute_0_418])).

cnf(refute_0_420, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * 3/2) +
      ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
         747667906887680/3119868895908289175433) * (pi * 20)) / ln(10)),
    inference(resolve, [], [refute_0_419, refute_0_353])).

cnf(refute_0_421, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * 3/2) +
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * 20) | ln(10) <= 0),
    inference(resolve, [], [refute_0_23, refute_0_420])).

cnf(refute_0_422, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) * (pi * 20)),
    inference(arithmetic, [], [refute_0_421])).

cnf(refute_0_423, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) | pi * 20 <= 0),
    inference(split, [], [refute_0_422])).

cnf(refute_0_424, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433) | pi <= 0),
    inference(arithmetic, [], [refute_0_423])).

cnf(refute_0_425, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     ln(sqrt(1325421053866224634595698711821825/295147905179352825856) *
        747667906887680/3119868895908289175433)),
    inference(resolve, [], [refute_0_424, refute_0_353])).

cnf(refute_0_426, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     -1 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 |
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433 <= 0),
    inference(resolve, [], [refute_0_22, refute_0_425])).

cnf(refute_0_427, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) <
     -1 +
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     747667906887680/3119868895908289175433),
    inference(arithmetic, [], [refute_0_426])).

cnf(refute_0_428, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1 +
      sqrt(1325421053866224634595698711821825/295147905179352825856) *
      747667906887680/3119868895908289175433) * (pi * 20) | pi * 20 <= 0),
    inference(resolve, [], [refute_0_17, refute_0_427])).

cnf(refute_0_429, plain,
    (pi * 20 +
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 14953358137753600/3119868895908289175433) | pi <= 0),
    inference(arithmetic, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (pi * 20 +
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) *
     (pi * 14953358137753600/3119868895908289175433)),
    inference(resolve, [], [refute_0_429, refute_0_353])).

cnf(refute_0_431, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) |
     pi * 14953358137753600/3119868895908289175433 <= 0),
    inference(split, [], [refute_0_430])).

cnf(refute_0_432, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     sqrt(1325421053866224634595698711821825/295147905179352825856) |
     pi <= 0), inference(arithmetic, [], [refute_0_431])).

cnf(refute_0_433, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     sqrt(1325421053866224634595698711821825/295147905179352825856)),
    inference(resolve, [], [refute_0_432, refute_0_353])).

cnf(refute_0_434, plain,
    (1325421053866224634595698711821825/295147905179352825856 < 0 |
     (pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     (2000000000000000000 +
      1325421053866224634595698711821825/295147905179352825856 *
      (3000000 +
       1325421053866224634595698711821825/295147905179352825856 *
       1/8000000)) /
     (4000000000000 +
      1325421053866224634595698711821825/295147905179352825856)),
    inference(resolve, [], [refute_0_16, refute_0_433])).

cnf(refute_0_435, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12)) /
     (pi * 14953358137753600/3119868895908289175433) <
     20062757577026320309797141508079595540476649613348567212816874129/
     9467448208080222069846420188471815866695581538881372160000),
    inference(arithmetic, [], [refute_0_434])).

cnf(refute_0_436, plain,
    (pi * 20 +
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     20062757577026320309797141508079595540476649613348567212816874129/
     9467448208080222069846420188471815866695581538881372160000 *
     (pi * 14953358137753600/3119868895908289175433) |
     pi * 14953358137753600/3119868895908289175433 <= 0),
    inference(resolve, [], [refute_0_11, refute_0_435])).

cnf(refute_0_437, plain,
    (pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 <
     ln(10) * (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     pi <= 0), inference(arithmetic, [], [refute_0_436])).

cnf(refute_0_438, plain,
    (pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 <
     ln(10) * (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(resolve, [], [refute_0_437, refute_0_353])).

cnf(refute_0_439, plain,
    (ln(10) <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12),
    inference(split, [], [refute_0_438])).

cnf(refute_0_440, plain,
    (ln(10) <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     pi * 3/8 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_439])).

cnf(refute_0_441, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) < pi * 3/8 |
     pi * 3/8 <= 836177642653679/2278501083899584),
    inference(subst, [], [refute_0_309])).

cnf(refute_0_442, plain,
    (arctan(836177642653679/2278501083899584) < pi * 3/8 |
     pi <= 836177642653679/854437906462344),
    inference(arithmetic, [], [refute_0_441])).

cnf(refute_0_443, plain, (836177642653679/854437906462344 < pi),
    inference(decision, [], [refute_0_303, refute_0_302, refute_0_26])).

cnf(refute_0_444, plain,
    (arctan(836177642653679/2278501083899584) < pi * 3/8),
    inference(resolve, [], [refute_0_442, refute_0_443])).

cnf(refute_0_445, plain,
    (ln(10) <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(resolve, [], [refute_0_440, refute_0_444])).

cnf(refute_0_446, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     10 <= 0), inference(resolve, [], [refute_0_10, refute_0_445])).

cnf(refute_0_447, plain,
    (231993/101125 <
     pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(arithmetic, [], [refute_0_446])).

cnf(refute_0_448, plain,
    (pi *
     330531229883534751204300800875837684990022151744938511783257971807/
     33579905434649109823542610325659540458906259758593207720057241600 <
     231993/101125 *
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12),
    inference(resolve, [], [refute_0_5, refute_0_447])).

cnf(refute_0_449, plain,
    (pi *
     15833834480630041832220569266204020314032293580577068249197350480063/
     21614713530174939011217907414420533002588781281411275945246445273088 <
     arctan(836177642653679/2278501083899584) |
     pi * 3/8 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_448])).

cnf(refute_0_450, plain,
    (pi *
     15833834480630041832220569266204020314032293580577068249197350480063/
     21614713530174939011217907414420533002588781281411275945246445273088 <
     arctan(836177642653679/2278501083899584)),
    inference(resolve, [], [refute_0_449, refute_0_444])).

cnf(refute_0_451, plain,
    (836177642653679/2278501083899584 < 0 |
     pi *
     15833834480630041832220569266204020314032293580577068249197350480063/
     21614713530174939011217907414420533002588781281411275945246445273088 <
     836177642653679/2278501083899584),
    inference(resolve, [], [refute_0_4, refute_0_450])).

cnf(refute_0_452, plain,
    (pi <
     282402190723379073056019145748099821613107290939741423614986486967898765596295168/
     563709516037533703038640108117479440113210287976519133736733098647501011601343653),
    inference(arithmetic, [], [refute_0_451])).

cnf(refute_0_453, plain,
    (282402190723379073056019145748099821613107290939741423614986486967898765596295168/
     563709516037533703038640108117479440113210287976519133736733098647501011601343653 <=
     pi),
    inference(decision, [], [refute_0_303, refute_0_302, refute_0_26])).

cnf(refute_0_454, plain, ($false),
    inference(resolve, [], [refute_0_453, refute_0_452])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-6-sqrt.tptp
