% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-5-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(atan_upper_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, 3 * X / (X ^ 2 + 3), Y) | lgen(R, arctan(X), Y))).

cnf(ln_lower_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/3 * (10 * X ^ 2 + 19 * X + 1) * (X - 1) /
            (X * (X ^ 2 + 6 * X + 3))) | lgen(R, Y, ln(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Fails_1_5', conjecture,
    (! [X] :
       (X = 79/64 =>
        12 / pi *
        arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
               (14325 * X ^ 4 + 406627 * X ^ 2 - 16660)) - 6 <
        4347/1000 *
        ln(4250 ^ 2 *
           (15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
            129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) /
           (3783025 * X ^ 6 + 217230815 * X ^ 4 + 3120782021 * X ^ 2 +
            69388900) ^ 2)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 79/64 =>
        12 / pi *
        arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
               (14325 * X ^ 4 + 406627 * X ^ 2 - 16660)) - 6 <
        4347/1000 *
        ln(4250 ^ 2 *
           (15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
            129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) /
           (3783025 * X ^ 6 + 217230815 * X ^ 4 + 3120782021 * X ^ 2 +
            69388900) ^ 2))),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Fails_1_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 79/64 =>
          12 / pi *
          arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                 (14325 * X ^ 4 + 406627 * X ^ 2 - 16660)) - 6 <
          4347/1000 *
          ln(4250 ^ 2 *
             (15132100 * X ^ 10 + 1024027285 * X ^ 8 +
              21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
              15328072984 * X ^ 2 + 277555600) /
             (3783025 * X ^ 6 + 217230815 * X ^ 4 + 3120782021 * X ^ 2 +
              69388900) ^ 2))), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (4347/1000 *
        ln(4250 ^ 2 *
           (15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
            129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) /
           (3783025 * X ^ 6 + 217230815 * X ^ 4 + 3120782021 * X ^ 2 +
            69388900) ^ 2) <=
        12 / pi *
        arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
               (14325 * X ^ 4 + 406627 * X ^ 2 - 16660)) - 6 & X = 79/64)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 &
     skoXC1 = 79/64), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 79/64),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (arctan(X_000082) < X_000083 | ~ lgen(1, arctan(X_000082), X_000083)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, 3 * X_000082 / (X_000082 ^ 2 + 3), X_000083) |
     0 <= X_000082 | lgen(1, arctan(X_000082), X_000083)),
    inference(subst, [], [atan_upper_bound_case_12])).

cnf(refute_0_2, plain,
    (arctan(X_000082) < X_000083 |
     ~ lgen(1, 3 * X_000082 / (X_000082 ^ 2 + 3), X_000083) |
     0 <= X_000082), inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (arctan(X_000082) < X_000083 |
     X_000083 <= X_000082 * 3 / (3 + X_000082 * X_000082) | 0 <= X_000082),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (arctan(-118851281795/113846536469728) <
     pi *
     705272101716410751172304370869866194715732610466925479967558157/
     233308080218443270775085651985034987714775078176957673472000000000 |
     0 <= -118851281795/113846536469728 |
     pi *
     705272101716410751172304370869866194715732610466925479967558157/
     233308080218443270775085651985034987714775078176957673472000000000 <=
     -118851281795/113846536469728 * 3 /
     (3 + -118851281795/113846536469728 * -118851281795/113846536469728)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000) <
     -240059561092143322228510396732189705305738729768140603014559921/
     174981060163832453081314238988776240786081308632718255104000000 |
     -240059561092143322228510396732189705305738729768140603014559921/
     174981060163832453081314238988776240786081308632718255104000000 *
     (pi * 4347/1000) <=
     pi * -6 + arctan(-118851281795/113846536469728) * 12 |
     pi * 4347/1000 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain, (skoXC1 = 79/64),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_8, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_9, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 79/64 | skoXC1 ^ 2 = 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_10, plain, (skoXC1 != 79/64 | skoXC1 ^ 2 = 79/64 ^ 2),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain, (skoXC1 ^ 2 = 79/64 ^ 2),
    inference(resolve, [], [refute_0_7, refute_0_10])).

cnf(refute_0_12, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_13, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 79/64 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_14, plain,
    (skoXC1 ^ 2 != 79/64 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_12, refute_0_13])).

cnf(refute_0_15, plain, (3120782021 * skoXC1 ^ 2 = 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_14])).

cnf(refute_0_16, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_17, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 79/64 ^ 2 |
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_18, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 79/64 ^ 2 |
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_21, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 79/64 | skoXC1 ^ 4 = 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_22, plain, (skoXC1 != 79/64 | skoXC1 ^ 4 = 79/64 ^ 4),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 ^ 4 = 79/64 ^ 4),
    inference(resolve, [], [refute_0_7, refute_0_22])).

cnf(refute_0_24, plain, (217230815 * skoXC1 ^ 4 = 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_25, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 79/64 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_26, plain,
    (skoXC1 ^ 4 != 79/64 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain, (217230815 * skoXC1 ^ 4 = 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_26])).

cnf(refute_0_28, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_29, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 79/64 ^ 4 |
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_30, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 79/64 ^ 4 |
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_28, refute_0_29])).

cnf(refute_0_31, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_27, refute_0_30])).

cnf(refute_0_32, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_33, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 79/64 | skoXC1 ^ 6 = 79/64 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_34, plain, (skoXC1 != 79/64 | skoXC1 ^ 6 = 79/64 ^ 6),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain, (skoXC1 ^ 6 = 79/64 ^ 6),
    inference(resolve, [], [refute_0_7, refute_0_34])).

cnf(refute_0_36, plain, (3783025 * skoXC1 ^ 6 = 3783025 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_37, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 79/64 ^ 6 | 3783025 * skoXC1 ^ 6 = 3783025 * 79/64 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_38, plain,
    (skoXC1 ^ 6 != 79/64 ^ 6 | 3783025 * skoXC1 ^ 6 = 3783025 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain, (3783025 * skoXC1 ^ 6 = 3783025 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_35, refute_0_38])).

cnf(refute_0_40, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_41, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 79/64 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_42, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 79/64 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_39, refute_0_42])).

cnf(refute_0_44, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_45, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_46, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_43, refute_0_49])).

cnf(refute_0_51, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_31, refute_0_50])).

cnf(refute_0_52, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_53, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_54, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_51, refute_0_54])).

cnf(refute_0_56, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_57, plain,
    (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_19, refute_0_57])).

cnf(refute_0_59, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900), introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 +
     69388900), introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 +
     69388900), inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 +
     69388900), inference(resolve, [], [refute_0_58, refute_0_61])).

cnf(refute_0_63, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 +
     69388900 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 + 3120782021 * 79/64 ^ 2 +
     69388900 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain,
    (4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_68, plain,
    (4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_69, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_66, refute_0_69])).

cnf(refute_0_71, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_72, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 79/64 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_73, plain,
    (skoXC1 ^ 2 != 79/64 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_71, refute_0_72])).

cnf(refute_0_74, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_73])).

cnf(refute_0_75, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_76, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 79/64 ^ 2 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_77, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 79/64 ^ 2 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_74, refute_0_77])).

cnf(refute_0_79, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_80, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 79/64 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_81, plain,
    (skoXC1 ^ 4 != 79/64 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_81])).

cnf(refute_0_83, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_84, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 79/64 ^ 4 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_85, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 79/64 ^ 4 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_82, refute_0_85])).

cnf(refute_0_87, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_88, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 79/64 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 79/64 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_89, plain,
    (skoXC1 ^ 6 != 79/64 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_35, refute_0_89])).

cnf(refute_0_91, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_92, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 79/64 ^ 6 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 79/64 ^ 6 |
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_90, refute_0_93])).

cnf(refute_0_95, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_96, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 79/64 | skoXC1 ^ 8 = 79/64 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_97, plain, (skoXC1 != 79/64 | skoXC1 ^ 8 = 79/64 ^ 8),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain, (skoXC1 ^ 8 = 79/64 ^ 8),
    inference(resolve, [], [refute_0_7, refute_0_97])).

cnf(refute_0_99, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_100, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 79/64 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 79/64 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_101, plain,
    (skoXC1 ^ 8 != 79/64 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_98, refute_0_101])).

cnf(refute_0_103, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 79/64 ^ 8 |
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 79/64 ^ 8 |
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_102, refute_0_105])).

cnf(refute_0_107, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 79/64 |
     skoXC1 ^ 10 = 79/64 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_109, plain, (skoXC1 != 79/64 | skoXC1 ^ 10 = 79/64 ^ 10),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain, (skoXC1 ^ 10 = 79/64 ^ 10),
    inference(resolve, [], [refute_0_7, refute_0_109])).

cnf(refute_0_111, plain, (15132100 * skoXC1 ^ 10 = 15132100 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * skoXC1 ^ 10 |
     skoXC1 ^ 10 != 79/64 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 79/64 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoXC1 ^ 10 != 79/64 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 79/64 ^ 10),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain, (15132100 * skoXC1 ^ 10 = 15132100 * 79/64 ^ 10),
    inference(resolve, [], [refute_0_110, refute_0_113])).

cnf(refute_0_115, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 79/64 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 79/64 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_120, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8),
    inference(resolve, [], [refute_0_106, refute_0_120])).

cnf(refute_0_122, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_123, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_124, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_121, refute_0_124])).

cnf(refute_0_126, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6), inference(subst, [], [refute_0_48])).

cnf(refute_0_127, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6),
    inference(resolve, [], [refute_0_94, refute_0_127])).

cnf(refute_0_129, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_130, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_131, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_128, refute_0_131])).

cnf(refute_0_133, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_134, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_132, refute_0_133])).

cnf(refute_0_135, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_86, refute_0_134])).

cnf(refute_0_136, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_137, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_138, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_136, refute_0_137])).

cnf(refute_0_139, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_135, refute_0_138])).

cnf(refute_0_140, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2), inference(subst, [], [refute_0_48])).

cnf(refute_0_141, plain,
    (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_78, refute_0_141])).

cnf(refute_0_143, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600), introduced(tautology, [refl])).

cnf(refute_0_144, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 + 277555600),
    introduced(tautology, [equality])).

cnf(refute_0_145, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_142, refute_0_145])).

cnf(refute_0_147, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_148, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 + 277555600 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_149, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
     21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
     15328072984 * 79/64 ^ 2 + 277555600 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_147, refute_0_148])).

cnf(refute_0_150, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_146, refute_0_149])).

cnf(refute_0_151, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_152, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_153, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_150, refute_0_153])).

cnf(refute_0_155, plain,
    (4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_156, plain,
    (4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_70, refute_0_156])).

cnf(refute_0_158, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
      21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
      15328072984 * 79/64 ^ 2 + 277555600) /
     (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
      3120782021 * 79/64 ^ 2 + 69388900) ^ 2 |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_163, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_164, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_161, refute_0_164])).

cnf(refute_0_166, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 <
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 <
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain, (406627 * skoXC1 ^ 2 = 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_169, plain,
    (406627 * skoXC1 ^ 2 != 406627 * skoXC1 ^ 2 | skoXC1 ^ 2 != 79/64 ^ 2 |
     406627 * skoXC1 ^ 2 = 406627 * 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_170, plain,
    (skoXC1 ^ 2 != 79/64 ^ 2 | 406627 * skoXC1 ^ 2 = 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain, (406627 * skoXC1 ^ 2 = 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_170])).

cnf(refute_0_172, plain,
    (14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_173, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 79/64 ^ 2 |
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_174, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 79/64 ^ 2 |
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_171, refute_0_174])).

cnf(refute_0_176, plain, (14325 * skoXC1 ^ 4 = 14325 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_177, plain,
    (14325 * skoXC1 ^ 4 != 14325 * skoXC1 ^ 4 | skoXC1 ^ 4 != 79/64 ^ 4 |
     14325 * skoXC1 ^ 4 = 14325 * 79/64 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_178, plain,
    (skoXC1 ^ 4 != 79/64 ^ 4 | 14325 * skoXC1 ^ 4 = 14325 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_176, refute_0_177])).

cnf(refute_0_179, plain, (14325 * skoXC1 ^ 4 = 14325 * 79/64 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_178])).

cnf(refute_0_180, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_181, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 79/64 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_182, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 79/64 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_179, refute_0_182])).

cnf(refute_0_184, plain,
    (14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_185, plain,
    (14325 * 79/64 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_183, refute_0_184])).

cnf(refute_0_186, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_185])).

cnf(refute_0_187, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660),
    introduced(tautology, [refl])).

cnf(refute_0_188, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660),
    introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660),
    inference(resolve, [], [refute_0_186, refute_0_189])).

cnf(refute_0_191, plain,
    ((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_192, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660 |
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660 |
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    ((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_190, refute_0_193])).

cnf(refute_0_195, plain, (169932 * skoXC1 = 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_196, plain,
    (169932 * skoXC1 != 169932 * skoXC1 | skoXC1 != 79/64 |
     169932 * skoXC1 = 169932 * 79/64), introduced(tautology, [equality])).

cnf(refute_0_197, plain,
    (skoXC1 != 79/64 | 169932 * skoXC1 = 169932 * 79/64),
    inference(resolve, [], [refute_0_195, refute_0_196])).

cnf(refute_0_198, plain, (169932 * skoXC1 = 169932 * 79/64),
    inference(resolve, [], [refute_0_7, refute_0_197])).

cnf(refute_0_199, plain,
    (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_200, plain,
    (169932 * skoXC1 != 169932 * 79/64 |
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 |
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    introduced(tautology, [equality])).

cnf(refute_0_201, plain,
    (169932 * skoXC1 != 169932 * 79/64 |
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    inference(resolve, [], [refute_0_198, refute_0_201])).

cnf(refute_0_203, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_204, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 79/64 | skoXC1 ^ 5 = 79/64 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_205, plain, (skoXC1 != 79/64 | skoXC1 ^ 5 = 79/64 ^ 5),
    inference(resolve, [], [refute_0_203, refute_0_204])).

cnf(refute_0_206, plain, (skoXC1 ^ 5 = 79/64 ^ 5),
    inference(resolve, [], [refute_0_7, refute_0_205])).

cnf(refute_0_207, plain, (3890 * skoXC1 ^ 5 = 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_208, plain,
    (3890 * skoXC1 ^ 5 != 3890 * skoXC1 ^ 5 | skoXC1 ^ 5 != 79/64 ^ 5 |
     3890 * skoXC1 ^ 5 = 3890 * 79/64 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_209, plain,
    (skoXC1 ^ 5 != 79/64 ^ 5 | 3890 * skoXC1 ^ 5 = 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_207, refute_0_208])).

cnf(refute_0_210, plain, (3890 * skoXC1 ^ 5 = 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_206, refute_0_209])).

cnf(refute_0_211, plain,
    (105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_212, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 79/64 ^ 5 |
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_213, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 79/64 ^ 5 |
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_211, refute_0_212])).

cnf(refute_0_214, plain,
    (105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_210, refute_0_213])).

cnf(refute_0_215, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_216, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 79/64 | skoXC1 ^ 3 = 79/64 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_217, plain, (skoXC1 != 79/64 | skoXC1 ^ 3 = 79/64 ^ 3),
    inference(resolve, [], [refute_0_215, refute_0_216])).

cnf(refute_0_218, plain, (skoXC1 ^ 3 = 79/64 ^ 3),
    inference(resolve, [], [refute_0_7, refute_0_217])).

cnf(refute_0_219, plain, (105247 * skoXC1 ^ 3 = 105247 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_220, plain,
    (105247 * skoXC1 ^ 3 != 105247 * skoXC1 ^ 3 | skoXC1 ^ 3 != 79/64 ^ 3 |
     105247 * skoXC1 ^ 3 = 105247 * 79/64 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_221, plain,
    (skoXC1 ^ 3 != 79/64 ^ 3 | 105247 * skoXC1 ^ 3 = 105247 * 79/64 ^ 3),
    inference(resolve, [], [refute_0_219, refute_0_220])).

cnf(refute_0_222, plain, (105247 * skoXC1 ^ 3 = 105247 * 79/64 ^ 3),
    inference(resolve, [], [refute_0_218, refute_0_221])).

cnf(refute_0_223, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_224, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 79/64 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_225, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 79/64 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_222, refute_0_225])).

cnf(refute_0_227, plain,
    (105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_228, plain,
    (105247 * 79/64 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5),
    inference(resolve, [], [refute_0_214, refute_0_228])).

cnf(refute_0_230, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_232, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_229, refute_0_232])).

cnf(refute_0_234, plain,
    (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_235, plain,
    (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    inference(resolve, [], [refute_0_233, refute_0_234])).

cnf(refute_0_236, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64),
    inference(resolve, [], [refute_0_202, refute_0_235])).

cnf(refute_0_237, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_238, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_239, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_236, refute_0_239])).

cnf(refute_0_241, plain,
    ((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    inference(subst, [], [refute_0_48])).

cnf(refute_0_242, plain,
    ((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_240, refute_0_241])).

cnf(refute_0_243, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_194, refute_0_242])).

cnf(refute_0_244, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_245, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_246, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
     (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_243, refute_0_246])).

cnf(refute_0_248, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_249, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_250, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_247, refute_0_250])).

cnf(refute_0_252, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 =
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6),
    introduced(tautology, [refl])).

cnf(refute_0_253, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 !=
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    introduced(tautology, [equality])).

cnf(refute_0_254, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_252, refute_0_253])).

cnf(refute_0_255, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 =
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_251, refute_0_254])).

cnf(refute_0_256, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 <
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) |
     12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 !=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6 |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    introduced(tautology, [equality])).

cnf(refute_0_257, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 <
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_255, refute_0_256])).

cnf(refute_0_258, plain,
    (12 / pi *
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) - 6 <
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_167, refute_0_257])).

cnf(refute_0_259, plain,
    (4347/1000 *
     ln(4250 ^ 2 *
        (15132100 * 79/64 ^ 10 + 1024027285 * 79/64 ^ 8 +
         21404723599 * 79/64 ^ 6 + 129098541721 * 79/64 ^ 4 +
         15328072984 * 79/64 ^ 2 + 277555600) /
        (3783025 * 79/64 ^ 6 + 217230815 * 79/64 ^ 4 +
         3120782021 * 79/64 ^ 2 + 69388900) ^ 2) <=
     12 / pi *
     arctan((105247 * 79/64 ^ 3 + 3890 * 79/64 ^ 5 - 169932 * 79/64) /
            (14325 * 79/64 ^ 4 + 406627 * 79/64 ^ 2 - 16660)) - 6),
    inference(resolve, [], [refute_0_6, refute_0_258])).

cnf(refute_0_260, plain,
    (arctan(-118851281795/113846536469728) * -12 / pi <=
     -6 + ln(94033498571776000000/367111367871840337041) * -4347/1000),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (-6 + ln(94033498571776000000/367111367871840337041) * -4347/1000 <
     arctan(-118851281795/113846536469728) * -12 / pi |
     arctan(-118851281795/113846536469728) * -12 <=
     (-6 + ln(94033498571776000000/367111367871840337041) * -4347/1000) *
     pi | pi <= 0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_262, plain,
    (arctan(-118851281795/113846536469728) * -12 <=
     (-6 + ln(94033498571776000000/367111367871840337041) * -4347/1000) *
     pi | pi <= 0), inference(resolve, [], [refute_0_260, refute_0_261])).

cnf(refute_0_263, plain,
    (ln(94033498571776000000/367111367871840337041) * (pi * 4347/1000) <=
     pi * -6 + arctan(-118851281795/113846536469728) * 12 | pi <= 0),
    inference(arithmetic, [], [refute_0_262])).

cnf(refute_0_264, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_265, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_266, plain, (0 < pi),
    inference(decision, [], [refute_0_7, refute_0_264, refute_0_265])).

cnf(refute_0_267, plain,
    (ln(94033498571776000000/367111367871840337041) * (pi * 4347/1000) <=
     pi * -6 + arctan(-118851281795/113846536469728) * 12),
    inference(resolve, [], [refute_0_263, refute_0_266])).

cnf(refute_0_268, plain,
    (pi * 4347/1000 <= 0 |
     ln(94033498571776000000/367111367871840337041) <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000)), inference(split, [], [refute_0_267])).

cnf(refute_0_269, plain,
    (ln(94033498571776000000/367111367871840337041) <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000) | pi <= 0),
    inference(arithmetic, [], [refute_0_268])).

cnf(refute_0_270, plain,
    (ln(94033498571776000000/367111367871840337041) <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000)),
    inference(resolve, [], [refute_0_269, refute_0_266])).

cnf(refute_0_271, plain,
    (X_000139 < ln(X_000138) | ~ lgen(1, X_000139, ln(X_000138))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_272, plain,
    (~ lgen(1, X_000139,
            1/3 * (10 * X_000138 ^ 2 + 19 * X_000138 + 1) *
            (X_000138 - 1) /
            (X_000138 * (X_000138 ^ 2 + 6 * X_000138 + 3))) |
     X_000138 <= 0 | lgen(1, X_000139, ln(X_000138))),
    inference(subst, [], [ln_lower_bound_cf5])).

cnf(refute_0_273, plain,
    (X_000139 < ln(X_000138) |
     ~ lgen(1, X_000139,
            1/3 * (10 * X_000138 ^ 2 + 19 * X_000138 + 1) *
            (X_000138 - 1) /
            (X_000138 * (X_000138 ^ 2 + 6 * X_000138 + 3))) |
     X_000138 <= 0), inference(resolve, [], [refute_0_272, refute_0_271])).

cnf(refute_0_274, plain,
    (X_000139 < ln(X_000138) | X_000138 <= 0 |
     (-1/3 + X_000138 * (-6 + X_000138 * (3 + X_000138 * 10/3))) /
     (X_000138 * (3 + X_000138 * (6 + X_000138))) <= X_000139),
    inference(arithmetic, [], [refute_0_273])).

cnf(refute_0_275, plain,
    ((pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000) < ln(94033498571776000000/367111367871840337041) |
     94033498571776000000/367111367871840337041 <= 0 |
     (-1/3 +
      94033498571776000000/367111367871840337041 *
      (-6 +
       94033498571776000000/367111367871840337041 *
       (3 + 94033498571776000000/367111367871840337041 * 10/3))) /
     (94033498571776000000/367111367871840337041 *
      (3 +
       94033498571776000000/367111367871840337041 *
       (6 + 94033498571776000000/367111367871840337041))) <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000)), inference(subst, [], [refute_0_274])).

cnf(refute_0_276, plain,
    (94033498571776000000/367111367871840337041 <= 0 |
     (-1/3 +
      94033498571776000000/367111367871840337041 *
      (-6 +
       94033498571776000000/367111367871840337041 *
       (3 + 94033498571776000000/367111367871840337041 * 10/3))) /
     (94033498571776000000/367111367871840337041 *
      (3 +
       94033498571776000000/367111367871840337041 *
       (6 + 94033498571776000000/367111367871840337041))) <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000)),
    inference(resolve, [], [refute_0_270, refute_0_275])).

cnf(refute_0_277, plain,
    (-240059561092143322228510396732189705305738729768140603014559921/
     174981060163832453081314238988776240786081308632718255104000000 <=
     (pi * -6 + arctan(-118851281795/113846536469728) * 12) /
     (pi * 4347/1000)), inference(arithmetic, [], [refute_0_276])).

cnf(refute_0_278, plain,
    (-240059561092143322228510396732189705305738729768140603014559921/
     174981060163832453081314238988776240786081308632718255104000000 *
     (pi * 4347/1000) <=
     pi * -6 + arctan(-118851281795/113846536469728) * 12 |
     pi * 4347/1000 <= 0),
    inference(resolve, [], [refute_0_277, refute_0_5])).

cnf(refute_0_279, plain,
    (pi *
     705272101716410751172304370869866194715732610466925479967558157/
     233308080218443270775085651985034987714775078176957673472000000000 <=
     arctan(-118851281795/113846536469728) | pi <= 0),
    inference(arithmetic, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (pi *
     705272101716410751172304370869866194715732610466925479967558157/
     233308080218443270775085651985034987714775078176957673472000000000 <=
     arctan(-118851281795/113846536469728)),
    inference(resolve, [], [refute_0_279, refute_0_266])).

cnf(refute_0_281, plain,
    (0 <= -118851281795/113846536469728 |
     pi *
     705272101716410751172304370869866194715732610466925479967558157/
     233308080218443270775085651985034987714775078176957673472000000000 <=
     -118851281795/113846536469728 * 3 /
     (3 + -118851281795/113846536469728 * -118851281795/113846536469728)),
    inference(resolve, [], [refute_0_280, refute_0_4])).

cnf(refute_0_282, plain,
    (pi <=
     -9470539666088802469263306934976354163073498688289720893062522152731538634755932160000000000/
     27423176748008909205964089598792357290612920439466819077880775039555429126290719347475990389),
    inference(arithmetic, [], [refute_0_281])).

cnf(refute_0_283, plain,
    (-9470539666088802469263306934976354163073498688289720893062522152731538634755932160000000000/
     27423176748008909205964089598792357290612920439466819077880775039555429126290719347475990389 <
     pi),
    inference(decision, [], [refute_0_7, refute_0_264, refute_0_265])).

cnf(refute_0_284, plain, ($false),
    inference(resolve, [], [refute_0_282, refute_0_283])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-5-weak.tptp
