% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-5.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_1_5', conjecture,
    (! [X] :
       (X = 15839/128 =>
        10 / ln(10) *
        ln(33600000000000 ^ 2 /
           (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
            10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000))) - 18)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 15839/128 =>
        10 / ln(10) *
        ln(33600000000000 ^ 2 /
           (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
            10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000))) - 18))),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 15839/128 =>
          10 / ln(10) *
          ln(33600000000000 ^ 2 /
             (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
              10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
          -12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000))) - 18))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000))) - 18) <=
        10 / ln(10) *
        ln(33600000000000 ^ 2 /
           (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
            10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) &
        X = 15839/128)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) & skoXC1 = 15839/128),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 15839/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain, (skoXC1 = 15839/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_2, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_3, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 15839/128 |
     skoXC1 ^ 2 = 15839/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_4, plain, (skoXC1 != 15839/128 | skoXC1 ^ 2 = 15839/128 ^ 2),
    inference(resolve, [], [refute_0_2, refute_0_3])).

cnf(refute_0_5, plain, (skoXC1 ^ 2 = 15839/128 ^ 2),
    inference(resolve, [], [refute_0_1, refute_0_4])).

cnf(refute_0_6, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_7, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_8, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_6, refute_0_7])).

cnf(refute_0_9, plain, (29 * skoXC1 ^ 2 = 29 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_5, refute_0_8])).

cnf(refute_0_10, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_11, plain,
    (29 * skoXC1 ^ 2 != 29 * 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_12, plain,
    (29 * skoXC1 ^ 2 != 29 * 15839/128 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 15839/128 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_9, refute_0_12])).

cnf(refute_0_14, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_15, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 15839/128 ^ 2 - 90000000 |
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_16, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 15839/128 ^ 2 - 90000000 |
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_13, refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_19, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 15839/128 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_20, plain,
    (skoXC1 != 15839/128 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_1, refute_0_20])).

cnf(refute_0_22, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_23, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_24, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_26, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (15839/128 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_21, refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     15839/128 * (29 * 15839/128 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_17, refute_0_28])).

cnf(refute_0_30, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_31, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_32, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     15839/128 * (29 * 15839/128 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_29, refute_0_32])).

cnf(refute_0_34, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_35, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 15839/128 |
     skoXC1 ^ 4 = 15839/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_36, plain, (skoXC1 != 15839/128 | skoXC1 ^ 4 = 15839/128 ^ 4),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain, (skoXC1 ^ 4 = 15839/128 ^ 4),
    inference(resolve, [], [refute_0_1, refute_0_36])).

cnf(refute_0_38, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_39, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 15839/128 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_40, plain,
    (skoXC1 ^ 4 != 15839/128 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain, (3 * skoXC1 ^ 4 = 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_37, refute_0_40])).

cnf(refute_0_42, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_43, plain,
    (3 * skoXC1 ^ 4 != 3 * 15839/128 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_44, plain,
    (3 * skoXC1 ^ 4 != 3 * 15839/128 ^ 4 |
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_41, refute_0_44])).

cnf(refute_0_46, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_47, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_5, refute_0_48])).

cnf(refute_0_50, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_51, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_52, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 15839/128 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_49, refute_0_52])).

cnf(refute_0_54, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_55, plain,
    (-130660000 * 15839/128 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_45, refute_0_55])).

cnf(refute_0_57, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_58, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_56, refute_0_59])).

cnf(refute_0_61, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_62, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_63, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_60, refute_0_63])).

cnf(refute_0_65, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_66, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_67, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_64, refute_0_67])).

cnf(refute_0_69, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_70, plain,
    (1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_33, refute_0_70])).

cnf(refute_0_72, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_73, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 + 2000000000000) /
     (15839/128 * (29 * 15839/128 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_71, refute_0_74])).

cnf(refute_0_76, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_77, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_78, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_75, refute_0_78])).

cnf(refute_0_80, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_81, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_82, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18),
    inference(resolve, [], [refute_0_80, refute_0_81])).

cnf(refute_0_83, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 =
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18),
    inference(resolve, [], [refute_0_79, refute_0_82])).

cnf(refute_0_84, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18)),
    introduced(tautology, [refl])).

cnf(refute_0_85, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18)),
    introduced(tautology, [equality])).

cnf(refute_0_86, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18 !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
             2000000000000) /
            (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18)),
    inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18)),
    inference(resolve, [], [refute_0_83, refute_0_86])).

cnf(refute_0_88, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_89, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_91, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 | skoXC1 ^ 2 != 15839/128 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_92, plain,
    (skoXC1 ^ 2 != 15839/128 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_5, refute_0_92])).

cnf(refute_0_94, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_95, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_96, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_94, refute_0_95])).

cnf(refute_0_97, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_93, refute_0_96])).

cnf(refute_0_98, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_99, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 15839/128 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_100, plain,
    (skoXC1 ^ 4 != 15839/128 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_37, refute_0_100])).

cnf(refute_0_102, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_103, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 15839/128 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_104, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 15839/128 ^ 4 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_101, refute_0_104])).

cnf(refute_0_106, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_107, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 15839/128 |
     skoXC1 ^ 6 = 15839/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_108, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 6 = 15839/128 ^ 6),
    inference(resolve, [], [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain, (skoXC1 ^ 6 = 15839/128 ^ 6),
    inference(resolve, [], [refute_0_1, refute_0_108])).

cnf(refute_0_110, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_111, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 15839/128 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_112, plain,
    (skoXC1 ^ 6 != 15839/128 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_109, refute_0_112])).

cnf(refute_0_114, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_115, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 15839/128 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_116, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 15839/128 ^ 6 |
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_113, refute_0_116])).

cnf(refute_0_118, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_119, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 15839/128 |
     skoXC1 ^ 8 = 15839/128 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_120, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 8 = 15839/128 ^ 8),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain, (skoXC1 ^ 8 = 15839/128 ^ 8),
    inference(resolve, [], [refute_0_1, refute_0_120])).

cnf(refute_0_122, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_123, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 15839/128 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_124, plain,
    (skoXC1 ^ 8 != 15839/128 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_121, refute_0_124])).

cnf(refute_0_126, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_127, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 15839/128 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_128, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 15839/128 ^ 8 |
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_125, refute_0_128])).

cnf(refute_0_130, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_131, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 15839/128 |
     skoXC1 ^ 10 = 15839/128 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_132, plain,
    (skoXC1 != 15839/128 | skoXC1 ^ 10 = 15839/128 ^ 10),
    inference(resolve, [], [refute_0_130, refute_0_131])).

cnf(refute_0_133, plain, (skoXC1 ^ 10 = 15839/128 ^ 10),
    inference(resolve, [], [refute_0_1, refute_0_132])).

cnf(refute_0_134, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (skoXC1 ^ 10 != 15839/128 ^ 10 | 9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain, (9 * skoXC1 ^ 10 = 9 * 15839/128 ^ 10),
    inference(resolve, [], [refute_0_133, refute_0_136])).

cnf(refute_0_138, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (9 * skoXC1 ^ 10 != 9 * 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (9 * skoXC1 ^ 10 != 9 * 15839/128 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_143, plain,
    (9 * 15839/128 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8),
    inference(resolve, [], [refute_0_129, refute_0_143])).

cnf(refute_0_145, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_146, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_147, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_144, refute_0_147])).

cnf(refute_0_149, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_150, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6),
    inference(resolve, [], [refute_0_117, refute_0_150])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_151, refute_0_154])).

cnf(refute_0_156, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_157, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_155, refute_0_156])).

cnf(refute_0_158, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4),
    inference(resolve, [], [refute_0_105, refute_0_157])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_159, refute_0_160])).

cnf(refute_0_162, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_158, refute_0_161])).

cnf(refute_0_163, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_164, plain,
    (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2),
    inference(resolve, [], [refute_0_97, refute_0_164])).

cnf(refute_0_166, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_167, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
      10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
      4000000000000000000000000 * 15839/128 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
     10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
     4000000000000000000000000 * 15839/128 ^ 2 |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
      10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
      4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_166, refute_0_167])).

cnf(refute_0_169, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
      10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
      4000000000000000000000000 * 15839/128 ^ 2)),
    inference(resolve, [], [refute_0_165, refute_0_168])).

cnf(refute_0_170, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_171, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
      10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
      4000000000000000000000000 * 15839/128 ^ 2) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_172, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
      10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
      4000000000000000000000000 * 15839/128 ^ 2) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_169, refute_0_172])).

cnf(refute_0_174, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_175, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2)) |
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_176, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2)) |
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_174, refute_0_175])).

cnf(refute_0_177, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_173, refute_0_176])).

cnf(refute_0_178, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) |
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2)) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_179, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) - 18) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_89, refute_0_179])).

cnf(refute_0_181, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 15839/128 ^ 2 + 3 * 15839/128 ^ 4 +
              2000000000000) /
             (15839/128 * (29 * 15839/128 ^ 2 - 90000000))) - 18) <=
     10 / ln(10) *
     ln(33600000000000 ^ 2 /
        (9 * 15839/128 ^ 10 + 309000000 * 15839/128 ^ 8 +
         10300000000000000 * 15839/128 ^ 6 + 10 ^ 22 * 15839/128 ^ 4 +
         4000000000000000000000000 * 15839/128 ^ 2))),
    inference(resolve, [], [refute_0_0, refute_0_180])).

cnf(refute_0_182, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * -10 /
     ln(10) <=
     (-216 + pi * -12 +
      arctan(-5533264095502723/3391227738231320302080) * -12) / pi),
    inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain,
    ((-216 + pi * -12 +
      arctan(-5533264095502723/3391227738231320302080) * -12) / pi <
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * -10 /
     ln(10) |
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * -10 /
     ln(10) * pi <=
     -216 + pi * -12 +
     arctan(-5533264095502723/3391227738231320302080) * -12 | pi <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_184, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * -10 /
     ln(10) * pi <=
     -216 + pi * -12 +
     arctan(-5533264095502723/3391227738231320302080) * -12 | pi <= 0),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * -10) /
     ln(10) <=
     -216 + pi * -12 +
     arctan(-5533264095502723/3391227738231320302080) * -12 | pi <= 0),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_187, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_188, plain, (0 < pi),
    inference(decision, [], [refute_0_1, refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * -10) /
     ln(10) <=
     -216 + pi * -12 +
     arctan(-5533264095502723/3391227738231320302080) * -12),
    inference(resolve, [], [refute_0_185, refute_0_188])).

cnf(refute_0_190, plain,
    (-216 + pi * -12 +
     arctan(-5533264095502723/3391227738231320302080) * -12 <
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * -10) /
     ln(10) |
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <=
     (-216 + pi * -12 +
      arctan(-5533264095502723/3391227738231320302080) * -12) * ln(10) |
     ln(10) <= 0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_191, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <=
     (-216 + pi * -12 +
      arctan(-5533264095502723/3391227738231320302080) * -12) * ln(10) |
     ln(10) <= 0), inference(resolve, [], [refute_0_189, refute_0_190])).

cnf(refute_0_192, plain,
    (ln(10) *
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10)),
    inference(arithmetic, [], [refute_0_191])).

cnf(refute_0_193, plain,
    (216 + pi * 12 +
     arctan(-5533264095502723/3391227738231320302080) * 12 < 0 |
     0 <
     216 + pi * 12 +
     arctan(-5533264095502723/3391227738231320302080) * 12 |
     0 <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10)),
    inference(split, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (arctan(-5533264095502723/3391227738231320302080) < -18 + pi * -1 |
     -18 + pi * -1 < arctan(-5533264095502723/3391227738231320302080) |
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <= 0), inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (~ lgen(0, X_000034, arctan(X_000033)) | X_000034 <= arctan(X_000033)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_196, plain,
    (~ lgen(0, X_000034, X_000033) | 0 <= X_000033 |
     lgen(0, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_197, plain,
    (~ lgen(0, X_000034, X_000033) | X_000034 <= arctan(X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_196, refute_0_195])).

cnf(refute_0_198, plain,
    (X_000033 < X_000034 | X_000034 <= arctan(X_000033) | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (-5533264095502723/3391227738231320302080 < -18 + pi * -1 |
     0 <= -5533264095502723/3391227738231320302080 |
     -18 + pi * -1 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(subst, [], [refute_0_198])).

cnf(refute_0_200, plain,
    (pi < -61042093754899669934717/3391227738231320302080 |
     -18 + pi * -1 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(arithmetic, [], [refute_0_199])).

cnf(refute_0_201, plain,
    (-61042093754899669934717/3391227738231320302080 <= pi),
    inference(decision, [], [refute_0_187, refute_0_186, refute_0_1])).

cnf(refute_0_202, plain,
    (-18 + pi * -1 <= arctan(-5533264095502723/3391227738231320302080)),
    inference(resolve, [], [refute_0_201, refute_0_200])).

cnf(refute_0_203, plain,
    (-18 + pi * -1 < arctan(-5533264095502723/3391227738231320302080) |
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <= 0),
    inference(resolve, [], [refute_0_202, refute_0_194])).

cnf(refute_0_204, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_205, plain,
    (~ lgen(1, X_000034, X_000033) | 0 <= X_000033 |
     lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_206, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_205, refute_0_204])).

cnf(refute_0_207, plain,
    (X_000034 < arctan(X_000033) | X_000033 <= X_000034 | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_206])).

cnf(refute_0_208, plain,
    (-18 +
     pi *
     -7139785999374718428962902603599363025325100197503108/
     3708020447920192134680327315561207724688032794857383 <
     arctan(-5533264095502723/3391227738231320302080) |
     -5533264095502723/3391227738231320302080 <=
     -18 +
     pi *
     -7139785999374718428962902603599363025325100197503108/
     3708020447920192134680327315561207724688032794857383 |
     0 <= -5533264095502723/3391227738231320302080),
    inference(subst, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (-137270622058181051771303011521526212025482696105829/
     137334090663710819802975085761526212025482696105829 <
     (-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) | 0 <= pi * -10 |
     -137270622058181051771303011521526212025482696105829/
     137334090663710819802975085761526212025482696105829 * (pi * -10) <=
     -972/5 + pi * -54/5 +
     arctan(-5533264095502723/3391227738231320302080) * -54/5),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_210, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_211, plain,
    (~ lgen(1, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_212, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, X_000047 - 1, X_000048) |
     X_000047 <= 0), inference(resolve, [], [refute_0_211, refute_0_210])).

cnf(refute_0_213, plain,
    (ln(X_000047) < X_000048 | X_000047 <= 0 | X_000048 <= -1 + X_000047),
    inference(arithmetic, [], [refute_0_212])).

cnf(refute_0_214, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) <
     (-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) |
     63468605529768031672074240000000000000000000000/
     137334090663710819802975085761526212025482696105829 <= 0 |
     (-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     -1 +
     63468605529768031672074240000000000000000000000/
     137334090663710819802975085761526212025482696105829),
    inference(subst, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) < 9/10 |
     9/10 *
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) |
     216 + pi * 12 +
     arctan(-5533264095502723/3391227738231320302080) * 12 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_216, plain,
    (216 + pi * 12 +
     arctan(-5533264095502723/3391227738231320302080) * 12 <= 0 |
     ln(10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12)),
    inference(split, [], [refute_0_192])).

cnf(refute_0_217, plain,
    (arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1 |
     ln(10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12)),
    inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (X_000069 < ln(X_000068) | ~ lgen(1, X_000069, ln(X_000068))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_219, plain,
    (~ lgen(1, X_000069, (X_000068 - 1) / X_000068) | X_000068 <= 0 |
     lgen(1, X_000069, ln(X_000068))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_220, plain,
    (X_000069 < ln(X_000068) |
     ~ lgen(1, X_000069, (X_000068 - 1) / X_000068) | X_000068 <= 0),
    inference(resolve, [], [refute_0_219, refute_0_218])).

cnf(refute_0_221, plain,
    (X_000069 < ln(X_000068) | X_000068 <= 0 |
     (-1 + X_000068) / X_000068 <= X_000069),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) < ln(10) |
     10 <= 0 |
     (-1 + 10) / 10 <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12)),
    inference(subst, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (10 <= 0 |
     (-1 + 10) / 10 <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_217, refute_0_222])).

cnf(refute_0_224, plain,
    (9/10 <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) /
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (9/10 *
     (216 + pi * 12 +
      arctan(-5533264095502723/3391227738231320302080) * 12) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) * (pi * 10) |
     216 + pi * 12 +
     arctan(-5533264095502723/3391227738231320302080) * 12 <= 0 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_224, refute_0_215])).

cnf(refute_0_226, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <=
     -972/5 + pi * -54/5 +
     arctan(-5533264095502723/3391227738231320302080) * -54/5 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(arithmetic, [], [refute_0_225])).

cnf(refute_0_227, plain,
    (0 <= pi * -10 |
     (-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(split, [], [refute_0_226])).

cnf(refute_0_228, plain,
    ((-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1 |
     pi <= 0), inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    ((-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_228, refute_0_188])).

cnf(refute_0_230, plain,
    (63468605529768031672074240000000000000000000000/
     137334090663710819802975085761526212025482696105829 <= 0 |
     (-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     -1 +
     63468605529768031672074240000000000000000000000/
     137334090663710819802975085761526212025482696105829 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_229, refute_0_214])).

cnf(refute_0_231, plain,
    ((-972/5 + pi * -54/5 +
      arctan(-5533264095502723/3391227738231320302080) * -54/5) /
     (pi * -10) <=
     -137270622058181051771303011521526212025482696105829/
     137334090663710819802975085761526212025482696105829 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(arithmetic, [], [refute_0_230])).

cnf(refute_0_232, plain,
    (0 <= pi * -10 |
     -137270622058181051771303011521526212025482696105829/
     137334090663710819802975085761526212025482696105829 * (pi * -10) <=
     -972/5 + pi * -54/5 +
     arctan(-5533264095502723/3391227738231320302080) * -54/5 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_231, refute_0_209])).

cnf(refute_0_233, plain,
    (arctan(-5533264095502723/3391227738231320302080) <=
     -18 +
     pi *
     -7139785999374718428962902603599363025325100197503108/
     3708020447920192134680327315561207724688032794857383 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1 |
     pi <= 0), inference(arithmetic, [], [refute_0_232])).

cnf(refute_0_234, plain,
    (arctan(-5533264095502723/3391227738231320302080) <=
     -18 +
     pi *
     -7139785999374718428962902603599363025325100197503108/
     3708020447920192134680327315561207724688032794857383 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_233, refute_0_188])).

cnf(refute_0_235, plain,
    (-5533264095502723/3391227738231320302080 <=
     -18 +
     pi *
     -7139785999374718428962902603599363025325100197503108/
     3708020447920192134680327315561207724688032794857383 |
     0 <= -5533264095502723/3391227738231320302080 |
     arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_234, refute_0_208])).

cnf(refute_0_236, plain,
    (arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1 |
     pi <=
     -4763460063283234150075964290088354315371994483603737203923359614810183/
     509557428417517377915287001292615699265363962533737463617282038097920),
    inference(arithmetic, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (-4763460063283234150075964290088354315371994483603737203923359614810183/
     509557428417517377915287001292615699265363962533737463617282038097920 <
     pi),
    inference(decision, [], [refute_0_1, refute_0_186, refute_0_187])).

cnf(refute_0_238, plain,
    (arctan(-5533264095502723/3391227738231320302080) <= -18 + pi * -1),
    inference(resolve, [], [refute_0_236, refute_0_237])).

cnf(refute_0_239, plain,
    (ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829) *
     (pi * -10) <= 0),
    inference(resolve, [], [refute_0_238, refute_0_203])).

cnf(refute_0_240, plain,
    (0 <= pi * -10 |
     0 / (pi * -10) <=
     ln(63468605529768031672074240000000000000000000000/
        137334090663710819802975085761526212025482696105829)),
    inference(split, [], [refute_0_239])).

cnf(refute_0_241, plain, (pi <= 0),
    inference(arithmetic, [], [refute_0_240])).

cnf(refute_0_242, plain, ($false),
    inference(resolve, [], [refute_0_241, refute_0_188])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-5.tptp
