% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_14, axiom,
    (X <= 1 | ~ lgen(R, Y, 157079/100000 - 1 / X) |
     lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_34, axiom,
    (X <= 1 |
     ~ lgen(R, Y,
            157079/100000 -
            1/15 * (64 + 735 * X ^ 2 + 945 * X ^ 4) /
            (X * (15 + 70 * X ^ 2 + 63 * X ^ 4))) |
     lgen(R, Y, arctan(X)))).

fof('Nichols_Plot_Disk_Drive_Meets_1_2', conjecture,
    (! [X] :
       ((354991/512 <= X & X <= 8808/5) =>
        157/200 <=
        arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
               (X * (29 * X ^ 2 - 90000000)))))).

fof(subgoal_0, plain,
    (! [X] :
       ((354991/512 <= X & X <= 8808/5) =>
        157/200 <=
        arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
               (X * (29 * X ^ 2 - 90000000))))),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_2'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((354991/512 <= X & X <= 8808/5) =>
          157/200 <=
          arctan(1/1140 *
                 (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                 (X * (29 * X ^ 2 - 90000000))))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
               (X * (29 * X ^ 2 - 90000000))) < 157/200 & X <= 8808/5 &
        354991/512 <= X)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) < 157/200 &
     354991/512 <= skoXC1 & skoXC1 <= 8808/5),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) < 157/200),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 8808/5),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (354991/512 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 <
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) |
     0 <= skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) |
     1 * (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <=
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (78579/100000 *
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) |
     -1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 +
     skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000))))))))))))))))))))))))) <= 0 |
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) <=
     78579/100000), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_2, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) < 157/200),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (arctan((100000000000/57 +
             skoXC1 *
             (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
            (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))) < 157/200),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (~ lgen(0, X_000104, arctan(X_000103)) | X_000104 <= arctan(X_000103)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_5, plain,
    (~ lgen(0, X_000104,
            157079/100000 -
            1/15 * (64 + 735 * X_000103 ^ 2 + 945 * X_000103 ^ 4) /
            (X_000103 * (15 + 70 * X_000103 ^ 2 + 63 * X_000103 ^ 4))) |
     X_000103 <= 1 | lgen(0, X_000104, arctan(X_000103))),
    inference(subst, [], [atan_lower_bound_case_34])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000104,
            157079/100000 -
            1/15 * (64 + 735 * X_000103 ^ 2 + 945 * X_000103 ^ 4) /
            (X_000103 * (15 + 70 * X_000103 ^ 2 + 63 * X_000103 ^ 4))) |
     X_000103 <= 1 | X_000104 <= arctan(X_000103)),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (157079/100000 +
     (64/15 + X_000103 * (X_000103 * (49 + X_000103 * (X_000103 * 63)))) /
     (X_000103 *
      (15 + X_000103 * (X_000103 * (70 + X_000103 * (X_000103 * 63))))) *
     -1 < X_000104 | X_000103 <= 1 | X_000104 <= arctan(X_000103)),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (157079/100000 +
     (64/15 +
      (100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      ((100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
       (49 +
        (100000000000/57 +
         skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
        (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
        ((100000000000/57 +
          skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
         (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 63)))) /
     ((100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      (15 +
       (100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
       ((100000000000/57 +
         skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
        (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
        (70 +
         (100000000000/57 +
          skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
         (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
         ((100000000000/57 +
           skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
          (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 63))))) * -1 <
     157/200 |
     157/200 <=
     arctan((100000000000/57 +
             skoXC1 *
             (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
            (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (157079/100000 +
     (64/15 +
      (100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      ((100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
       (49 +
        (100000000000/57 +
         skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
        (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
        ((100000000000/57 +
          skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
         (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 63)))) /
     ((100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      (15 +
       (100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
       ((100000000000/57 +
         skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
        (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
        (70 +
         (100000000000/57 +
          skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
         (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
         ((100000000000/57 +
           skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
          (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 63))))) * -1 <
     157/200 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_8, refute_0_3])).

cnf(refute_0_10, plain,
    (78579/100000 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1 |
     skoXC1 * (skoXC1 * 29) = 90000000 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (65610000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-84564000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (40872600000000000000 +
             skoXC1 *
             (skoXC1 *
              (-8780040000000 + skoXC1 * (skoXC1 * 707281))))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (8100000000000000 +
       skoXC1 * (skoXC1 * (-5220000000 + skoXC1 * (skoXC1 * 841))))) = 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (157/200 * (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380))) |
     0 <= skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) |
     157/200 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_12, plain,
    (~ lgen(0, X_000034, arctan(X_000033)) | X_000034 <= arctan(X_000033)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (~ lgen(0, X_000034, X_000033) | 0 <= X_000033 |
     lgen(0, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_14, plain,
    (~ lgen(0, X_000034, X_000033) | X_000034 <= arctan(X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (X_000033 < X_000034 | X_000034 <= arctan(X_000033) | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 157/200 |
     0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) |
     157/200 <=
     arctan((100000000000/57 +
             skoXC1 *
             (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
            (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))))),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 157/200 |
     0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))),
    inference(resolve, [], [refute_0_16, refute_0_3])).

cnf(refute_0_18, plain,
    (157/200 * (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380))) |
     0 <= skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) |
     0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))),
    inference(resolve, [], [refute_0_11, refute_0_17])).

cnf(refute_0_19, plain,
    (-100000000000/57 <
     skoXC1 *
     (70650000 +
      skoXC1 * (-6533000/57 + skoXC1 * (-4553/200 + skoXC1 * 1/380))) |
     0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) |
     skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 <= 8808/5),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_21, plain, (354991/512 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (70650000 +
      skoXC1 * (-6533000/57 + skoXC1 * (-4553/200 + skoXC1 * 1/380))) <=
     -100000000000/57 |
     skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) |
     skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_19])).

cnf(refute_0_24, plain,
    (0 < skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29))),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_25, plain,
    (0 <=
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) <
     157/200 *
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) |
     157/200 <=
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) /
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) |
     10000000000000000000000/3249 +
     skoXC1 *
     (skoXC1 *
      (77644100000000000000/3249 +
       skoXC1 *
       (skoXC1 *
        (-8169251000000/3249 +
         skoXC1 *
         (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400))))))) <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000062, arctan(X_000061)) | X_000062 <= arctan(X_000061)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3)) |
     lgen(0, X_000062, arctan(X_000061))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_29, plain,
    (X_000061 < 0 |
     ~ lgen(0, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3)) |
     X_000062 <= arctan(X_000061)),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (X_000061 * 3 / (3 + X_000061 * X_000061) < X_000062 | X_000061 < 0 |
     X_000062 <= arctan(X_000061)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 3 /
     (3 +
      (100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      ((100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))))) < 157/200 |
     157/200 <=
     arctan((100000000000/57 +
             skoXC1 *
             (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
            (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))))),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) * 3 /
     (3 +
      (100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) *
      ((100000000000/57 +
        skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
       (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))))) < 157/200),
    inference(resolve, [], [refute_0_31, refute_0_3])).

cnf(refute_0_33, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) /
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) <
     157/200 | skoXC1 * (skoXC1 * 29) = 90000000 |
     skoXC1 *
     (skoXC1 *
      (8100000000000000 +
       skoXC1 * (skoXC1 * (-5220000000 + skoXC1 * (skoXC1 * 841))))) = 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) <
     78579/100000 *
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     0 <=
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380))) |
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) <=
     78579/100000), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_35, plain,
    (~ lgen(0, X_000048, arctan(X_000047)) | X_000048 <= arctan(X_000047)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000048, 157079/100000 - 1 / X_000047) | X_000047 <= 1 |
     lgen(0, X_000048, arctan(X_000047))),
    inference(subst, [], [atan_lower_bound_case_14])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000048, 157079/100000 - 1 / X_000047) | X_000047 <= 1 |
     X_000048 <= arctan(X_000047)),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (157079/100000 + 1 / X_000047 * -1 < X_000048 | X_000047 <= 1 |
     X_000048 <= arctan(X_000047)),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (157079/100000 +
     1 /
     ((100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))) * -1 < 157/200 |
     157/200 <=
     arctan((100000000000/57 +
             skoXC1 *
             (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
            (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (157079/100000 +
     1 /
     ((100000000000/57 +
       skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
      (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)))) * -1 < 157/200 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_39, refute_0_3])).

cnf(refute_0_41, plain,
    (78579/100000 <
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (78579/100000 <
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     0 <= skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) |
     1 * (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <=
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))),
    inference(resolve, [], [refute_0_41, refute_0_0])).

cnf(refute_0_43, plain,
    (78579/100000 <
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     skoXC1 *
     (-90000000 +
      skoXC1 * (6533000/57 + skoXC1 * (29 + skoXC1 * -1/380))) <=
     100000000000/57 | skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (100000000000/57 <
     skoXC1 *
     (-90000000 +
      skoXC1 * (6533000/57 + skoXC1 * (29 + skoXC1 * -1/380))) |
     skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_45, plain,
    (78579/100000 <
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (78579/100000 <
     skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) /
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380))))),
    inference(resolve, [], [refute_0_45, refute_0_24])).

cnf(refute_0_47, plain,
    (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) <
     78579/100000 *
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) |
     0 <=
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))),
    inference(resolve, [], [refute_0_34, refute_0_46])).

cnf(refute_0_48, plain,
    (-26193000000/19 <
     skoXC1 *
     (90000000 +
      skoXC1 *
      (-171118869/1900 + skoXC1 * (-29 + skoXC1 * 78579/38000000))) |
     skoXC1 * (skoXC1 * (6533000/57 + skoXC1 * (skoXC1 * -1/380))) <=
     100000000000/57), inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (100000000000/57 <
     skoXC1 * (skoXC1 * (6533000/57 + skoXC1 * (skoXC1 * -1/380))) |
     -26193000000/19 <
     skoXC1 *
     (90000000 +
      skoXC1 *
      (-171118869/1900 + skoXC1 * (-29 + skoXC1 * 78579/38000000)))),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_50, plain,
    (-26193000000/19 <
     skoXC1 *
     (90000000 +
      skoXC1 *
      (-171118869/1900 + skoXC1 * (-29 + skoXC1 * 78579/38000000)))),
    inference(resolve, [], [refute_0_48, refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 *
     (skoXC1 *
      (8100000000000000 +
       skoXC1 * (skoXC1 * (-5220000000 + skoXC1 * (skoXC1 * 841))))) != 0 |
     skoXC1 * (skoXC1 * 29) = 90000000 | skoXC1 = 0),
    inference(decision, [], [refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_52, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) /
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) <
     157/200 | skoXC1 * (skoXC1 * 29) = 90000000 | skoXC1 = 0),
    inference(resolve, [], [refute_0_33, refute_0_51])).

cnf(refute_0_53, plain, (skoXC1 * (skoXC1 * 29) != 90000000 | skoXC1 = 0),
    inference(decision, [], [refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_54, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) /
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) <
     157/200 | skoXC1 = 0),
    inference(resolve, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain, (skoXC1 != 0),
    inference(decision, [], [refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_56, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) /
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) <
     157/200), inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     skoXC1 *
     (-9000000000000000000/19 +
      skoXC1 *
      (skoXC1 *
       (590870000000000/19 +
        skoXC1 *
        (skoXC1 * (-202957000/19 + skoXC1 * (skoXC1 * 87/380)))))) <
     157/200 *
     (10000000000000000000000/3249 +
      skoXC1 *
      (skoXC1 *
       (77644100000000000000/3249 +
        skoXC1 *
        (skoXC1 *
         (-8169251000000/3249 +
          skoXC1 *
          (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400)))))))) |
     10000000000000000000000/3249 +
     skoXC1 *
     (skoXC1 *
      (77644100000000000000/3249 +
       skoXC1 *
       (skoXC1 *
        (-8169251000000/3249 +
         skoXC1 *
         (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400))))))) <=
     0), inference(resolve, [], [refute_0_26, refute_0_56])).

cnf(refute_0_58, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     -7850000000000000000000/3249 <
     skoXC1 *
     (9000000000000000000/19 +
      skoXC1 *
      (60950618500000000000/3249 +
       skoXC1 *
       (-590870000000000/19 +
        skoXC1 *
        (-6412862035000/3249 +
         skoXC1 *
         (202957000/19 +
          skoXC1 *
          (326420113/216600 +
           skoXC1 * (-87/380 + skoXC1 * 157/28880000))))))) |
     skoXC1 *
     (skoXC1 *
      (77644100000000000000/3249 +
       skoXC1 *
       (skoXC1 *
        (-8169251000000/3249 +
         skoXC1 *
         (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400))))))) <=
     -10000000000000000000000/3249),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (-10000000000000000000000/3249 <
     skoXC1 *
     (skoXC1 *
      (77644100000000000000/3249 +
       skoXC1 *
       (skoXC1 *
        (-8169251000000/3249 +
         skoXC1 *
         (skoXC1 * (2079109/1083 + skoXC1 * (skoXC1 * 1/144400))))))) |
     -7850000000000000000000/3249 <
     skoXC1 *
     (9000000000000000000/19 +
      skoXC1 *
      (60950618500000000000/3249 +
       skoXC1 *
       (-590870000000000/19 +
        skoXC1 *
        (-6412862035000/3249 +
         skoXC1 *
         (202957000/19 +
          skoXC1 *
          (326420113/216600 +
           skoXC1 * (-87/380 + skoXC1 * 157/28880000)))))))),
    inference(decision, [], [refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_60, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) < 0 |
     -7850000000000000000000/3249 <
     skoXC1 *
     (9000000000000000000/19 +
      skoXC1 *
      (60950618500000000000/3249 +
       skoXC1 *
       (-590870000000000/19 +
        skoXC1 *
        (-6412862035000/3249 +
         skoXC1 *
         (202957000/19 +
          skoXC1 *
          (326420113/216600 +
           skoXC1 * (-87/380 + skoXC1 * 157/28880000)))))))),
    inference(resolve, [], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (-7850000000000000000000/3249 <
     skoXC1 *
     (9000000000000000000/19 +
      skoXC1 *
      (60950618500000000000/3249 +
       skoXC1 *
       (-590870000000000/19 +
        skoXC1 *
        (-6412862035000/3249 +
         skoXC1 *
         (202957000/19 +
          skoXC1 *
          (326420113/216600 +
           skoXC1 * (-87/380 + skoXC1 * 157/28880000)))))))),
    inference(resolve, [], [refute_0_25, refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (65610000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-84564000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (40872600000000000000 +
             skoXC1 *
             (skoXC1 *
              (-8780040000000 + skoXC1 * (skoXC1 * 707281))))))))))) != 0 |
     skoXC1 * (skoXC1 * 29) = 90000000 |
     skoXC1 *
     (skoXC1 *
      (8100000000000000 +
       skoXC1 * (skoXC1 * (-5220000000 + skoXC1 * (skoXC1 * 841))))) = 0 |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_61, refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_63, plain,
    (78579/100000 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1 |
     skoXC1 * (skoXC1 * 29) = 90000000 |
     skoXC1 *
     (skoXC1 *
      (8100000000000000 +
       skoXC1 * (skoXC1 * (-5220000000 + skoXC1 * (skoXC1 * 841))))) = 0 |
     skoXC1 = 0), inference(resolve, [], [refute_0_10, refute_0_62])).

cnf(refute_0_64, plain,
    (78579/100000 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1 |
     skoXC1 * (skoXC1 * 29) = 90000000 | skoXC1 = 0),
    inference(resolve, [], [refute_0_63, refute_0_51])).

cnf(refute_0_65, plain,
    (78579/100000 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1 | skoXC1 = 0),
    inference(resolve, [], [refute_0_64, refute_0_53])).

cnf(refute_0_66, plain,
    (78579/100000 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) /
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_65, refute_0_55])).

cnf(refute_0_67, plain,
    (78579/100000 *
     (-1890000000000000000000000000000000000000000000000000000000000000000000000000000/
      2476099 +
      skoXC1 *
      (skoXC1 *
       (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
        2476099 +
        skoXC1 *
        (skoXC1 *
         (-2105831213695000000000000000000000000000000000000000000000000000000000000/
          2476099 +
          skoXC1 *
          (skoXC1 *
           (3873192586474466450000000000000000000000000000000000000000000000000000/
            66854673 +
            skoXC1 *
            (skoXC1 *
             (27298035567934532024500000000000000000000000000000000000000000000/
              66854673 +
              skoXC1 *
              (skoXC1 *
               (6037192745891231780429000000000000000000000000000000000000/
                66854673 +
                skoXC1 *
                (skoXC1 *
                 (-25687521996357455345524700000000000000000000000000000/
                  66854673 +
                  skoXC1 *
                  (skoXC1 *
                   (11371458976596733414176070000000000000000000000/
                    66854673 +
                    skoXC1 *
                    (skoXC1 *
                     (-1706231753425725563862139000000000000000/66854673 +
                      skoXC1 *
                      (skoXC1 *
                       (71132938024799332185890000000000/66854673 +
                        skoXC1 *
                        (skoXC1 *
                         (-88449140115525021685000/2476099 +
                          skoXC1 *
                          (skoXC1 *
                           (9249321356054667/9904396 +
                            skoXC1 *
                            (skoXC1 *
                             (-17697340451/990439600 +
                              skoXC1 *
                              (skoXC1 *
                               1536507/
                               7923516800000)))))))))))))))))))))))))) <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (skoXC1 *
                             44558703/20851360000)))))))))))))))))))))))) |
     -1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 +
     skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000))))))))))))))))))))))))) <= 0 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_1, refute_0_66])).

cnf(refute_0_68, plain,
    (-1485143100000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 <
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (3856149306765000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (1654741109409394050000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (-338168444725085665749500000000000000000000000000000000000000000000000/
           7428297 +
           skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (-2383391485436363991059095000000000000000000000000000000000000000/
             7428297 +
             skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (-527107298643763446749255990000000000000000000000000000000/
               7428297 +
               skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (2242777545501969426217761557000000000000000000000000/
                 7428297 +
                 skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (-992842083246660794391712671700000000000000000/7428297 +
                   skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (148971094391600098980803356090000000000/7428297 +
                     skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (-6210616818945229693150055900000/7428297 +
                       skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (1390048996227568135797123/49521980 +
                         skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (-726802422837419678193/990439600000 +
                           skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (1390639315299129/99043960000000 +
                             skoXC1 *
                             (44558703/20851360000 +
                              skoXC1 *
                              -120737183553/
                              792351680000000000))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000))))))))))))))))))))))))) <=
     1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000))))))))))))))))))))))))) <=
     1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 |
     skoXC1 *
     (5103000000000000000000000000000000000000000000000000000000000000000000000000/
      130321 +
      skoXC1 *
      (3856149306765000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (9105084450000000000000000000000000000000000000000000000000000000000000000/
        130321 +
        skoXC1 *
        (1654741109409394050000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (1141023334523760000000000000000000000000000000000000000000000000000000/
          130321 +
          skoXC1 *
          (-338168444725085665749500000000000000000000000000000000000000000000000/
           7428297 +
           skoXC1 *
           (35229454505508097000000000000000000000000000000000000000000000000/
            130321 +
            skoXC1 *
            (-2383391485436363991059095000000000000000000000000000000000000000/
             7428297 +
             skoXC1 *
             (155725628153753376670000000000000000000000000000000000000000/
              1172889 +
              skoXC1 *
              (-527107298643763446749255990000000000000000000000000000000/
               7428297 +
               skoXC1 *
               (-562841426895750963758000000000000000000000000000000000/
                1172889 +
                skoXC1 *
                (2242777545501969426217761557000000000000000000000000/
                 7428297 +
                 skoXC1 *
                 (116180975427488395845400000000000000000000000000/390963 +
                  skoXC1 *
                  (-992842083246660794391712671700000000000000000/7428297 +
                   skoXC1 *
                   (-90701515638045346820380000000000000000000/1172889 +
                    skoXC1 *
                    (148971094391600098980803356090000000000/7428297 +
                     skoXC1 *
                     (9440136668353309849000000000000/1083 +
                      skoXC1 *
                      (-6210616818945229693150055900000/7428297 +
                       skoXC1 *
                       (-403467848473662506831000000/1172889 +
                        skoXC1 *
                        (1390048996227568135797123/49521980 +
                         skoXC1 *
                         (6428869873502592428/651605 +
                          skoXC1 *
                          (-726802422837419678193/990439600000 +
                           skoXC1 *
                           (-10267827021431/52128400 +
                            skoXC1 *
                            (1390639315299129/99043960000000 +
                             skoXC1 *
                             (44558703/20851360000 +
                              skoXC1 *
                              -120737183553/
                              792351680000000000))))))))))))))))))))))))) <=
     -1485143100000000000000000000000000000000000000000000000000000000000000000000000/
     2476099),
    inference(decision, [],
              [refute_0_61, refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_70, plain,
    (skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000))))))))))))))))))))))))) <=
     1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 |
     (100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_69, refute_0_68])).

cnf(refute_0_71, plain,
    (1890000000000000000000000000000000000000000000000000000000000000000000000000000/
     2476099 <
     skoXC1 *
     (skoXC1 *
      (-4907353500000000000000000000000000000000000000000000000000000000000000000000/
       2476099 +
       skoXC1 *
       (skoXC1 *
        (-2105831213695000000000000000000000000000000000000000000000000000000000000/
         2476099 +
         skoXC1 *
         (skoXC1 *
          (3873192586474466450000000000000000000000000000000000000000000000000000/
           66854673 +
           skoXC1 *
           (skoXC1 *
            (27298035567934532024500000000000000000000000000000000000000000000/
             66854673 +
             skoXC1 *
             (skoXC1 *
              (6037192745891231780429000000000000000000000000000000000000/
               66854673 +
               skoXC1 *
               (skoXC1 *
                (-25687521996357455345524700000000000000000000000000000/
                 66854673 +
                 skoXC1 *
                 (skoXC1 *
                  (11371458976596733414176070000000000000000000000/
                   66854673 +
                   skoXC1 *
                   (skoXC1 *
                    (-1706231753425725563862139000000000000000/66854673 +
                     skoXC1 *
                     (skoXC1 *
                      (71132938024799332185890000000000/66854673 +
                       skoXC1 *
                       (skoXC1 *
                        (-88449140115525021685000/2476099 +
                         skoXC1 *
                         (skoXC1 *
                          (9249321356054667/9904396 +
                           skoXC1 *
                           (skoXC1 *
                            (-17697340451/990439600 +
                             skoXC1 *
                             (skoXC1 *
                              1536507/
                              7923516800000)))))))))))))))))))))))))),
    inference(decision, [],
              [refute_0_61, refute_0_50, refute_0_20, refute_0_21])).

cnf(refute_0_72, plain,
    ((100000000000/57 +
      skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))) /
     (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <= 1),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (0 <= skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29)) |
     1 * (skoXC1 * (-90000000 + skoXC1 * (skoXC1 * 29))) <=
     100000000000/57 +
     skoXC1 * (skoXC1 * (-6533000/57 + skoXC1 * (skoXC1 * 1/380)))),
    inference(resolve, [], [refute_0_72, refute_0_0])).

cnf(refute_0_74, plain,
    (skoXC1 *
     (-90000000 +
      skoXC1 * (6533000/57 + skoXC1 * (29 + skoXC1 * -1/380))) <=
     100000000000/57 | skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (skoXC1 * (90000000 + skoXC1 * (skoXC1 * -29)) <= 0),
    inference(resolve, [], [refute_0_74, refute_0_44])).

cnf(refute_0_76, plain, ($false),
    inference(resolve, [], [refute_0_75, refute_0_24])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-2.tptp
