% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-8.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_14, axiom,
    (X <= 1 | ~ lgen(R, Y, 157079/100000 - 1 / X) |
     lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_1_8', conjecture,
    (! [X] :
       (X = 354991/512 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            4000000000000000 * X * (X + 10 ^ 30) /
            (X ^ 2 + 6000000000000000000000000000000 * X + 10 ^ 60)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 354991/512 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_8'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 354991/512 =>
          20 / ln(10) *
          ln(36400000000000 /
             sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2)) <
          -12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18 <=
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) & X = 354991/512)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) &
     skoXC1 = 354991/512), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 354991/512),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (X_000056 < arctan(X_000055) | ~ lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1 |
     lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [atan_lower_bound_case_14])).

cnf(refute_0_2, plain,
    (X_000056 < arctan(X_000055) |
     ~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000056 < arctan(X_000055) | X_000055 <= 1 |
     157079/100000 + 1 / X_000055 * -1 <= X_000056),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (pi *
     -54076248728958887986448302458966633355972231314007660375461397699724470542679015266422733095607393428984388573203724199619/
     3920497722881471140184559797121600022450219928947303652767431525515523120941336359649715592435318364419158592261139434755000 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     -54076248728958887986448302458966633355972231314007660375461397699724470542679015266422733095607393428984388573203724199619/
     3920497722881471140184559797121600022450219928947303652767431525515523120941336359649715592435318364419158592261139434755000),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <
     73995764428903618260668580467154347090393780584872071905566040940841132682655741069536854292779194018814934022825061331524000000000000000/
     18525560954031318313265767858100033726439248851683148009120667587600031696975361764127555653194798504703733505706265332881 |
     73995764428903618260668580467154347090393780584872071905566040940841132682655741069536854292779194018814934022825061331524000000000000000/
     18525560954031318313265767858100033726439248851683148009120667587600031696975361764127555653194798504703733505706265332881 *
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <= pi * 728000000000000 |
     pi * 5/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 270/7 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (45/14 <
     (pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) * 12)) |
     0 <=
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * -6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) |
     45/14 *
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <=
     pi * 728000000000000 +
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) * (pi * -20)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_7, plain,
    ((36400000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * -1) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (36400000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * -1) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_8, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain, (skoXC1 = 354991/512),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_10, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_11, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 354991/512 |
     skoXC1 ^ 2 = 354991/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_12, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain, (skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_9, refute_0_12])).

cnf(refute_0_14, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_15, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_16, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain, (29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_13, refute_0_16])).

cnf(refute_0_18, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_19, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_20, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_17, refute_0_20])).

cnf(refute_0_22, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_23, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_24, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_21, refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_27, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_28, plain,
    (skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_9, refute_0_28])).

cnf(refute_0_30, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_31, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_32, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_34, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_29, refute_0_35])).

cnf(refute_0_37, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_25, refute_0_36])).

cnf(refute_0_38, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_39, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_40, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_37, refute_0_40])).

cnf(refute_0_42, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_43, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 354991/512 |
     skoXC1 ^ 4 = 354991/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_44, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain, (skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_9, refute_0_44])).

cnf(refute_0_46, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_47, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 354991/512 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain, (3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_45, refute_0_48])).

cnf(refute_0_50, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_51, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_52, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_49, refute_0_52])).

cnf(refute_0_54, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_55, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_56, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_13, refute_0_56])).

cnf(refute_0_58, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_57, refute_0_60])).

cnf(refute_0_62, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_63, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_53, refute_0_63])).

cnf(refute_0_65, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_66, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_67, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_64, refute_0_67])).

cnf(refute_0_69, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_70, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_71, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_68, refute_0_71])).

cnf(refute_0_73, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_75, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_72, refute_0_75])).

cnf(refute_0_77, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_78, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_41, refute_0_78])).

cnf(refute_0_80, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_81, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_82, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_80, refute_0_81])).

cnf(refute_0_83, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_79, refute_0_82])).

cnf(refute_0_84, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_85, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_86, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_83, refute_0_86])).

cnf(refute_0_88, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_89, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_90, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_88, refute_0_89])).

cnf(refute_0_91, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_87, refute_0_90])).

cnf(refute_0_92, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_93, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_94, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_91, refute_0_94])).

cnf(refute_0_96, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_97, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_99, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_100, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_13, refute_0_100])).

cnf(refute_0_102, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_103, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_104, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_101, refute_0_104])).

cnf(refute_0_106, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_107, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_108, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_45, refute_0_108])).

cnf(refute_0_110, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_111, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_112, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_109, refute_0_112])).

cnf(refute_0_114, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_115, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 354991/512 |
     skoXC1 ^ 6 = 354991/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_116, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain, (skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_9, refute_0_116])).

cnf(refute_0_118, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_119, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_120, plain,
    (skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_117, refute_0_120])).

cnf(refute_0_122, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_123, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_124, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_121, refute_0_124])).

cnf(refute_0_126, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_127, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 354991/512 |
     skoXC1 ^ 8 = 354991/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_128, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain, (skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_9, refute_0_128])).

cnf(refute_0_130, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_131, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_130, refute_0_131])).

cnf(refute_0_133, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_129, refute_0_132])).

cnf(refute_0_134, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_133, refute_0_136])).

cnf(refute_0_138, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 354991/512 |
     skoXC1 ^ 10 = 354991/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain, (skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_9, refute_0_140])).

cnf(refute_0_142, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_143, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_144, plain,
    (skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain, (9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_141, refute_0_144])).

cnf(refute_0_146, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_147, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_145, refute_0_148])).

cnf(refute_0_150, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_151, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_137, refute_0_151])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_155, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_152, refute_0_155])).

cnf(refute_0_157, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_158, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_125, refute_0_158])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_162, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_159, refute_0_162])).

cnf(refute_0_164, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_165, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_163, refute_0_164])).

cnf(refute_0_166, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_113, refute_0_165])).

cnf(refute_0_167, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_169, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_166, refute_0_169])).

cnf(refute_0_171, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_172, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_105, refute_0_172])).

cnf(refute_0_174, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_175, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_176, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_174, refute_0_175])).

cnf(refute_0_177, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_173, refute_0_176])).

cnf(refute_0_178, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_179, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_180, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_177, refute_0_180])).

cnf(refute_0_182, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_183, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_184, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_181, refute_0_184])).

cnf(refute_0_186, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_187, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_188, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_185, refute_0_188])).

cnf(refute_0_190, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2)) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_191, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_189, refute_0_190])).

cnf(refute_0_192, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_97, refute_0_191])).

cnf(refute_0_193, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
             10300000000000000 * 354991/512 ^ 6 +
             10 ^ 22 * 354991/512 ^ 4 +
             4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_8, refute_0_192])).

cnf(refute_0_194, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * -20 / ln(10) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     pi | pi = 0), inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    ((pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     pi <
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * -20 / ln(10) |
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_196, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0 | pi = 0),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0 | pi = 0), inference(arithmetic, [], [refute_0_196])).

cnf(refute_0_198, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_199, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_200, plain, (0 < pi | pi = 0),
    inference(decision, [], [refute_0_9, refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi = 0), inference(resolve, [], [refute_0_197, refute_0_200])).

cnf(refute_0_202, plain, (pi != 0),
    inference(decision, [], [refute_0_9, refute_0_198, refute_0_199])).

cnf(refute_0_203, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12),
    inference(resolve, [], [refute_0_201, refute_0_202])).

cnf(refute_0_204, plain,
    (pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 <
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) / ln(10) |
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) *
     ln(10) | ln(10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_205, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) *
     ln(10) | ln(10) <= 0),
    inference(resolve, [], [refute_0_203, refute_0_204])).

cnf(refute_0_206, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    (0 <= pi * -20 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224))),
    inference(split, [], [refute_0_206])).

cnf(refute_0_208, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) | pi <= 0),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain, (0 < pi),
    inference(decision, [], [refute_0_9, refute_0_198, refute_0_199])).

cnf(refute_0_210, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224))),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_212, plain,
    (~ lgen(1, X_000048 - 1, X_000049) | X_000048 <= 0 |
     lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_213, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, X_000048 - 1, X_000049) |
     X_000048 <= 0), inference(resolve, [], [refute_0_212, refute_0_211])).

cnf(refute_0_214, plain,
    (ln(X_000048) < X_000049 | X_000048 <= 0 | X_000049 <= -1 + X_000048),
    inference(arithmetic, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (ln(36400000000000 /
        sqrt(4300420612494076722707453206704419591401332323901082498509609/
             1237940039285380274899124224)) <
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) |
     36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224)),
    inference(subst, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_210, refute_0_215])).

cnf(refute_0_217, plain,
    (36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * -1) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (0 <
     36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) |
     36400000000000 <=
     0 *
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_219, plain,
    (0 <
     36400000000000 /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * -1) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_217, refute_0_219])).

cnf(refute_0_221, plain,
    (0 <= pi * -20 |
     (36400000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * -1) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(resolve, [], [refute_0_220, refute_0_7])).

cnf(refute_0_222, plain,
    ((pi * -728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * 20)) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi <= 0), inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    ((pi * -728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * 20)) /
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(resolve, [], [refute_0_222, refute_0_209])).

cnf(refute_0_224, plain,
    (0 <=
     pi * -6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 12 |
     -((pi * -728000000000000 +
       sqrt(4300420612494076722707453206704419591401332323901082498509609/
            1237940039285380274899124224) * (pi * 20)) /
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224)) /
     (pi * -6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) <=
     ln(10)), inference(split, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (pi * 1/2 <=
     arctan(4131275018871243593760317/4131263554720769638778880) |
     (pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <= ln(10)), inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (arctan(X_000041) < X_000042 | ~ lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_227, plain,
    (X_000041 < 0 | ~ lgen(1, X_000041, X_000042) |
     lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_228, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 |
     ~ lgen(1, X_000041, X_000042)),
    inference(resolve, [], [refute_0_227, refute_0_226])).

cnf(refute_0_229, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 | X_000042 <= X_000041),
    inference(arithmetic, [], [refute_0_228])).

cnf(refute_0_230, plain,
    (4131275018871243593760317/4131263554720769638778880 < 0 |
     arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2 |
     pi * 1/2 <= 4131275018871243593760317/4131263554720769638778880),
    inference(subst, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2 | pi <= 4131275018871243593760317/2065631777360384819389440),
    inference(arithmetic, [], [refute_0_230])).

cnf(refute_0_232, plain,
    (4131275018871243593760317/2065631777360384819389440 < pi),
    inference(decision, [], [refute_0_199, refute_0_198, refute_0_9])).

cnf(refute_0_233, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2), inference(resolve, [], [refute_0_231, refute_0_232])).

cnf(refute_0_234, plain,
    ((pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <= ln(10)),
    inference(resolve, [], [refute_0_225, refute_0_233])).

cnf(refute_0_235, plain,
    (ln(X_000111) < X_000112 | ~ lgen(1, ln(X_000111), X_000112)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_236, plain,
    (~ lgen(1, 1/2 * (X_000111 + 5) * (X_000111 - 1) / (2 * X_000111 + 1),
            X_000112) | X_000111 <= 0 | lgen(1, ln(X_000111), X_000112)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_237, plain,
    (ln(X_000111) < X_000112 |
     ~ lgen(1, 1/2 * (X_000111 + 5) * (X_000111 - 1) / (2 * X_000111 + 1),
            X_000112) | X_000111 <= 0),
    inference(resolve, [], [refute_0_236, refute_0_235])).

cnf(refute_0_238, plain,
    (ln(X_000111) < X_000112 | X_000111 <= 0 |
     X_000112 <=
     (-5/2 + X_000111 * (2 + X_000111 * 1/2)) / (1 + X_000111 * 2)),
    inference(arithmetic, [], [refute_0_237])).

cnf(refute_0_239, plain,
    (ln(10) <
     (pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) * 12)) |
     10 <= 0 |
     (pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <= (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(subst, [], [refute_0_238])).

cnf(refute_0_240, plain,
    (10 <= 0 |
     (pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <= (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(resolve, [], [refute_0_234, refute_0_239])).

cnf(refute_0_241, plain,
    ((pi * 728000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <= 45/14), inference(arithmetic, [], [refute_0_240])).

cnf(refute_0_242, plain,
    (0 <=
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * -6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) |
     45/14 *
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12)) <=
     pi * 728000000000000 +
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) * (pi * -20)),
    inference(resolve, [], [refute_0_241, refute_0_6])).

cnf(refute_0_243, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <= pi * 728000000000000 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) <=
     0), inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_245, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 <=
     0 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     0 /
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(split, [], [refute_0_244])).

cnf(refute_0_246, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi * 1/2 <=
     arctan(4131275018871243593760317/4131263554720769638778880)),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (4131275018871243593760317/4131263554720769638778880 < 0 |
     0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi * 1/2 <= 4131275018871243593760317/4131263554720769638778880),
    inference(resolve, [], [refute_0_246, refute_0_230])).

cnf(refute_0_248, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi <= 4131275018871243593760317/2065631777360384819389440),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(resolve, [], [refute_0_248, refute_0_232])).

cnf(refute_0_250, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) *
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <= pi * 728000000000000),
    inference(resolve, [], [refute_0_243, refute_0_249])).

cnf(refute_0_251, plain,
    (pi * 5/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 270/7 <=
     0 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(split, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi * -1/54 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(arithmetic, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi * -1/54), inference(subst, [], [refute_0_3])).

cnf(refute_0_254, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     pi <=
     -6366875216864375931832498519161/206563750943562179688015850000),
    inference(arithmetic, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (-6366875216864375931832498519161/206563750943562179688015850000 < pi),
    inference(decision, [], [refute_0_9, refute_0_198, refute_0_199])).

cnf(refute_0_256, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880)),
    inference(resolve, [], [refute_0_254, refute_0_255])).

cnf(refute_0_257, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(resolve, [], [refute_0_252, refute_0_256])).

cnf(refute_0_258, plain,
    (X_000126 < sqrt(X_000125) | ~ lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_259, plain,
    (X_000125 < 0 |
     ~ lgen(1, X_000126,
            4000000000000000 * X_000125 * (X_000125 + 10 ^ 30) /
            (X_000125 ^ 2 + 6000000000000000000000000000000 * X_000125 +
             10 ^ 60)) | lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_260, plain,
    (X_000125 < 0 | X_000126 < sqrt(X_000125) |
     ~ lgen(1, X_000126,
            4000000000000000 * X_000125 * (X_000125 + 10 ^ 30) /
            (X_000125 ^ 2 + 6000000000000000000000000000000 * X_000125 +
             10 ^ 60))),
    inference(resolve, [], [refute_0_259, refute_0_258])).

cnf(refute_0_261, plain,
    (X_000125 < 0 | X_000126 < sqrt(X_000125) |
     X_000125 *
     (4000000000000000000000000000000000000000000000 +
      X_000125 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      X_000125 * (6000000000000000000000000000000 + X_000125)) <=
     X_000126), inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224 < 0 |
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224) |
     4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000000 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224 *
      (6000000000000000000000000000000 +
       4300420612494076722707453206704419591401332323901082498509609/
       1237940039285380274899124224)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(subst, [], [refute_0_261])).

cnf(refute_0_263, plain,
    (4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224 < 0 |
     4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000000 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224 *
      (6000000000000000000000000000000 +
       4300420612494076722707453206704419591401332323901082498509609/
       1237940039285380274899124224)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(resolve, [], [refute_0_257, refute_0_262])).

cnf(refute_0_264, plain,
    (73995764428903618260668580467154347090393780584872071905566040940841132682655741069536854292779194018814934022825061331524000000000000000/
     18525560954031318313265767858100033726439248851683148009120667587600031696975361764127555653194798504703733505706265332881 <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7)), inference(arithmetic, [], [refute_0_263])).

cnf(refute_0_265, plain,
    (73995764428903618260668580467154347090393780584872071905566040940841132682655741069536854292779194018814934022825061331524000000000000000/
     18525560954031318313265767858100033726439248851683148009120667587600031696975361764127555653194798504703733505706265332881 *
     (pi * 5/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270/7) <= pi * 728000000000000 |
     pi * 5/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 270/7 <=
     0), inference(resolve, [], [refute_0_264, refute_0_5])).

cnf(refute_0_266, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi * -1/54 |
     arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     -54076248728958887986448302458966633355972231314007660375461397699724470542679015266422733095607393428984388573203724199619/
     3920497722881471140184559797121600022450219928947303652767431525515523120941336359649715592435318364419158592261139434755000),
    inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     -54076248728958887986448302458966633355972231314007660375461397699724470542679015266422733095607393428984388573203724199619/
     3920497722881471140184559797121600022450219928947303652767431525515523120941336359649715592435318364419158592261139434755000),
    inference(resolve, [], [refute_0_266, refute_0_256])).

cnf(refute_0_268, plain,
    (4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     -54076248728958887986448302458966633355972231314007660375461397699724470542679015266422733095607393428984388573203724199619/
     3920497722881471140184559797121600022450219928947303652767431525515523120941336359649715592435318364419158592261139434755000),
    inference(resolve, [], [refute_0_267, refute_0_4])).

cnf(refute_0_269, plain,
    (pi <=
     -184898665108053767263134782268763414650328202002365073763011566535331896754631766631927926255433177439609531026254221492879765970904998932926504272893/
     4468077109764313847343382208413808233187736380442843787226561135824949002292925698070011736823654101973565855039341712032385987552232855266974384460),
    inference(arithmetic, [], [refute_0_268])).

cnf(refute_0_270, plain,
    (-184898665108053767263134782268763414650328202002365073763011566535331896754631766631927926255433177439609531026254221492879765970904998932926504272893/
     4468077109764313847343382208413808233187736380442843787226561135824949002292925698070011736823654101973565855039341712032385987552232855266974384460 <
     pi),
    inference(decision, [], [refute_0_9, refute_0_198, refute_0_199])).

cnf(refute_0_271, plain, ($false),
    inference(resolve, [], [refute_0_269, refute_0_270])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-8.tptp
