% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-2-5.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, -157079/100000 - 1 / X, Y) |
     lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_31, axiom,
    (-1 <= X |
     ~ lgen(R,
            -157079/100000 -
            1/15 * (64 + 735 * X ^ 2 + 945 * X ^ 4) /
            (X * (15 + 70 * X ^ 2 + 63 * X ^ 4)), Y) |
     lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (6 * X ^ 4 + 481 * X ^ 3 + 1881 * X ^ 2 + 1281 * X + 131) *
            (X - 1) / (5 * X ^ 4 + 40 * X ^ 3 + 60 * X ^ 2 + 20 * X + 1),
            Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_2_5', conjecture,
    (! [X] :
       (X = 9745/512 =>
        20 / ln(10) *
        ln(33600000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            400000000000000 * X * (X + 10 ^ 28) /
            (X ^ 2 + 60000000000000000000000000000 * X + 10 ^ 56)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 9745/512 =>
        20 / ln(10) *
        ln(33600000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_2_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 9745/512 =>
          20 / ln(10) *
          ln(33600000000000 /
             sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2)) <
          12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) + 6)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_1, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_2, plain,
    (? [X] :
       (12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6 <=
        20 / ln(10) *
        ln(33600000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) & X = 9745/512)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_3, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) & skoXC1 = 9745/512),
    inference(skolemize, [], [normalize_0_2])).

fof(normalize_0_4, plain, (skoXC1 = 9745/512),
    inference(conjunct, [], [normalize_0_3])).

fof(normalize_0_5, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_3])).

cnf(refute_0_0, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, X_000034, X_000033) | 0 <= X_000033 |
     lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_2, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000034 < arctan(X_000033) | X_000033 <= X_000034 | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (pi * -1/2 < arctan(-214698021844321510003/214688297512456402944) |
     -214698021844321510003/214688297512456402944 <= pi * -1/2 |
     0 <= -214698021844321510003/214688297512456402944),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (45/14 <
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     0 <=
     pi * -6 + arctan(-214698021844321510003/214688297512456402944) * -12 |
     45/14 *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20), inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_6, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20), introduced(tautology, [assume])).

cnf(refute_0_7, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     0 <=
     pi * -6 + arctan(-214698021844321510003/214688297512456402944) * -12 |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     ln(10)), inference(split, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     ln(10) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (pi * -1/2 < arctan(-214698021844321510003/214688297512456402944) |
     pi <= 214698021844321510003/107344148756228201472),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_10, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_0])).

cnf(refute_0_11, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_1])).

cnf(refute_0_12, plain, (skoXC1 = 9745/512),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_13, plain, (214698021844321510003/107344148756228201472 < pi),
    inference(decision, [], [refute_0_10, refute_0_11, refute_0_12])).

cnf(refute_0_14, plain,
    (pi * -1/2 < arctan(-214698021844321510003/214688297512456402944)),
    inference(resolve, [], [refute_0_9, refute_0_13])).

cnf(refute_0_15, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     ln(10)), inference(resolve, [], [refute_0_8, refute_0_14])).

cnf(refute_0_16, plain,
    (ln(X_000103) < X_000104 | ~ lgen(1, ln(X_000103), X_000104)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain,
    (~ lgen(1, 1/2 * (X_000103 + 5) * (X_000103 - 1) / (2 * X_000103 + 1),
            X_000104) | X_000103 <= 0 | lgen(1, ln(X_000103), X_000104)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_18, plain,
    (ln(X_000103) < X_000104 |
     ~ lgen(1, 1/2 * (X_000103 + 5) * (X_000103 - 1) / (2 * X_000103 + 1),
            X_000104) | X_000103 <= 0),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (ln(X_000103) < X_000104 | X_000103 <= 0 |
     X_000104 <=
     (-5/2 + X_000103 * (2 + X_000103 * 1/2)) / (1 + X_000103 * 2)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (ln(10) <
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     10 <= 0 |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     10 <= 0 |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(resolve, [], [refute_0_15, refute_0_20])).

cnf(refute_0_22, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     pi * -20 /
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     45/14), inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     0 <=
     pi * -6 + arctan(-214698021844321510003/214688297512456402944) * -12 |
     45/14 *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20), inference(resolve, [], [refute_0_22, refute_0_5])).

cnf(refute_0_24, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     pi * 1/54 <= arctan(-214698021844321510003/214688297512456402944) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (arctan(X_000061) < X_000062 | ~ lgen(1, arctan(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_26, plain,
    (~ lgen(1, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     lgen(1, arctan(X_000061), X_000062)),
    inference(subst, [], [atan_upper_bound_case_11])).

cnf(refute_0_27, plain,
    (arctan(X_000061) < X_000062 |
     ~ lgen(1, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061),
    inference(resolve, [], [refute_0_26, refute_0_25])).

cnf(refute_0_28, plain,
    (arctan(X_000061) < X_000062 |
     X_000062 <= -157079/100000 + 1 / X_000061 * -1 | -1 <= X_000061),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (arctan(-214698021844321510003/214688297512456402944) < pi * 1/54 |
     -1 <= -214698021844321510003/214688297512456402944 |
     pi * 1/54 <=
     -157079/100000 +
     1 / -214698021844321510003/214688297512456402944 * -1),
    inference(subst, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (arctan(-214698021844321510003/214688297512456402944) < pi * 1/54 |
     pi <= -330904462195040530734753399/10734901092216075500150000),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (-330904462195040530734753399/10734901092216075500150000 < pi),
    inference(decision, [], [refute_0_12, refute_0_11, refute_0_10])).

cnf(refute_0_32, plain,
    (arctan(-214698021844321510003/214688297512456402944) < pi * 1/54),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(resolve, [], [refute_0_24, refute_0_32])).

cnf(refute_0_34, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     -214698021844321510003/214688297512456402944 <= pi * -1/2 |
     0 <= -214698021844321510003/214688297512456402944),
    inference(resolve, [], [refute_0_33, refute_0_4])).

cnf(refute_0_35, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) |
     pi <= 214698021844321510003/107344148756228201472),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (pi * -20 <
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12)),
    inference(resolve, [], [refute_0_35, refute_0_13])).

cnf(refute_0_37, plain,
    (pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <
     226558021396570388307297241258246181577573066887293213283983647340156831066313418283506650474432400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 |
     226558021396570388307297241258246181577573066887293213283983647340156831066313418283506650474432400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 *
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <=
     pi * 672000000000000 |
     pi * 20 +
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_38, plain,
    (33600000000000 +
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) * -1 <
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (ln(10) *
      (pi * 6 +
       arctan(-214698021844321510003/214688297512456402944) * 12)) /
     (pi * -20) | 0 <= pi * -20 |
     (33600000000000 +
      sqrt(3419047384141490680217341264051304690187335712978515625/
           1237940039285380274899124224) * -1) * (pi * -20) <=
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (ln(10) *
      (pi * 6 +
       arctan(-214698021844321510003/214688297512456402944) * 12))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_39, plain,
    ((33600000000000 +
      sqrt(3419047384141490680217341264051304690187335712978515625/
           1237940039285380274899124224) * -1) /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) |
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) *
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <=
     33600000000000 +
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) * -1 |
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_40, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_41, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_42, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 9745/512 |
     skoXC1 ^ 2 = 9745/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_43, plain, (skoXC1 != 9745/512 | skoXC1 ^ 2 = 9745/512 ^ 2),
    inference(resolve, [], [refute_0_41, refute_0_42])).

cnf(refute_0_44, plain, (skoXC1 ^ 2 = 9745/512 ^ 2),
    inference(resolve, [], [refute_0_12, refute_0_43])).

cnf(refute_0_45, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_46, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 9745/512 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 9745/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_47, plain,
    (skoXC1 ^ 2 != 9745/512 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain, (29 * skoXC1 ^ 2 = 29 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_44, refute_0_47])).

cnf(refute_0_49, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_50, plain,
    (29 * skoXC1 ^ 2 != 29 * 9745/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 9745/512 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_51, plain,
    (29 * skoXC1 ^ 2 != 29 * 9745/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 9745/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 9745/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_48, refute_0_51])).

cnf(refute_0_53, plain,
    (9745/512 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_54, plain,
    (9745/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 9745/512 ^ 2 - 90000000 |
     9745/512 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_55, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 9745/512 ^ 2 - 90000000 |
     9745/512 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (9745/512 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_52, refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_58, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 9745/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_59, plain,
    (skoXC1 != 9745/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_12, refute_0_59])).

cnf(refute_0_61, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_62, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_63, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (9745/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * 9745/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (9745/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * 9745/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_60, refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     9745/512 * (29 * 9745/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_56, refute_0_67])).

cnf(refute_0_69, plain,
    (1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_70, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * 9745/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_71, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     9745/512 * (29 * 9745/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_68, refute_0_71])).

cnf(refute_0_73, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_74, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 9745/512 |
     skoXC1 ^ 4 = 9745/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_75, plain, (skoXC1 != 9745/512 | skoXC1 ^ 4 = 9745/512 ^ 4),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain, (skoXC1 ^ 4 = 9745/512 ^ 4),
    inference(resolve, [], [refute_0_12, refute_0_75])).

cnf(refute_0_77, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_78, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 9745/512 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 9745/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_79, plain,
    (skoXC1 ^ 4 != 9745/512 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain, (3 * skoXC1 ^ 4 = 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_76, refute_0_79])).

cnf(refute_0_81, plain,
    (-130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_82, plain,
    (3 * skoXC1 ^ 4 != 3 * 9745/512 ^ 4 |
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_83, plain,
    (3 * skoXC1 ^ 4 != 3 * 9745/512 ^ 4 |
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (-130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_80, refute_0_83])).

cnf(refute_0_85, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 9745/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 9745/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (skoXC1 ^ 2 != 9745/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_44, refute_0_87])).

cnf(refute_0_89, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 9745/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 9745/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (-130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_94, plain,
    (-130660000 * 9745/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_84, refute_0_94])).

cnf(refute_0_96, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_97, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_98, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_96, refute_0_97])).

cnf(refute_0_99, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_95, refute_0_98])).

cnf(refute_0_100, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_101, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_102, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_99, refute_0_102])).

cnf(refute_0_104, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_105, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_106, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_103, refute_0_106])).

cnf(refute_0_108, plain,
    (1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_109, plain,
    (1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_72, refute_0_109])).

cnf(refute_0_111, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 + 2000000000000) /
     (9745/512 * (29 * 9745/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_110, refute_0_113])).

cnf(refute_0_115, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
             2000000000000) /
            (9745/512 * (29 * 9745/512 ^ 2 - 90000000))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_118, refute_0_121])).

cnf(refute_0_123, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_128, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_130, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 | skoXC1 ^ 2 != 9745/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 9745/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_131, plain,
    (skoXC1 ^ 2 != 9745/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_44, refute_0_131])).

cnf(refute_0_133, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_134, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 9745/512 ^ 2 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_135, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 9745/512 ^ 2 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_132, refute_0_135])).

cnf(refute_0_137, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_138, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 9745/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 9745/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_139, plain,
    (skoXC1 ^ 4 != 9745/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_76, refute_0_139])).

cnf(refute_0_141, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_142, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 9745/512 ^ 4 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_143, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 9745/512 ^ 4 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_140, refute_0_143])).

cnf(refute_0_145, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_146, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 9745/512 |
     skoXC1 ^ 6 = 9745/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_147, plain, (skoXC1 != 9745/512 | skoXC1 ^ 6 = 9745/512 ^ 6),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain, (skoXC1 ^ 6 = 9745/512 ^ 6),
    inference(resolve, [], [refute_0_12, refute_0_147])).

cnf(refute_0_149, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_150, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 9745/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 9745/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_151, plain,
    (skoXC1 ^ 6 != 9745/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_148, refute_0_151])).

cnf(refute_0_153, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_154, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 9745/512 ^ 6 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_155, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 9745/512 ^ 6 |
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_152, refute_0_155])).

cnf(refute_0_157, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_158, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 9745/512 |
     skoXC1 ^ 8 = 9745/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_159, plain, (skoXC1 != 9745/512 | skoXC1 ^ 8 = 9745/512 ^ 8),
    inference(resolve, [], [refute_0_157, refute_0_158])).

cnf(refute_0_160, plain, (skoXC1 ^ 8 = 9745/512 ^ 8),
    inference(resolve, [], [refute_0_12, refute_0_159])).

cnf(refute_0_161, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_162, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 9745/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 9745/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_163, plain,
    (skoXC1 ^ 8 != 9745/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_160, refute_0_163])).

cnf(refute_0_165, plain,
    (9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 9745/512 ^ 8 |
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 9745/512 ^ 8 |
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_164, refute_0_167])).

cnf(refute_0_169, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_170, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 9745/512 |
     skoXC1 ^ 10 = 9745/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_171, plain,
    (skoXC1 != 9745/512 | skoXC1 ^ 10 = 9745/512 ^ 10),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain, (skoXC1 ^ 10 = 9745/512 ^ 10),
    inference(resolve, [], [refute_0_12, refute_0_171])).

cnf(refute_0_173, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_174, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 9745/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 9745/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_175, plain,
    (skoXC1 ^ 10 != 9745/512 ^ 10 | 9 * skoXC1 ^ 10 = 9 * 9745/512 ^ 10),
    inference(resolve, [], [refute_0_173, refute_0_174])).

cnf(refute_0_176, plain, (9 * skoXC1 ^ 10 = 9 * 9745/512 ^ 10),
    inference(resolve, [], [refute_0_172, refute_0_175])).

cnf(refute_0_177, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_178, plain,
    (9 * skoXC1 ^ 10 != 9 * 9745/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_179, plain,
    (9 * skoXC1 ^ 10 != 9 * 9745/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_176, refute_0_179])).

cnf(refute_0_181, plain,
    (9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_182, plain,
    (9 * 9745/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8),
    inference(resolve, [], [refute_0_168, refute_0_182])).

cnf(refute_0_184, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_185, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_186, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_184, refute_0_185])).

cnf(refute_0_187, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_183, refute_0_186])).

cnf(refute_0_188, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_189, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6),
    inference(resolve, [], [refute_0_156, refute_0_189])).

cnf(refute_0_191, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_192, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_190, refute_0_193])).

cnf(refute_0_195, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_196, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4),
    inference(resolve, [], [refute_0_144, refute_0_196])).

cnf(refute_0_198, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_199, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_200, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_197, refute_0_200])).

cnf(refute_0_202, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(subst, [], [refute_0_65])).

cnf(refute_0_203, plain,
    (9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_201, refute_0_202])).

cnf(refute_0_204, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2),
    inference(resolve, [], [refute_0_136, refute_0_203])).

cnf(refute_0_205, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_206, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_207, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
     10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
     4000000000000000000000000 * 9745/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    inference(resolve, [], [refute_0_205, refute_0_206])).

cnf(refute_0_208, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    inference(resolve, [], [refute_0_204, refute_0_207])).

cnf(refute_0_209, plain,
    (33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_210, plain,
    (33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2) |
     33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 /
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_211, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2) |
     33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 /
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    inference(resolve, [], [refute_0_209, refute_0_210])).

cnf(refute_0_212, plain,
    (33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 /
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2)),
    inference(resolve, [], [refute_0_208, refute_0_211])).

cnf(refute_0_213, plain,
    (ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_214, plain,
    (33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 /
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2) |
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_215, plain,
    (33600000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 /
     sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
          10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
          4000000000000000000000000 * 9745/512 ^ 2) |
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_213, refute_0_214])).

cnf(refute_0_216, plain,
    (ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_212, refute_0_215])).

cnf(refute_0_217, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_218, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2)) |
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_219, plain,
    (ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2)) |
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_217, refute_0_218])).

cnf(refute_0_220, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_216, refute_0_219])).

cnf(refute_0_221, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 |
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2)) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_222, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_220, refute_0_221])).

cnf(refute_0_223, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_128, refute_0_222])).

cnf(refute_0_224, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 9745/512 ^ 2 + 3 * 9745/512 ^ 4 +
              2000000000000) /
             (9745/512 * (29 * 9745/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(33600000000000 /
        sqrt(9 * 9745/512 ^ 10 + 309000000 * 9745/512 ^ 8 +
             10300000000000000 * 9745/512 ^ 6 + 10 ^ 22 * 9745/512 ^ 4 +
             4000000000000000000000000 * 9745/512 ^ 2))),
    inference(resolve, [], [refute_0_40, refute_0_223])).

cnf(refute_0_225, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * -20 / ln(10) <=
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     pi | pi = 0), inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    ((pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     pi <
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * -20 / ln(10) |
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12 |
     pi <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_227, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12 |
     pi <= 0 | pi = 0),
    inference(resolve, [], [refute_0_225, refute_0_226])).

cnf(refute_0_228, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12 |
     pi <= 0 | pi = 0), inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain, (0 < pi | pi = 0),
    inference(decision, [], [refute_0_12, refute_0_11, refute_0_10])).

cnf(refute_0_230, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12 |
     pi = 0), inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain, (pi != 0),
    inference(decision, [], [refute_0_12, refute_0_11, refute_0_10])).

cnf(refute_0_232, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12 <
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) / ln(10) |
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) *
     ln(10) | ln(10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_234, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) *
     ln(10) | ln(10) <= 0),
    inference(resolve, [], [refute_0_232, refute_0_233])).

cnf(refute_0_235, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12)),
    inference(arithmetic, [], [refute_0_234])).

cnf(refute_0_236, plain,
    (0 <= pi * -20 |
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224))),
    inference(split, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) | pi <= 0),
    inference(arithmetic, [], [refute_0_236])).

cnf(refute_0_238, plain, (0 < pi),
    inference(decision, [], [refute_0_12, refute_0_11, refute_0_10])).

cnf(refute_0_239, plain,
    (ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224))),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_241, plain,
    (~ lgen(1, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(1, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_242, plain,
    (ln(X_000047) < X_000048 | ~ lgen(1, X_000047 - 1, X_000048) |
     X_000047 <= 0), inference(resolve, [], [refute_0_241, refute_0_240])).

cnf(refute_0_243, plain,
    (ln(X_000047) < X_000048 | X_000047 <= 0 | X_000048 <= -1 + X_000047),
    inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (ln(33600000000000 /
        sqrt(3419047384141490680217341264051304690187335712978515625/
             1237940039285380274899124224)) <
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) |
     33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     -1 +
     33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224)),
    inference(subst, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     -1 +
     33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_239, refute_0_244])).

cnf(refute_0_246, plain,
    (33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     (33600000000000 +
      sqrt(3419047384141490680217341264051304690187335712978515625/
           1237940039285380274899124224) * -1) /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (0 <
     33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) |
     33600000000000 <=
     0 *
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) |
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_248, plain,
    (0 <
     33600000000000 /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) <=
     (33600000000000 +
      sqrt(3419047384141490680217341264051304690187335712978515625/
           1237940039285380274899124224) * -1) /
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_246, refute_0_248])).

cnf(refute_0_250, plain,
    (ln(10) *
     (pi * 6 + arctan(-214698021844321510003/214688297512456402944) * 12) /
     (pi * -20) *
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <=
     33600000000000 +
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) * -1 |
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <= 0),
    inference(resolve, [], [refute_0_249, refute_0_39])).

cnf(refute_0_251, plain,
    (sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (ln(10) *
      (pi * 6 +
       arctan(-214698021844321510003/214688297512456402944) * 12)) /
     (pi * -20) <=
     33600000000000 +
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) * -1),
    inference(arithmetic, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (0 <= pi * -20 |
     (33600000000000 +
      sqrt(3419047384141490680217341264051304690187335712978515625/
           1237940039285380274899124224) * -1) * (pi * -20) <=
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (ln(10) *
      (pi * 6 +
       arctan(-214698021844321510003/214688297512456402944) * 12))),
    inference(resolve, [], [refute_0_251, refute_0_38])).

cnf(refute_0_253, plain,
    (sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <=
     pi * 672000000000000 | pi <= 0),
    inference(arithmetic, [], [refute_0_252])).

cnf(refute_0_254, plain,
    (sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) *
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <=
     pi * 672000000000000),
    inference(resolve, [], [refute_0_253, refute_0_238])).

cnf(refute_0_255, plain,
    (pi * 20 +
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <= 0 |
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(split, [], [refute_0_254])).

cnf(refute_0_256, plain,
    (ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20 |
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(arithmetic, [], [refute_0_255])).

cnf(refute_0_257, plain,
    (sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(resolve, [], [refute_0_256, refute_0_36])).

cnf(refute_0_258, plain,
    (X_000118 < sqrt(X_000117) | ~ lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_259, plain,
    (X_000117 < 0 |
     ~ lgen(1, X_000118,
            400000000000000 * X_000117 * (X_000117 + 10 ^ 28) /
            (X_000117 ^ 2 + 60000000000000000000000000000 * X_000117 +
             10 ^ 56)) | lgen(1, X_000118, sqrt(X_000117))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_260, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     ~ lgen(1, X_000118,
            400000000000000 * X_000117 * (X_000117 + 10 ^ 28) /
            (X_000117 ^ 2 + 60000000000000000000000000000 * X_000117 +
             10 ^ 56))),
    inference(resolve, [], [refute_0_259, refute_0_258])).

cnf(refute_0_261, plain,
    (X_000117 < 0 | X_000118 < sqrt(X_000117) |
     X_000117 *
     (4000000000000000000000000000000000000000000 +
      X_000117 * 400000000000000) /
     (100000000000000000000000000000000000000000000000000000000 +
      X_000117 * (60000000000000000000000000000 + X_000117)) <= X_000118),
    inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (3419047384141490680217341264051304690187335712978515625/
     1237940039285380274899124224 < 0 |
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <
     sqrt(3419047384141490680217341264051304690187335712978515625/
          1237940039285380274899124224) |
     3419047384141490680217341264051304690187335712978515625/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000 +
      3419047384141490680217341264051304690187335712978515625/
      1237940039285380274899124224 * 400000000000000) /
     (100000000000000000000000000000000000000000000000000000000 +
      3419047384141490680217341264051304690187335712978515625/
      1237940039285380274899124224 *
      (60000000000000000000000000000 +
       3419047384141490680217341264051304690187335712978515625/
       1237940039285380274899124224)) <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(subst, [], [refute_0_261])).

cnf(refute_0_263, plain,
    (3419047384141490680217341264051304690187335712978515625/
     1237940039285380274899124224 < 0 |
     3419047384141490680217341264051304690187335712978515625/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000 +
      3419047384141490680217341264051304690187335712978515625/
      1237940039285380274899124224 * 400000000000000) /
     (100000000000000000000000000000000000000000000000000000000 +
      3419047384141490680217341264051304690187335712978515625/
      1237940039285380274899124224 *
      (60000000000000000000000000000 +
       3419047384141490680217341264051304690187335712978515625/
       1237940039285380274899124224)) <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(resolve, [], [refute_0_257, refute_0_262])).

cnf(refute_0_264, plain,
    (226558021396570388307297241258246181577573066887293213283983647340156831066313418283506650474432400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <=
     pi * 672000000000000 /
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12))),
    inference(arithmetic, [], [refute_0_263])).

cnf(refute_0_265, plain,
    (226558021396570388307297241258246181577573066887293213283983647340156831066313418283506650474432400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 *
     (pi * 20 +
      ln(10) *
      (pi * -6 +
       arctan(-214698021844321510003/214688297512456402944) * -12)) <=
     pi * 672000000000000 |
     pi * 20 +
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <= 0),
    inference(resolve, [], [refute_0_264, refute_0_37])).

cnf(refute_0_266, plain,
    (ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 |
     ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20), inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (82449/35630 <
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     0 <=
     pi *
     -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
     arctan(-214698021844321510003/214688297512456402944) *
     -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 |
     82449/35630 *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_268, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081),
    introduced(tautology, [assume])).

cnf(refute_0_269, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     0 <=
     pi *
     -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
     arctan(-214698021844321510003/214688297512456402944) *
     -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     ln(10)), inference(split, [], [refute_0_268])).

cnf(refute_0_270, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     ln(10) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_269])).

cnf(refute_0_271, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     ln(10)), inference(resolve, [], [refute_0_270, refute_0_14])).

cnf(refute_0_272, plain,
    (ln(X_000208) < X_000209 | ~ lgen(1, ln(X_000208), X_000209)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_273, plain,
    (~ lgen(1,
            1/30 *
            (6 * X_000208 ^ 4 + 481 * X_000208 ^ 3 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 + 131) * (X_000208 - 1) /
            (5 * X_000208 ^ 4 + 40 * X_000208 ^ 3 + 60 * X_000208 ^ 2 +
             20 * X_000208 + 1), X_000209) | X_000208 <= 0 |
     lgen(1, ln(X_000208), X_000209)),
    inference(subst, [], [ln_upper_bound_cf9])).

cnf(refute_0_274, plain,
    (ln(X_000208) < X_000209 |
     ~ lgen(1,
            1/30 *
            (6 * X_000208 ^ 4 + 481 * X_000208 ^ 3 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 + 131) * (X_000208 - 1) /
            (5 * X_000208 ^ 4 + 40 * X_000208 ^ 3 + 60 * X_000208 ^ 2 +
             20 * X_000208 + 1), X_000209) | X_000208 <= 0),
    inference(resolve, [], [refute_0_273, refute_0_272])).

cnf(refute_0_275, plain,
    (ln(X_000208) < X_000209 | X_000208 <= 0 |
     X_000209 <=
     (-131/30 +
      X_000208 *
      (-115/3 +
       X_000208 *
       (-20 + X_000208 * (140/3 + X_000208 * (95/6 + X_000208 * 1/5))))) /
     (1 +
      X_000208 * (20 + X_000208 * (60 + X_000208 * (40 + X_000208 * 5))))),
    inference(arithmetic, [], [refute_0_274])).

cnf(refute_0_276, plain,
    (ln(10) <
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     10 <= 0 |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     (-131/30 +
      10 * (-115/3 + 10 * (-20 + 10 * (140/3 + 10 * (95/6 + 10 * 1/5))))) /
     (1 + 10 * (20 + 10 * (60 + 10 * (40 + 10 * 5))))),
    inference(subst, [], [refute_0_275])).

cnf(refute_0_277, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     10 <= 0 |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     (-131/30 +
      10 * (-115/3 + 10 * (-20 + 10 * (140/3 + 10 * (95/6 + 10 * 1/5))))) /
     (1 + 10 * (20 + 10 * (60 + 10 * (40 + 10 * 5))))),
    inference(resolve, [], [refute_0_271, refute_0_276])).

cnf(refute_0_278, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 /
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     82449/35630), inference(arithmetic, [], [refute_0_277])).

cnf(refute_0_279, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     0 <=
     pi *
     -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
     arctan(-214698021844321510003/214688297512456402944) *
     -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 |
     82449/35630 *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) <=
     pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081),
    inference(resolve, [], [refute_0_278, refute_0_267])).

cnf(refute_0_280, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     pi *
     -5535811740558404735260924274686650421790758550784262603279532577844611713184009389255521260761616847/
     22237478935864085649462321719644213601058716418798140645299009213748322100698184552448618839245805890 <=
     arctan(-214698021844321510003/214688297512456402944) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (arctan(X_000159) < X_000160 | ~ lgen(1, arctan(X_000159), X_000160)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_282, plain,
    (~ lgen(1,
            -157079/100000 -
            1/15 * (64 + 735 * X_000159 ^ 2 + 945 * X_000159 ^ 4) /
            (X_000159 * (15 + 70 * X_000159 ^ 2 + 63 * X_000159 ^ 4)),
            X_000160) | -1 <= X_000159 |
     lgen(1, arctan(X_000159), X_000160)),
    inference(subst, [], [atan_upper_bound_case_31])).

cnf(refute_0_283, plain,
    (arctan(X_000159) < X_000160 |
     ~ lgen(1,
            -157079/100000 -
            1/15 * (64 + 735 * X_000159 ^ 2 + 945 * X_000159 ^ 4) /
            (X_000159 * (15 + 70 * X_000159 ^ 2 + 63 * X_000159 ^ 4)),
            X_000160) | -1 <= X_000159),
    inference(resolve, [], [refute_0_282, refute_0_281])).

cnf(refute_0_284, plain,
    (arctan(X_000159) < X_000160 |
     X_000160 <=
     -157079/100000 +
     (64/15 + X_000159 * (X_000159 * (49 + X_000159 * (X_000159 * 63)))) /
     (X_000159 *
      (15 + X_000159 * (X_000159 * (70 + X_000159 * (X_000159 * 63))))) *
     -1 | -1 <= X_000159), inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (arctan(-214698021844321510003/214688297512456402944) <
     pi *
     -5535811740558404735260924274686650421790758550784262603279532577844611713184009389255521260761616847/
     22237478935864085649462321719644213601058716418798140645299009213748322100698184552448618839245805890 |
     -1 <= -214698021844321510003/214688297512456402944 |
     pi *
     -5535811740558404735260924274686650421790758550784262603279532577844611713184009389255521260761616847/
     22237478935864085649462321719644213601058716418798140645299009213748322100698184552448618839245805890 <=
     -157079/100000 +
     (64/15 +
      -214698021844321510003/214688297512456402944 *
      (-214698021844321510003/214688297512456402944 *
       (49 +
        -214698021844321510003/214688297512456402944 *
        (-214698021844321510003/214688297512456402944 * 63)))) /
     (-214698021844321510003/214688297512456402944 *
      (15 +
       -214698021844321510003/214688297512456402944 *
       (-214698021844321510003/214688297512456402944 *
        (70 +
         -214698021844321510003/214688297512456402944 *
         (-214698021844321510003/214688297512456402944 * 63))))) * -1),
    inference(subst, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (arctan(-214698021844321510003/214688297512456402944) <
     pi *
     -5535811740558404735260924274686650421790758550784262603279532577844611713184009389255521260761616847/
     22237478935864085649462321719644213601058716418798140645299009213748322100698184552448618839245805890 |
     1309826232640786766986567199566899323086105367618344332613423084535502804696904937596226916332473358208684821485657289523031471770664615240597318151205623253389409414256039285548736145841904107487233699174471/
     415254312868219273589641472343058613987395424965769177986030395379048379279879930191885599020356715444941550927690438622927370059470052454379001775211929328542318780478088454633640720754777358320646976270000 <=
     pi), inference(arithmetic, [], [refute_0_285])).

cnf(refute_0_287, plain,
    (pi <
     1309826232640786766986567199566899323086105367618344332613423084535502804696904937596226916332473358208684821485657289523031471770664615240597318151205623253389409414256039285548736145841904107487233699174471/
     415254312868219273589641472343058613987395424965769177986030395379048379279879930191885599020356715444941550927690438622927370059470052454379001775211929328542318780478088454633640720754777358320646976270000),
    inference(decision, [], [refute_0_10, refute_0_11, refute_0_12])).

cnf(refute_0_288, plain,
    (arctan(-214698021844321510003/214688297512456402944) <
     pi *
     -5535811740558404735260924274686650421790758550784262603279532577844611713184009389255521260761616847/
     22237478935864085649462321719644213601058716418798140645299009213748322100698184552448618839245805890),
    inference(resolve, [], [refute_0_286, refute_0_287])).

cnf(refute_0_289, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     arctan(-214698021844321510003/214688297512456402944) <= pi * -1/2),
    inference(resolve, [], [refute_0_280, refute_0_288])).

cnf(refute_0_290, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     -214698021844321510003/214688297512456402944 <= pi * -1/2 |
     0 <= -214698021844321510003/214688297512456402944),
    inference(resolve, [], [refute_0_289, refute_0_4])).

cnf(refute_0_291, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081) |
     pi <= 214698021844321510003/107344148756228201472),
    inference(arithmetic, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (pi *
     -1579453030926923153013190703849716539368091062555073303711740613320736944109571582953841836691601568000000000000/
     4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 <
     ln(10) *
     (pi *
      -1359348128379422329843783447549477089465438401323759279703901884040940986397880509701039902846594400000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081 +
      arctan(-214698021844321510003/214688297512456402944) *
      -2718696256758844659687566895098954178930876802647518559407803768081881972795761019402079805693188800000000000000/
      4392421721732864007638026966242867696701443861891057979118946924824999519667703545708766626186081)),
    inference(resolve, [], [refute_0_291, refute_0_13])).

cnf(refute_0_293, plain,
    (ln(10) *
     (pi * -6 +
      arctan(-214698021844321510003/214688297512456402944) * -12) <=
     pi * -20), inference(resolve, [], [refute_0_266, refute_0_292])).

cnf(refute_0_294, plain, ($false),
    inference(resolve, [], [refute_0_293, refute_0_36])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-2-5.tptp
