% SZS output start CNFRefutation for Nichols-Plot-Simple-Example-4.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

cnf(log_def, axiom, (log(X) = 868589/2000000 * ln(X))).

fof(nichols_plot_4, conjecture,
    (! [X] :
       (X = 97/125 =>
        10 *
        log((640000 * X ^ 16 + 7684800 * X ^ 14 + 42990418 * X ^ 12 +
             136160432 * X ^ 10 + 338091528 * X ^ 8 - 21346562 * X ^ 6 -
             87425842 * X ^ 4 - 4998610 * X ^ 2 + 10204661) /
            (X ^ 6 + 5 * X ^ 4 + 60 * X ^ 2 + 16) ^ 2) <
        55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55))))).

cnf(log_upper_bound_scaled_cf3, axiom,
    (X <= 0 |
     ~ lgen(R, 1/40000 * (X + 50000) * (X - 10000) / (X + 5000) + 461/50,
            Y) | lgen(R, ln(X), Y))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 97/125 =>
        10 *
        log((640000 * X ^ 16 + 7684800 * X ^ 14 + 42990418 * X ^ 12 +
             136160432 * X ^ 10 + 338091528 * X ^ 8 - 21346562 * X ^ 6 -
             87425842 * X ^ 4 - 4998610 * X ^ 2 + 10204661) /
            (X ^ 6 + 5 * X ^ 4 + 60 * X ^ 2 + 16) ^ 2) <
        55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55)))), inference(strip, [], [nichols_plot_4])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 97/125 =>
          10 *
          log((640000 * X ^ 16 + 7684800 * X ^ 14 + 42990418 * X ^ 12 +
               136160432 * X ^ 10 + 338091528 * X ^ 8 - 21346562 * X ^ 6 -
               87425842 * X ^ 4 - 4998610 * X ^ 2 + 10204661) /
              (X ^ 6 + 5 * X ^ 4 + 60 * X ^ 2 + 16) ^ 2) <
          55 +
          8 *
          arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
                 (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 -
                  1597 * X ^ 2 + 55)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (55 +
        8 *
        arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 - 3194) /
               (800 * X ^ 8 + 4803 * X ^ 6 + 12054 * X ^ 4 - 1597 * X ^ 2 +
                55)) <=
        10 *
        log((640000 * X ^ 16 + 7684800 * X ^ 14 + 42990418 * X ^ 12 +
             136160432 * X ^ 10 + 338091528 * X ^ 8 - 21346562 * X ^ 6 -
             87425842 * X ^ 4 - 4998610 * X ^ 2 + 10204661) /
            (X ^ 6 + 5 * X ^ 4 + 60 * X ^ 2 + 16) ^ 2) & X = 97/125)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) &
     skoXC1 = 97/125), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 97/125),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (X_000063 < arctan(X_000062) | ~ lgen(1, X_000063, arctan(X_000062))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (X_000062 < 0 |
     ~ lgen(1, X_000063, 3 * X_000062 / (X_000062 ^ 2 + 3)) |
     lgen(1, X_000063, arctan(X_000062))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_2, plain,
    (X_000062 < 0 | X_000063 < arctan(X_000062) |
     ~ lgen(1, X_000063, 3 * X_000062 / (X_000062 ^ 2 + 3))),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000062 < 0 | X_000063 < arctan(X_000062) |
     X_000062 * 3 / (3 + X_000062 * X_000062) <= X_000063),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (6317656210882992500/11010804368800679627 < 0 |
     -492772850186751082174258706259547880510056903518817731428359508983077312229088952451/
     260960699184861893241698938008806317823405914321122921161381426904100000000000000000 <
     arctan(6317656210882992500/11010804368800679627) |
     6317656210882992500/11010804368800679627 * 3 /
     (3 +
      6317656210882992500/11010804368800679627 *
      6317656210882992500/11010804368800679627) <=
     -492772850186751082174258706259547880510056903518817731428359508983077312229088952451/
     260960699184861893241698938008806317823405914321122921161381426904100000000000000000),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) <=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_6, plain, (skoXC1 = 97/125),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_7, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_8, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 97/125 |
     skoXC1 ^ 2 = 97/125 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_9, plain, (skoXC1 != 97/125 | skoXC1 ^ 2 = 97/125 ^ 2),
    inference(resolve, [], [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain, (skoXC1 ^ 2 = 97/125 ^ 2),
    inference(resolve, [], [refute_0_6, refute_0_9])).

cnf(refute_0_11, plain, (1597 * skoXC1 ^ 2 = 1597 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_12, plain,
    (1597 * skoXC1 ^ 2 != 1597 * skoXC1 ^ 2 | skoXC1 ^ 2 != 97/125 ^ 2 |
     1597 * skoXC1 ^ 2 = 1597 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_13, plain,
    (skoXC1 ^ 2 != 97/125 ^ 2 | 1597 * skoXC1 ^ 2 = 1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_12])).

cnf(refute_0_14, plain, (1597 * skoXC1 ^ 2 = 1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_10, refute_0_13])).

cnf(refute_0_15, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    (1597 * skoXC1 ^ 2 != 1597 * 97/125 ^ 2 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (1597 * skoXC1 ^ 2 != 1597 * 97/125 ^ 2 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 97/125 |
     skoXC1 ^ 4 = 97/125 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_21, plain, (skoXC1 != 97/125 | skoXC1 ^ 4 = 97/125 ^ 4),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain, (skoXC1 ^ 4 = 97/125 ^ 4),
    inference(resolve, [], [refute_0_6, refute_0_21])).

cnf(refute_0_23, plain, (12054 * skoXC1 ^ 4 = 12054 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_24, plain,
    (12054 * skoXC1 ^ 4 != 12054 * skoXC1 ^ 4 | skoXC1 ^ 4 != 97/125 ^ 4 |
     12054 * skoXC1 ^ 4 = 12054 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (skoXC1 ^ 4 != 97/125 ^ 4 | 12054 * skoXC1 ^ 4 = 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain, (12054 * skoXC1 ^ 4 = 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (12054 * skoXC1 ^ 4 != 12054 * 97/125 ^ 4 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_29, plain,
    (12054 * skoXC1 ^ 4 != 12054 * 97/125 ^ 4 |
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 97/125 |
     skoXC1 ^ 6 = 97/125 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_33, plain, (skoXC1 != 97/125 | skoXC1 ^ 6 = 97/125 ^ 6),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain, (skoXC1 ^ 6 = 97/125 ^ 6),
    inference(resolve, [], [refute_0_6, refute_0_33])).

cnf(refute_0_35, plain, (4803 * skoXC1 ^ 6 = 4803 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_36, plain,
    (4803 * skoXC1 ^ 6 != 4803 * skoXC1 ^ 6 | skoXC1 ^ 6 != 97/125 ^ 6 |
     4803 * skoXC1 ^ 6 = 4803 * 97/125 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    (skoXC1 ^ 6 != 97/125 ^ 6 | 4803 * skoXC1 ^ 6 = 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (4803 * skoXC1 ^ 6 = 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_34, refute_0_37])).

cnf(refute_0_39, plain,
    (800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_40, plain,
    (4803 * skoXC1 ^ 6 != 4803 * 97/125 ^ 6 |
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 |
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_41, plain,
    (4803 * skoXC1 ^ 6 != 4803 * 97/125 ^ 6 |
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_39, refute_0_40])).

cnf(refute_0_42, plain,
    (800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_38, refute_0_41])).

cnf(refute_0_43, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 97/125 |
     skoXC1 ^ 8 = 97/125 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_45, plain, (skoXC1 != 97/125 | skoXC1 ^ 8 = 97/125 ^ 8),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain, (skoXC1 ^ 8 = 97/125 ^ 8),
    inference(resolve, [], [refute_0_6, refute_0_45])).

cnf(refute_0_47, plain, (800 * skoXC1 ^ 8 = 800 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (800 * skoXC1 ^ 8 != 800 * skoXC1 ^ 8 | skoXC1 ^ 8 != 97/125 ^ 8 |
     800 * skoXC1 ^ 8 = 800 * 97/125 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (skoXC1 ^ 8 != 97/125 ^ 8 | 800 * skoXC1 ^ 8 = 800 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (800 * skoXC1 ^ 8 = 800 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_46, refute_0_49])).

cnf(refute_0_51, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (800 * skoXC1 ^ 8 != 800 * 97/125 ^ 8 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (800 * skoXC1 ^ 8 != 800 * 97/125 ^ 8 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_56, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_57, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain, (Y0 != X0 | Y0 != Z | X0 = Z),
    introduced(tautology, [equality])).

cnf(refute_0_59, plain, (X0 != Y0 | Y0 != Z | X0 = Z),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (800 * 97/125 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_54, refute_0_60])).

cnf(refute_0_62, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_42, refute_0_61])).

cnf(refute_0_63, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_68, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_30, refute_0_68])).

cnf(refute_0_70, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_71, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_72, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2), inference(subst, [], [refute_0_59])).

cnf(refute_0_75, plain,
    (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_75])).

cnf(refute_0_77, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 =
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55), introduced(tautology, [refl])).

cnf(refute_0_78, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 !=
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 + 55), introduced(tautology, [equality])).

cnf(refute_0_79, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 |
     800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 + 55),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 =
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 + 55),
    inference(resolve, [], [refute_0_76, refute_0_79])).

cnf(refute_0_81, plain,
    ((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55)), introduced(tautology, [refl])).

cnf(refute_0_82, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 + 55 |
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) |
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)), introduced(tautology, [equality])).

cnf(refute_0_83, plain,
    (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
     1597 * skoXC1 ^ 2 + 55 !=
     800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
     1597 * 97/125 ^ 2 + 55 |
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    ((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)),
    inference(resolve, [], [refute_0_80, refute_0_83])).

cnf(refute_0_85, plain, (755 * skoXC1 ^ 2 = 755 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (755 * skoXC1 ^ 2 != 755 * skoXC1 ^ 2 | skoXC1 ^ 2 != 97/125 ^ 2 |
     755 * skoXC1 ^ 2 = 755 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (skoXC1 ^ 2 != 97/125 ^ 2 | 755 * skoXC1 ^ 2 = 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain, (755 * skoXC1 ^ 2 = 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_10, refute_0_87])).

cnf(refute_0_89, plain,
    (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (755 * skoXC1 ^ 2 != 755 * 97/125 ^ 2 |
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 |
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (755 * skoXC1 ^ 2 != 755 * 97/125 ^ 2 |
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain, (14382 * skoXC1 ^ 4 = 14382 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_94, plain,
    (14382 * skoXC1 ^ 4 != 14382 * skoXC1 ^ 4 | skoXC1 ^ 4 != 97/125 ^ 4 |
     14382 * skoXC1 ^ 4 = 14382 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_95, plain,
    (skoXC1 ^ 4 != 97/125 ^ 4 | 14382 * skoXC1 ^ 4 = 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain, (14382 * skoXC1 ^ 4 = 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_22, refute_0_95])).

cnf(refute_0_97, plain,
    (797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    (14382 * skoXC1 ^ 4 != 14382 * 97/125 ^ 4 |
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 |
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    (14382 * skoXC1 ^ 4 != 14382 * 97/125 ^ 4 |
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain, (797 * skoXC1 ^ 6 = 797 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_102, plain,
    (797 * skoXC1 ^ 6 != 797 * skoXC1 ^ 6 | skoXC1 ^ 6 != 97/125 ^ 6 |
     797 * skoXC1 ^ 6 = 797 * 97/125 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_103, plain,
    (skoXC1 ^ 6 != 97/125 ^ 6 | 797 * skoXC1 ^ 6 = 797 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain, (797 * skoXC1 ^ 6 = 797 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_34, refute_0_103])).

cnf(refute_0_105, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_106, plain,
    (797 * skoXC1 ^ 6 != 797 * 97/125 ^ 6 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_107, plain,
    (797 * skoXC1 ^ 6 != 797 * 97/125 ^ 6 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_104, refute_0_107])).

cnf(refute_0_109, plain,
    (797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_110, plain,
    (797 * 97/125 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_100, refute_0_110])).

cnf(refute_0_112, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_113, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_114, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_112, refute_0_113])).

cnf(refute_0_115, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_111, refute_0_114])).

cnf(refute_0_116, plain,
    (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_117, plain,
    (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_92, refute_0_117])).

cnf(refute_0_119, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 =
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 !=
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194),
    introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 |
     797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 =
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194),
    inference(resolve, [], [refute_0_118, refute_0_121])).

cnf(refute_0_123, plain,
    ((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55)), introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194 |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55)), introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194 !=
     797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194 |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55)),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    ((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55)),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    ((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55) |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)), inference(subst, [], [refute_0_59])).

cnf(refute_0_128, plain,
    ((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55) |
     (797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)),
    inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain,
    ((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) =
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55)),
    inference(resolve, [], [refute_0_84, refute_0_128])).

cnf(refute_0_130, plain,
    (arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55))), introduced(tautology, [refl])).

cnf(refute_0_131, plain,
    ((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55) |
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))), introduced(tautology, [equality])).

cnf(refute_0_132, plain,
    ((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 - 3194) /
     (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
      1597 * skoXC1 ^ 2 + 55) !=
     (797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 - 3194) /
     (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
      1597 * 97/125 ^ 2 + 55) |
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_130, refute_0_131])).

cnf(refute_0_133, plain,
    (arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_129, refute_0_132])).

cnf(refute_0_134, plain,
    (8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55))), introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))), introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_133, refute_0_136])).

cnf(refute_0_138, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55))), introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))), introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) =
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55))),
    inference(resolve, [], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) <
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) !=
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_143, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) <
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain, (60 * skoXC1 ^ 2 = 60 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_145, plain,
    (60 * skoXC1 ^ 2 != 60 * skoXC1 ^ 2 | skoXC1 ^ 2 != 97/125 ^ 2 |
     60 * skoXC1 ^ 2 = 60 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_146, plain,
    (skoXC1 ^ 2 != 97/125 ^ 2 | 60 * skoXC1 ^ 2 = 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain, (60 * skoXC1 ^ 2 = 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_10, refute_0_146])).

cnf(refute_0_148, plain,
    (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_149, plain,
    (60 * skoXC1 ^ 2 != 60 * 97/125 ^ 2 |
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 |
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_150, plain,
    (60 * skoXC1 ^ 2 != 60 * 97/125 ^ 2 |
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_147, refute_0_150])).

cnf(refute_0_152, plain, (5 * skoXC1 ^ 4 = 5 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_153, plain,
    (5 * skoXC1 ^ 4 != 5 * skoXC1 ^ 4 | skoXC1 ^ 4 != 97/125 ^ 4 |
     5 * skoXC1 ^ 4 = 5 * 97/125 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_154, plain,
    (skoXC1 ^ 4 != 97/125 ^ 4 | 5 * skoXC1 ^ 4 = 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain, (5 * skoXC1 ^ 4 = 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_22, refute_0_154])).

cnf(refute_0_156, plain,
    (97/125 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_157, plain,
    (5 * skoXC1 ^ 4 != 5 * 97/125 ^ 4 |
     97/125 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * skoXC1 ^ 4 |
     97/125 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_158, plain,
    (5 * skoXC1 ^ 4 != 5 * 97/125 ^ 4 |
     97/125 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (97/125 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_155, refute_0_158])).

cnf(refute_0_160, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = skoXC1 ^ 6 + 5 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_161, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 != skoXC1 ^ 6 + 5 * skoXC1 ^ 4 |
     skoXC1 ^ 6 != 97/125 ^ 6 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_162, plain,
    (skoXC1 ^ 6 != 97/125 ^ 6 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_34, refute_0_162])).

cnf(refute_0_164, plain,
    (97/125 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * 97/125 ^ 4 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * skoXC1 ^ 4 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_165, plain,
    (97/125 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * 97/125 ^ 4 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_163, refute_0_164])).

cnf(refute_0_166, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 = 97/125 ^ 6 + 5 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_159, refute_0_165])).

cnf(refute_0_167, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_168, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 !=
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * 97/125 ^ 4 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_169, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 != 97/125 ^ 6 + 5 * 97/125 ^ 4 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_166, refute_0_169])).

cnf(refute_0_171, plain,
    (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_172, plain,
    (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_151, refute_0_172])).

cnf(refute_0_174, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 =
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16),
    introduced(tautology, [refl])).

cnf(refute_0_175, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 !=
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16),
    introduced(tautology, [equality])).

cnf(refute_0_176, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 |
     skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16),
    inference(resolve, [], [refute_0_174, refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 =
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16),
    inference(resolve, [], [refute_0_173, refute_0_176])).

cnf(refute_0_178, plain,
    ((skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_179, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16 |
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 |
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_180, plain,
    (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16 !=
     97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16 |
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    ((skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_177, refute_0_180])).

cnf(refute_0_182, plain,
    ((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_183, plain,
    ((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 |
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_184, plain,
    ((skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    ((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_181, refute_0_184])).

cnf(refute_0_186, plain, (4998610 * skoXC1 ^ 2 = 4998610 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_187, plain,
    (4998610 * skoXC1 ^ 2 != 4998610 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 97/125 ^ 2 |
     4998610 * skoXC1 ^ 2 = 4998610 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_188, plain,
    (skoXC1 ^ 2 != 97/125 ^ 2 |
     4998610 * skoXC1 ^ 2 = 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain, (4998610 * skoXC1 ^ 2 = 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_10, refute_0_188])).

cnf(refute_0_190, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_191, plain,
    (4998610 * skoXC1 ^ 2 != 4998610 * 97/125 ^ 2 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_192, plain,
    (4998610 * skoXC1 ^ 2 != 4998610 * 97/125 ^ 2 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_190, refute_0_191])).

cnf(refute_0_193, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_189, refute_0_192])).

cnf(refute_0_194, plain, (87425842 * skoXC1 ^ 4 = 87425842 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_195, plain,
    (87425842 * skoXC1 ^ 4 != 87425842 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 97/125 ^ 4 |
     87425842 * skoXC1 ^ 4 = 87425842 * 97/125 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_196, plain,
    (skoXC1 ^ 4 != 97/125 ^ 4 |
     87425842 * skoXC1 ^ 4 = 87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain, (87425842 * skoXC1 ^ 4 = 87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_22, refute_0_196])).

cnf(refute_0_198, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4), introduced(tautology, [refl])).

cnf(refute_0_199, plain,
    (87425842 * skoXC1 ^ 4 != 87425842 * 97/125 ^ 4 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_200, plain,
    (87425842 * skoXC1 ^ 4 != 87425842 * 97/125 ^ 4 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_197, refute_0_200])).

cnf(refute_0_202, plain, (21346562 * skoXC1 ^ 6 = 21346562 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_203, plain,
    (21346562 * skoXC1 ^ 6 != 21346562 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 97/125 ^ 6 |
     21346562 * skoXC1 ^ 6 = 21346562 * 97/125 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_204, plain,
    (skoXC1 ^ 6 != 97/125 ^ 6 |
     21346562 * skoXC1 ^ 6 = 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain, (21346562 * skoXC1 ^ 6 = 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_34, refute_0_204])).

cnf(refute_0_206, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_207, plain,
    (21346562 * skoXC1 ^ 6 != 21346562 * 97/125 ^ 6 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_208, plain,
    (21346562 * skoXC1 ^ 6 != 21346562 * 97/125 ^ 6 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_205, refute_0_208])).

cnf(refute_0_210, plain, (338091528 * skoXC1 ^ 8 = 338091528 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_211, plain,
    (338091528 * skoXC1 ^ 8 != 338091528 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 97/125 ^ 8 |
     338091528 * skoXC1 ^ 8 = 338091528 * 97/125 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_212, plain,
    (skoXC1 ^ 8 != 97/125 ^ 8 |
     338091528 * skoXC1 ^ 8 = 338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain, (338091528 * skoXC1 ^ 8 = 338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_46, refute_0_212])).

cnf(refute_0_214, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8), introduced(tautology, [refl])).

cnf(refute_0_215, plain,
    (338091528 * skoXC1 ^ 8 != 338091528 * 97/125 ^ 8 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_216, plain,
    (338091528 * skoXC1 ^ 8 != 338091528 * 97/125 ^ 8 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_213, refute_0_216])).

cnf(refute_0_218, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_219, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 97/125 |
     skoXC1 ^ 10 = 97/125 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_220, plain, (skoXC1 != 97/125 | skoXC1 ^ 10 = 97/125 ^ 10),
    inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain, (skoXC1 ^ 10 = 97/125 ^ 10),
    inference(resolve, [], [refute_0_6, refute_0_220])).

cnf(refute_0_222, plain,
    (136160432 * skoXC1 ^ 10 = 136160432 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_223, plain,
    (136160432 * skoXC1 ^ 10 != 136160432 * skoXC1 ^ 10 |
     skoXC1 ^ 10 != 97/125 ^ 10 |
     136160432 * skoXC1 ^ 10 = 136160432 * 97/125 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_224, plain,
    (skoXC1 ^ 10 != 97/125 ^ 10 |
     136160432 * skoXC1 ^ 10 = 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    (136160432 * skoXC1 ^ 10 = 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_221, refute_0_224])).

cnf(refute_0_226, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_227, plain,
    (136160432 * skoXC1 ^ 10 != 136160432 * 97/125 ^ 10 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_228, plain,
    (136160432 * skoXC1 ^ 10 != 136160432 * 97/125 ^ 10 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_225, refute_0_228])).

cnf(refute_0_230, plain, (skoXC1 ^ 12 = skoXC1 ^ 12),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (skoXC1 ^ 12 != skoXC1 ^ 12 | skoXC1 != 97/125 |
     skoXC1 ^ 12 = 97/125 ^ 12), introduced(tautology, [equality])).

cnf(refute_0_232, plain, (skoXC1 != 97/125 | skoXC1 ^ 12 = 97/125 ^ 12),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain, (skoXC1 ^ 12 = 97/125 ^ 12),
    inference(resolve, [], [refute_0_6, refute_0_232])).

cnf(refute_0_234, plain, (42990418 * skoXC1 ^ 12 = 42990418 * skoXC1 ^ 12),
    introduced(tautology, [refl])).

cnf(refute_0_235, plain,
    (42990418 * skoXC1 ^ 12 != 42990418 * skoXC1 ^ 12 |
     skoXC1 ^ 12 != 97/125 ^ 12 |
     42990418 * skoXC1 ^ 12 = 42990418 * 97/125 ^ 12),
    introduced(tautology, [equality])).

cnf(refute_0_236, plain,
    (skoXC1 ^ 12 != 97/125 ^ 12 |
     42990418 * skoXC1 ^ 12 = 42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain, (42990418 * skoXC1 ^ 12 = 42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_233, refute_0_236])).

cnf(refute_0_238, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12), introduced(tautology, [refl])).

cnf(refute_0_239, plain,
    (42990418 * skoXC1 ^ 12 != 42990418 * 97/125 ^ 12 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12), introduced(tautology, [equality])).

cnf(refute_0_240, plain,
    (42990418 * skoXC1 ^ 12 != 42990418 * 97/125 ^ 12 |
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_238, refute_0_239])).

cnf(refute_0_241, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_237, refute_0_240])).

cnf(refute_0_242, plain, (skoXC1 ^ 14 = skoXC1 ^ 14),
    introduced(tautology, [refl])).

cnf(refute_0_243, plain,
    (skoXC1 ^ 14 != skoXC1 ^ 14 | skoXC1 != 97/125 |
     skoXC1 ^ 14 = 97/125 ^ 14), introduced(tautology, [equality])).

cnf(refute_0_244, plain, (skoXC1 != 97/125 | skoXC1 ^ 14 = 97/125 ^ 14),
    inference(resolve, [], [refute_0_242, refute_0_243])).

cnf(refute_0_245, plain, (skoXC1 ^ 14 = 97/125 ^ 14),
    inference(resolve, [], [refute_0_6, refute_0_244])).

cnf(refute_0_246, plain, (7684800 * skoXC1 ^ 14 = 7684800 * skoXC1 ^ 14),
    introduced(tautology, [refl])).

cnf(refute_0_247, plain,
    (7684800 * skoXC1 ^ 14 != 7684800 * skoXC1 ^ 14 |
     skoXC1 ^ 14 != 97/125 ^ 14 |
     7684800 * skoXC1 ^ 14 = 7684800 * 97/125 ^ 14),
    introduced(tautology, [equality])).

cnf(refute_0_248, plain,
    (skoXC1 ^ 14 != 97/125 ^ 14 |
     7684800 * skoXC1 ^ 14 = 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_246, refute_0_247])).

cnf(refute_0_249, plain, (7684800 * skoXC1 ^ 14 = 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_245, refute_0_248])).

cnf(refute_0_250, plain,
    (640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14),
    introduced(tautology, [refl])).

cnf(refute_0_251, plain,
    (7684800 * skoXC1 ^ 14 != 7684800 * 97/125 ^ 14 |
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 |
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    introduced(tautology, [equality])).

cnf(refute_0_252, plain,
    (7684800 * skoXC1 ^ 14 != 7684800 * 97/125 ^ 14 |
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_250, refute_0_251])).

cnf(refute_0_253, plain,
    (640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_249, refute_0_252])).

cnf(refute_0_254, plain, (skoXC1 ^ 16 = skoXC1 ^ 16),
    introduced(tautology, [refl])).

cnf(refute_0_255, plain,
    (skoXC1 ^ 16 != skoXC1 ^ 16 | skoXC1 != 97/125 |
     skoXC1 ^ 16 = 97/125 ^ 16), introduced(tautology, [equality])).

cnf(refute_0_256, plain, (skoXC1 != 97/125 | skoXC1 ^ 16 = 97/125 ^ 16),
    inference(resolve, [], [refute_0_254, refute_0_255])).

cnf(refute_0_257, plain, (skoXC1 ^ 16 = 97/125 ^ 16),
    inference(resolve, [], [refute_0_6, refute_0_256])).

cnf(refute_0_258, plain, (640000 * skoXC1 ^ 16 = 640000 * skoXC1 ^ 16),
    introduced(tautology, [refl])).

cnf(refute_0_259, plain,
    (640000 * skoXC1 ^ 16 != 640000 * skoXC1 ^ 16 |
     skoXC1 ^ 16 != 97/125 ^ 16 |
     640000 * skoXC1 ^ 16 = 640000 * 97/125 ^ 16),
    introduced(tautology, [equality])).

cnf(refute_0_260, plain,
    (skoXC1 ^ 16 != 97/125 ^ 16 |
     640000 * skoXC1 ^ 16 = 640000 * 97/125 ^ 16),
    inference(resolve, [], [refute_0_258, refute_0_259])).

cnf(refute_0_261, plain, (640000 * skoXC1 ^ 16 = 640000 * 97/125 ^ 16),
    inference(resolve, [], [refute_0_257, refute_0_260])).

cnf(refute_0_262, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14),
    introduced(tautology, [refl])).

cnf(refute_0_263, plain,
    (640000 * skoXC1 ^ 16 != 640000 * 97/125 ^ 16 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14),
    introduced(tautology, [equality])).

cnf(refute_0_264, plain,
    (640000 * skoXC1 ^ 16 != 640000 * 97/125 ^ 16 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14),
    inference(resolve, [], [refute_0_261, refute_0_264])).

cnf(refute_0_266, plain,
    (640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_267, plain,
    (640000 * 97/125 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_265, refute_0_266])).

cnf(refute_0_268, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14),
    inference(resolve, [], [refute_0_253, refute_0_267])).

cnf(refute_0_269, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12), introduced(tautology, [refl])).

cnf(refute_0_270, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12), introduced(tautology, [equality])).

cnf(refute_0_271, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12),
    inference(resolve, [], [refute_0_269, refute_0_270])).

cnf(refute_0_272, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12),
    inference(resolve, [], [refute_0_268, refute_0_271])).

cnf(refute_0_273, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12), inference(subst, [], [refute_0_59])).

cnf(refute_0_274, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_272, refute_0_273])).

cnf(refute_0_275, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12),
    inference(resolve, [], [refute_0_241, refute_0_274])).

cnf(refute_0_276, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_277, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_278, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10),
    inference(resolve, [], [refute_0_276, refute_0_277])).

cnf(refute_0_279, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10),
    inference(resolve, [], [refute_0_275, refute_0_278])).

cnf(refute_0_280, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_281, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_279, refute_0_280])).

cnf(refute_0_282, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10),
    inference(resolve, [], [refute_0_229, refute_0_281])).

cnf(refute_0_283, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8), introduced(tautology, [refl])).

cnf(refute_0_284, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_285, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_283, refute_0_284])).

cnf(refute_0_286, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_282, refute_0_285])).

cnf(refute_0_287, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8), inference(subst, [], [refute_0_59])).

cnf(refute_0_288, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_286, refute_0_287])).

cnf(refute_0_289, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8),
    inference(resolve, [], [refute_0_217, refute_0_288])).

cnf(refute_0_290, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_291, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_292, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_290, refute_0_291])).

cnf(refute_0_293, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_289, refute_0_292])).

cnf(refute_0_294, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_295, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_293, refute_0_294])).

cnf(refute_0_296, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6),
    inference(resolve, [], [refute_0_209, refute_0_295])).

cnf(refute_0_297, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4), introduced(tautology, [refl])).

cnf(refute_0_298, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_299, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_297, refute_0_298])).

cnf(refute_0_300, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_296, refute_0_299])).

cnf(refute_0_301, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4), inference(subst, [], [refute_0_59])).

cnf(refute_0_302, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_300, refute_0_301])).

cnf(refute_0_303, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4),
    inference(resolve, [], [refute_0_201, refute_0_302])).

cnf(refute_0_304, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_305, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_306, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_304, refute_0_305])).

cnf(refute_0_307, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_303, refute_0_306])).

cnf(refute_0_308, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_309, plain,
    (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_307, refute_0_308])).

cnf(refute_0_310, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2),
    inference(resolve, [], [refute_0_193, refute_0_309])).

cnf(refute_0_311, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 =
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661),
    introduced(tautology, [refl])).

cnf(refute_0_312, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 !=
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661),
    introduced(tautology, [equality])).

cnf(refute_0_313, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 |
     640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661),
    inference(resolve, [], [refute_0_311, refute_0_312])).

cnf(refute_0_314, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 =
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661),
    inference(resolve, [], [refute_0_310, refute_0_313])).

cnf(refute_0_315, plain,
    ((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_316, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_317, plain,
    (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
     42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
     338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
     87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661 !=
     640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
     42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
     338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
     87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_315, refute_0_316])).

cnf(refute_0_318, plain,
    ((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_314, refute_0_317])).

cnf(refute_0_319, plain,
    ((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_320, plain,
    ((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     (640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_318, refute_0_319])).

cnf(refute_0_321, plain,
    ((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 =
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2),
    inference(resolve, [], [refute_0_185, refute_0_320])).

cnf(refute_0_322, plain,
    (ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_323, plain,
    ((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) |
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_324, plain,
    ((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
      42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
      338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
      87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
     (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2 !=
     (640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
      42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
      338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
      87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
     (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2 |
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_322, refute_0_323])).

cnf(refute_0_325, plain,
    (ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_321, refute_0_324])).

cnf(refute_0_326, plain,
    (868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_327, plain,
    (868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) |
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_328, plain,
    (ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_326, refute_0_327])).

cnf(refute_0_329, plain,
    (868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_325, refute_0_328])).

cnf(refute_0_330, plain,
    (log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    inference(subst, [], [log_def])).

cnf(refute_0_331, plain,
    (868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) |
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_332, plain,
    (868589/2000000 *
     ln((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
         42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
         338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
         87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
        (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_330, refute_0_331])).

cnf(refute_0_333, plain,
    (log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)),
    inference(resolve, [], [refute_0_329, refute_0_332])).

cnf(refute_0_334, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_335, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) |
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_336, plain,
    (log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     868589/2000000 *
     ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
         42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
         338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
         87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
        (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2) |
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    inference(resolve, [], [refute_0_334, refute_0_335])).

cnf(refute_0_337, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) =
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    inference(resolve, [], [refute_0_333, refute_0_336])).

cnf(refute_0_338, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) <
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) !=
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2)) |
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_339, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) <
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    inference(resolve, [], [refute_0_337, refute_0_338])).

cnf(refute_0_340, plain,
    (10 *
     log((640000 * skoXC1 ^ 16 + 7684800 * skoXC1 ^ 14 +
          42990418 * skoXC1 ^ 12 + 136160432 * skoXC1 ^ 10 +
          338091528 * skoXC1 ^ 8 - 21346562 * skoXC1 ^ 6 -
          87425842 * skoXC1 ^ 4 - 4998610 * skoXC1 ^ 2 + 10204661) /
         (skoXC1 ^ 6 + 5 * skoXC1 ^ 4 + 60 * skoXC1 ^ 2 + 16) ^ 2) <
     55 +
     8 *
     arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 + 755 * skoXC1 ^ 2 -
             3194) /
            (800 * skoXC1 ^ 8 + 4803 * skoXC1 ^ 6 + 12054 * skoXC1 ^ 4 -
             1597 * skoXC1 ^ 2 + 55)) |
     55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    inference(resolve, [], [refute_0_143, refute_0_339])).

cnf(refute_0_341, plain,
    (55 +
     8 *
     arctan((797 * 97/125 ^ 6 + 14382 * 97/125 ^ 4 + 755 * 97/125 ^ 2 -
             3194) /
            (800 * 97/125 ^ 8 + 4803 * 97/125 ^ 6 + 12054 * 97/125 ^ 4 -
             1597 * 97/125 ^ 2 + 55)) <=
     10 *
     (868589/2000000 *
      ln((640000 * 97/125 ^ 16 + 7684800 * 97/125 ^ 14 +
          42990418 * 97/125 ^ 12 + 136160432 * 97/125 ^ 10 +
          338091528 * 97/125 ^ 8 - 21346562 * 97/125 ^ 6 -
          87425842 * 97/125 ^ 4 - 4998610 * 97/125 ^ 2 + 10204661) /
         (97/125 ^ 6 + 5 * 97/125 ^ 4 + 60 * 97/125 ^ 2 + 16) ^ 2))),
    inference(resolve, [], [refute_0_5, refute_0_340])).

cnf(refute_0_342, plain,
    (11000000/868589 +
     arctan(6317656210882992500/11010804368800679627) * 1600000/868589 <=
     ln(161150592846908583096118035092149109129/
        16675142311240119654438073014062500)),
    inference(arithmetic, [], [refute_0_341])).

cnf(refute_0_343, plain,
    (ln(X_000090) < X_000091 | ~ lgen(1, ln(X_000090), X_000091)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_344, plain,
    (~ lgen(1,
            1/40000 * (X_000090 + 50000) * (X_000090 - 10000) /
            (X_000090 + 5000) + 461/50, X_000091) | X_000090 <= 0 |
     lgen(1, ln(X_000090), X_000091)),
    inference(subst, [], [log_upper_bound_scaled_cf3])).

cnf(refute_0_345, plain,
    (ln(X_000090) < X_000091 |
     ~ lgen(1,
            1/40000 * (X_000090 + 50000) * (X_000090 - 10000) /
            (X_000090 + 5000) + 461/50, X_000091) | X_000090 <= 0),
    inference(resolve, [], [refute_0_344, refute_0_343])).

cnf(refute_0_346, plain,
    (ln(X_000090) < X_000091 | X_000090 <= 0 |
     X_000091 <=
     461/50 +
     (-12500 + X_000090 * (1 + X_000090 * 1/40000)) / (5000 + X_000090)),
    inference(arithmetic, [], [refute_0_345])).

cnf(refute_0_347, plain,
    (ln(161150592846908583096118035092149109129/
        16675142311240119654438073014062500) <
     11000000/868589 +
     arctan(6317656210882992500/11010804368800679627) * 1600000/868589 |
     161150592846908583096118035092149109129/
     16675142311240119654438073014062500 <= 0 |
     11000000/868589 +
     arctan(6317656210882992500/11010804368800679627) * 1600000/868589 <=
     461/50 +
     (-12500 +
      161150592846908583096118035092149109129/
      16675142311240119654438073014062500 *
      (1 +
       161150592846908583096118035092149109129/
       16675142311240119654438073014062500 * 1/40000)) /
     (5000 +
      161150592846908583096118035092149109129/
      16675142311240119654438073014062500)),
    inference(subst, [], [refute_0_346])).

cnf(refute_0_348, plain,
    (161150592846908583096118035092149109129/
     16675142311240119654438073014062500 <= 0 |
     11000000/868589 +
     arctan(6317656210882992500/11010804368800679627) * 1600000/868589 <=
     461/50 +
     (-12500 +
      161150592846908583096118035092149109129/
      16675142311240119654438073014062500 *
      (1 +
       161150592846908583096118035092149109129/
       16675142311240119654438073014062500 * 1/40000)) /
     (5000 +
      161150592846908583096118035092149109129/
      16675142311240119654438073014062500)),
    inference(resolve, [], [refute_0_342, refute_0_347])).

cnf(refute_0_349, plain,
    (arctan(6317656210882992500/11010804368800679627) <=
     -492772850186751082174258706259547880510056903518817731428359508983077312229088952451/
     260960699184861893241698938008806317823405914321122921161381426904100000000000000000),
    inference(arithmetic, [], [refute_0_348])).

cnf(refute_0_350, plain,
    (6317656210882992500/11010804368800679627 < 0 |
     6317656210882992500/11010804368800679627 * 3 /
     (3 +
      6317656210882992500/11010804368800679627 *
      6317656210882992500/11010804368800679627) <=
     -492772850186751082174258706259547880510056903518817731428359508983077312229088952451/
     260960699184861893241698938008806317823405914321122921161381426904100000000000000000),
    inference(resolve, [], [refute_0_349, refute_0_4])).

cnf(refute_0_351, plain, ($false),
    inference(arithmetic, [], [refute_0_350])).
% SZS output end CNFRefutation for Nichols-Plot-Simple-Example-4.tptp
