% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-2-5.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(atan_upper_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, -157079/100000 - 1 / X, Y) |
     lgen(R, arctan(X), Y))).

cnf(ln_lower_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, Y, (X - 1) / X) | lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Fails_2_5', conjecture,
    (! [X] :
       (X = 57/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 *
                    sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                         21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                         15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900)))))).

cnf(sqrt_upper_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R,
            1/368000 * (X ^ 2 + 50784000000 * X + 71639296000000000000) /
            (X + 8464000000), Y) | lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            368000 * X * (X + 8464000000) /
            (X ^ 2 + 50784000000 * X + 71639296000000000000)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 57/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 *
                    sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                         21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                         15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900))))),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Fails_2_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 57/128 =>
          interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 *
                      sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                           21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                           15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900))))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (~ interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 *
                      sqrt(15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                           21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                           15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900))) & X = 57/128)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900))) &
     skoXC1 = 57/128), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 57/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (~ lgen(0, arctan(X_000061), X_000062) | arctan(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     lgen(0, arctan(X_000061), X_000062)),
    inference(subst, [], [atan_upper_bound_case_11])).

cnf(refute_0_2, plain,
    (~ lgen(0, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     arctan(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000062 < -157079/100000 + 1 / X_000061 * -1 | -1 <= X_000061 |
     arctan(X_000061) <= X_000062),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (pi *
     96057565852967449038779211488856321334913015619562767230897531760799/
     117647976619860341692992894213599115464964858459764158056849733582848 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     arctan(-1139208930514575/1108767509542208) <=
     pi *
     96057565852967449038779211488856321334913015619562767230897531760799/
     117647976619860341692992894213599115464964858459764158056849733582848),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 <
     231993/101125 *
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     0 <= pi * -6 + arctan(-1139208930514575/1108767509542208) * 12 |
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) <=
     231993/101125), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000223, ln(X_000222)) | X_000223 <= ln(X_000222)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (~ lgen(0, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0 |
     lgen(0, X_000223, ln(X_000222))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0 | X_000223 <= ln(X_000222)),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    ((-1/5 +
      X_000222 *
      (-95/6 +
       X_000222 *
       (-140/3 +
        X_000222 * (20 + X_000222 * (115/3 + X_000222 * 131/30))))) /
     (X_000222 *
      (5 +
       X_000222 * (40 + X_000222 * (60 + X_000222 * (20 + X_000222))))) <
     X_000223 | X_000222 <= 0 | X_000223 <= ln(X_000222)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     10 <= 0 |
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) <=
     ln(10)), inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (pi * 20 +
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     50506321767834152681255141268608666866792908661439445722000321/
     545819465729604363683111646639431975876487640999002112000 *
     (pi * 373833953443840000/3064655361465834062209) |
     50506321767834152681255141268608666866792908661439445722000321/
     545819465729604363683111646639431975876487640999002112000 <=
     (pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) |
     pi * 373833953443840000/3064655361465834062209 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain,
    (~ lgen(0, sqrt(X_000103), X_000104) | sqrt(X_000103) <= X_000104),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (X_000103 < 0 |
     ~ lgen(0,
            1/368000 *
            (X_000103 ^ 2 + 50784000000 * X_000103 +
             71639296000000000000) / (X_000103 + 8464000000), X_000104) |
     lgen(0, sqrt(X_000103), X_000104)),
    inference(subst, [], [sqrt_upper_bound_1_scaled])).

cnf(refute_0_14, plain,
    (X_000103 < 0 |
     ~ lgen(0,
            1/368000 *
            (X_000103 ^ 2 + 50784000000 * X_000103 +
             71639296000000000000) / (X_000103 + 8464000000), X_000104) |
     sqrt(X_000103) <= X_000104),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (X_000103 < 0 |
     X_000104 <
     (194672000000000 + X_000103 * (138000 + X_000103 * 1/368000)) /
     (8464000000 + X_000103) | sqrt(X_000103) <= X_000104),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (2527162377580592078908020945889/295147905179352825856 < 0 |
     (pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     (194672000000000 +
      2527162377580592078908020945889/295147905179352825856 *
      (138000 +
       2527162377580592078908020945889/295147905179352825856 * 1/368000)) /
     (8464000000 + 2527162377580592078908020945889/295147905179352825856) |
     sqrt(2527162377580592078908020945889/295147905179352825856) <=
     (pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209)),
    inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) * (pi * 20) |
     pi * 20 <= 0 |
     -1 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <=
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_18, plain,
    (~ lgen(0, ln(X_000047), X_000048) | ln(X_000047) <= X_000048),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(0, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_20, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     ln(X_000047) <= X_000048),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (X_000048 < -1 + X_000047 | X_000047 <= 0 | ln(X_000047) <= X_000048),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     -1 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0 |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) <=
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20)), inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) * ln(10) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20) |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20) / ln(10) <=
     pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_24, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) * pi |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_25, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_26, plain, (skoXC1 = 57/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 57/128 |
     skoXC1 ^ 2 = 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_29, plain, (skoXC1 != 57/128 | skoXC1 ^ 2 = 57/128 ^ 2),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain, (skoXC1 ^ 2 = 57/128 ^ 2),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain, (406627 * skoXC1 ^ 2 = 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (406627 * skoXC1 ^ 2 != 406627 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 | 406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 | 406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain, (406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_33])).

cnf(refute_0_35, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_36, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 57/128 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 57/128 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_34, refute_0_37])).

cnf(refute_0_39, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_40, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 57/128 |
     skoXC1 ^ 4 = 57/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_41, plain, (skoXC1 != 57/128 | skoXC1 ^ 4 = 57/128 ^ 4),
    inference(resolve, [], [refute_0_39, refute_0_40])).

cnf(refute_0_42, plain, (skoXC1 ^ 4 = 57/128 ^ 4),
    inference(resolve, [], [refute_0_26, refute_0_41])).

cnf(refute_0_43, plain, (14325 * skoXC1 ^ 4 = 14325 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (14325 * skoXC1 ^ 4 != 14325 * skoXC1 ^ 4 | skoXC1 ^ 4 != 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 | 14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain, (14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_46, refute_0_49])).

cnf(refute_0_51, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_52, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_55, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_50, refute_0_56])).

cnf(refute_0_58, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_38, refute_0_57])).

cnf(refute_0_59, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660),
    introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_58, refute_0_61])).

cnf(refute_0_63, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660 |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660 |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain, (169932 * skoXC1 = 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_68, plain,
    (169932 * skoXC1 != 169932 * skoXC1 | skoXC1 != 57/128 |
     169932 * skoXC1 = 169932 * 57/128),
    introduced(tautology, [equality])).

cnf(refute_0_69, plain,
    (skoXC1 != 57/128 | 169932 * skoXC1 = 169932 * 57/128),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain, (169932 * skoXC1 = 169932 * 57/128),
    inference(resolve, [], [refute_0_26, refute_0_69])).

cnf(refute_0_71, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_72, plain,
    (169932 * skoXC1 != 169932 * 57/128 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    introduced(tautology, [equality])).

cnf(refute_0_73, plain,
    (169932 * skoXC1 != 169932 * 57/128 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_71, refute_0_72])).

cnf(refute_0_74, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_70, refute_0_73])).

cnf(refute_0_75, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_76, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 57/128 |
     skoXC1 ^ 5 = 57/128 ^ 5), introduced(tautology, [equality])).

cnf(refute_0_77, plain, (skoXC1 != 57/128 | skoXC1 ^ 5 = 57/128 ^ 5),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain, (skoXC1 ^ 5 = 57/128 ^ 5),
    inference(resolve, [], [refute_0_26, refute_0_77])).

cnf(refute_0_79, plain, (3890 * skoXC1 ^ 5 = 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_80, plain,
    (3890 * skoXC1 ^ 5 != 3890 * skoXC1 ^ 5 | skoXC1 ^ 5 != 57/128 ^ 5 |
     3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_81, plain,
    (skoXC1 ^ 5 != 57/128 ^ 5 | 3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain, (3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_78, refute_0_81])).

cnf(refute_0_83, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_84, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 57/128 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_85, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 57/128 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_82, refute_0_85])).

cnf(refute_0_87, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_88, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 57/128 |
     skoXC1 ^ 3 = 57/128 ^ 3), introduced(tautology, [equality])).

cnf(refute_0_89, plain, (skoXC1 != 57/128 | skoXC1 ^ 3 = 57/128 ^ 3),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain, (skoXC1 ^ 3 = 57/128 ^ 3),
    inference(resolve, [], [refute_0_26, refute_0_89])).

cnf(refute_0_91, plain, (105247 * skoXC1 ^ 3 = 105247 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_92, plain,
    (105247 * skoXC1 ^ 3 != 105247 * skoXC1 ^ 3 |
     skoXC1 ^ 3 != 57/128 ^ 3 | 105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (skoXC1 ^ 3 != 57/128 ^ 3 | 105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain, (105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    inference(resolve, [], [refute_0_90, refute_0_93])).

cnf(refute_0_95, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_96, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 57/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_97, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 57/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_94, refute_0_97])).

cnf(refute_0_99, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_100, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_86, refute_0_100])).

cnf(refute_0_102, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_103, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_104, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_101, refute_0_104])).

cnf(refute_0_106, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_107, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_74, refute_0_107])).

cnf(refute_0_109, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_110, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_111, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_108, refute_0_111])).

cnf(refute_0_113, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_114, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_112, refute_0_113])).

cnf(refute_0_115, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_66, refute_0_114])).

cnf(refute_0_116, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_117, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_118, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_116, refute_0_117])).

cnf(refute_0_119, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_115, refute_0_118])).

cnf(refute_0_120, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_121, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_122, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_119, refute_0_122])).

cnf(refute_0_124, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_125, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_126, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_123, refute_0_126])).

cnf(refute_0_128, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_129, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_130, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_127, refute_0_130])).

cnf(refute_0_132, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_123, refute_0_136])).

cnf(refute_0_138, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18),
    introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    introduced(tautology, [equality])).

cnf(refute_0_143, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                        21404723599 * skoXC1 ^ 6 +
                        129098541721 * skoXC1 ^ 4 +
                        15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_145, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_146, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_146])).

cnf(refute_0_148, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_149, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 57/128 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_150, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 57/128 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_147, refute_0_150])).

cnf(refute_0_152, plain, (217230815 * skoXC1 ^ 4 = 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_153, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 57/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_154, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain, (217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_154])).

cnf(refute_0_156, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_157, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 57/128 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_158, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 57/128 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_155, refute_0_158])).

cnf(refute_0_160, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_161, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 57/128 |
     skoXC1 ^ 6 = 57/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_162, plain, (skoXC1 != 57/128 | skoXC1 ^ 6 = 57/128 ^ 6),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain, (skoXC1 ^ 6 = 57/128 ^ 6),
    inference(resolve, [], [refute_0_26, refute_0_162])).

cnf(refute_0_164, plain, (3783025 * skoXC1 ^ 6 = 3783025 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_165, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_166, plain,
    (skoXC1 ^ 6 != 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_164, refute_0_165])).

cnf(refute_0_167, plain, (3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_163, refute_0_166])).

cnf(refute_0_168, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_169, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_170, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_167, refute_0_170])).

cnf(refute_0_172, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_173, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_171, refute_0_172])).

cnf(refute_0_174, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_159, refute_0_173])).

cnf(refute_0_175, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_176, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_177, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_174, refute_0_177])).

cnf(refute_0_179, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2), inference(subst, [], [refute_0_55])).

cnf(refute_0_180, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_178, refute_0_179])).

cnf(refute_0_181, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_151, refute_0_180])).

cnf(refute_0_182, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900), introduced(tautology, [refl])).

cnf(refute_0_183, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    introduced(tautology, [equality])).

cnf(refute_0_184, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_181, refute_0_184])).

cnf(refute_0_186, plain,
    (4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)), introduced(tautology, [refl])).

cnf(refute_0_187, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900 |
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    introduced(tautology, [equality])).

cnf(refute_0_188, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900 |
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_185, refute_0_188])).

cnf(refute_0_190, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_191, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_192, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_190, refute_0_191])).

cnf(refute_0_193, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_30, refute_0_192])).

cnf(refute_0_194, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_195, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 57/128 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_196, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 57/128 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_193, refute_0_196])).

cnf(refute_0_198, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_199, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 57/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_200, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_200])).

cnf(refute_0_202, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_203, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 57/128 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_204, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 57/128 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_201, refute_0_204])).

cnf(refute_0_206, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_207, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 57/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_208, plain,
    (skoXC1 ^ 6 != 57/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_163, refute_0_208])).

cnf(refute_0_210, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_211, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 57/128 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_212, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 57/128 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_209, refute_0_212])).

cnf(refute_0_214, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_215, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 57/128 |
     skoXC1 ^ 8 = 57/128 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_216, plain, (skoXC1 != 57/128 | skoXC1 ^ 8 = 57/128 ^ 8),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain, (skoXC1 ^ 8 = 57/128 ^ 8),
    inference(resolve, [], [refute_0_26, refute_0_216])).

cnf(refute_0_218, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_219, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 57/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_220, plain,
    (skoXC1 ^ 8 != 57/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_217, refute_0_220])).

cnf(refute_0_222, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_223, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 57/128 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_224, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 57/128 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_221, refute_0_224])).

cnf(refute_0_226, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_227, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 57/128 |
     skoXC1 ^ 10 = 57/128 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_228, plain, (skoXC1 != 57/128 | skoXC1 ^ 10 = 57/128 ^ 10),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain, (skoXC1 ^ 10 = 57/128 ^ 10),
    inference(resolve, [], [refute_0_26, refute_0_228])).

cnf(refute_0_230, plain, (15132100 * skoXC1 ^ 10 = 15132100 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * skoXC1 ^ 10 |
     skoXC1 ^ 10 != 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_232, plain,
    (skoXC1 ^ 10 != 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain, (15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    inference(resolve, [], [refute_0_229, refute_0_232])).

cnf(refute_0_234, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_235, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_236, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_234, refute_0_235])).

cnf(refute_0_237, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_233, refute_0_236])).

cnf(refute_0_238, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_239, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_225, refute_0_239])).

cnf(refute_0_241, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_242, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_243, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_241, refute_0_242])).

cnf(refute_0_244, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_240, refute_0_243])).

cnf(refute_0_245, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6), inference(subst, [], [refute_0_55])).

cnf(refute_0_246, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_213, refute_0_246])).

cnf(refute_0_248, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_249, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_250, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_247, refute_0_250])).

cnf(refute_0_252, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_253, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_251, refute_0_252])).

cnf(refute_0_254, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_205, refute_0_253])).

cnf(refute_0_255, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_256, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_257, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_255, refute_0_256])).

cnf(refute_0_258, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_254, refute_0_257])).

cnf(refute_0_259, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2), inference(subst, [], [refute_0_55])).

cnf(refute_0_260, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_258, refute_0_259])).

cnf(refute_0_261, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_197, refute_0_260])).

cnf(refute_0_262, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600), introduced(tautology, [refl])).

cnf(refute_0_263, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    introduced(tautology, [equality])).

cnf(refute_0_264, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_261, refute_0_264])).

cnf(refute_0_266, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_267, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600 |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_268, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600 |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_266, refute_0_267])).

cnf(refute_0_269, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_265, refute_0_268])).

cnf(refute_0_270, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_271, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) |
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_272, plain,
    (sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_270, refute_0_271])).

cnf(refute_0_273, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_269, refute_0_272])).

cnf(refute_0_274, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)), introduced(tautology, [refl])).

cnf(refute_0_275, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    introduced(tautology, [equality])).

cnf(refute_0_276, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_274, refute_0_275])).

cnf(refute_0_277, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_273, refute_0_276])).

cnf(refute_0_278, plain,
    (4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_279, plain,
    (4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) |
     4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_277, refute_0_278])).

cnf(refute_0_280, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_189, refute_0_279])).

cnf(refute_0_281, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_282, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_283, plain,
    (4250 *
     sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
          21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
          15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 *
     sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
          21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
          15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_281, refute_0_282])).

cnf(refute_0_284, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_280, refute_0_283])).

cnf(refute_0_285, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_286, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_287, plain,
    (ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_285, refute_0_286])).

cnf(refute_0_288, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_284, refute_0_287])).

cnf(refute_0_289, plain,
    (20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
             21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
             15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 *
        sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
             21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
             15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900)) |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                        21404723599 * 57/128 ^ 6 +
                        129098541721 * 57/128 ^ 4 +
                        15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    introduced(tautology, [equality])).

cnf(refute_0_290, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                        21404723599 * 57/128 ^ 6 +
                        129098541721 * 57/128 ^ 4 +
                        15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_288, refute_0_289])).

cnf(refute_0_291, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                        21404723599 * 57/128 ^ 6 +
                        129098541721 * 57/128 ^ 4 +
                        15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_290, refute_0_143])).

cnf(refute_0_292, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                        21404723599 * 57/128 ^ 6 +
                        129098541721 * 57/128 ^ 4 +
                        15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900))) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 *
                 sqrt(15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                      21404723599 * skoXC1 ^ 6 +
                      129098541721 * skoXC1 ^ 4 +
                      15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_291, refute_0_133])).

cnf(refute_0_293, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 *
                   sqrt(15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                        21404723599 * 57/128 ^ 6 +
                        129098541721 * 57/128 ^ 4 +
                        15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900)))),
    inference(resolve, [], [refute_0_292, refute_0_25])).

cnf(refute_0_294, plain,
    (~ interval(0,
                (pi * -6 +
                 arctan(-1139208930514575/1108767509542208) * 12) / pi, 0,
                (pi * 6 +
                 arctan(-1139208930514575/1108767509542208) * -12) / pi,
                ln(sqrt(2527162377580592078908020945889/
                        295147905179352825856) *
                   18691697672192000/3064655361465834062209) * 20 /
                ln(10)) | pi = 0),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (~ lgen(0,
            ln(sqrt(2527162377580592078908020945889/
                    295147905179352825856) *
               18691697672192000/3064655361465834062209) * 20 / ln(10),
            (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
            pi,
            ln(sqrt(2527162377580592078908020945889/
                    295147905179352825856) *
               18691697672192000/3064655361465834062209) * 20 / ln(10)) |
     interval(0,
              (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
              pi, 0,
              (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
              pi,
              ln(sqrt(2527162377580592078908020945889/
                      295147905179352825856) *
                 18691697672192000/3064655361465834062209) * 20 / ln(10))),
    inference(subst, [], [interval_intro])).

cnf(refute_0_296, plain,
    (~ lgen(0,
            ln(sqrt(2527162377580592078908020945889/
                    295147905179352825856) *
               18691697672192000/3064655361465834062209) * 20 / ln(10),
            (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
            pi,
            ln(sqrt(2527162377580592078908020945889/
                    295147905179352825856) *
               18691697672192000/3064655361465834062209) * 20 / ln(10)) |
     pi = 0), inference(resolve, [], [refute_0_295, refute_0_294])).

cnf(refute_0_297, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) <
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi |
     pi = 0), inference(arithmetic, [], [refute_0_296])).

cnf(refute_0_298, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_299, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_300, plain, (pi != 0),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_301, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) <
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi),
    inference(resolve, [], [refute_0_297, refute_0_300])).

cnf(refute_0_302, plain,
    (arctan(X_000061) < X_000062 | ~ lgen(1, arctan(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_303, plain,
    (~ lgen(1, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     lgen(1, arctan(X_000061), X_000062)),
    inference(subst, [], [atan_upper_bound_case_11])).

cnf(refute_0_304, plain,
    (arctan(X_000061) < X_000062 |
     ~ lgen(1, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061),
    inference(resolve, [], [refute_0_303, refute_0_302])).

cnf(refute_0_305, plain,
    (arctan(X_000061) < X_000062 |
     X_000062 <= -157079/100000 + 1 / X_000061 * -1 | -1 <= X_000061),
    inference(arithmetic, [], [refute_0_304])).

cnf(refute_0_306, plain,
    (arctan(-1139208930514575/1108767509542208) < pi * 1707229/1391958 |
     -1 <= -1139208930514575/1108767509542208 |
     pi * 1707229/1391958 <=
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1),
    inference(subst, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) < 231993/101125 |
     231993/101125 *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209 |
     pi * 112150186033152000/3064655361465834062209 +
     arctan(-1139208930514575/1108767509542208) *
     -224300372066304000/3064655361465834062209 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_308, plain,
    (X_000223 < ln(X_000222) | ~ lgen(1, X_000223, ln(X_000222))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_309, plain,
    (~ lgen(1, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0 |
     lgen(1, X_000223, ln(X_000222))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_310, plain,
    (X_000223 < ln(X_000222) |
     ~ lgen(1, X_000223,
            1/30 *
            (6 + 481 * X_000222 + 1881 * X_000222 ^ 2 +
             1281 * X_000222 ^ 3 + 131 * X_000222 ^ 4) * (X_000222 - 1) /
            (X_000222 *
             (5 + 40 * X_000222 + 60 * X_000222 ^ 2 + 20 * X_000222 ^ 3 +
              X_000222 ^ 4))) | X_000222 <= 0),
    inference(resolve, [], [refute_0_309, refute_0_308])).

cnf(refute_0_311, plain,
    (X_000223 < ln(X_000222) | X_000222 <= 0 |
     (-1/5 +
      X_000222 *
      (-95/6 +
       X_000222 *
       (-140/3 +
        X_000222 * (20 + X_000222 * (115/3 + X_000222 * 131/30))))) /
     (X_000222 *
      (5 +
       X_000222 * (40 + X_000222 * (60 + X_000222 * (20 + X_000222))))) <=
     X_000223), inference(arithmetic, [], [refute_0_310])).

cnf(refute_0_312, plain,
    (pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) < ln(10) | 10 <= 0 |
     (-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(subst, [], [refute_0_311])).

cnf(refute_0_313, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209),
    introduced(tautology, [assume])).

cnf(refute_0_314, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     pi * 112150186033152000/3064655361465834062209 +
     arctan(-1139208930514575/1108767509542208) *
     -224300372066304000/3064655361465834062209 <= 0 |
     ln(10) <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(split, [], [refute_0_313])).

cnf(refute_0_315, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208) |
     ln(10) <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(arithmetic, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (arctan(-1139208930514575/1108767509542208) < pi * 1/2 |
     -1 <= -1139208930514575/1108767509542208 |
     pi * 1/2 <=
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1),
    inference(subst, [], [refute_0_305])).

cnf(refute_0_317, plain,
    (arctan(-1139208930514575/1108767509542208) < pi * 1/2 |
     pi <= -2722761945683125057/2278417861029150000),
    inference(arithmetic, [], [refute_0_316])).

cnf(refute_0_318, plain, (-2722761945683125057/2278417861029150000 < pi),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_319, plain,
    (arctan(-1139208930514575/1108767509542208) < pi * 1/2),
    inference(resolve, [], [refute_0_317, refute_0_318])).

cnf(refute_0_320, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     ln(10) <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(resolve, [], [refute_0_315, refute_0_319])).

cnf(refute_0_321, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) | 10 <= 0 |
     (-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(resolve, [], [refute_0_320, refute_0_312])).

cnf(refute_0_322, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     231993/101125 <=
     pi * -373833953443840000/3064655361465834062209 /
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(arithmetic, [], [refute_0_321])).

cnf(refute_0_323, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     231993/101125 *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209 |
     pi * 112150186033152000/3064655361465834062209 +
     arctan(-1139208930514575/1108767509542208) *
     -224300372066304000/3064655361465834062209 <= 0),
    inference(resolve, [], [refute_0_322, refute_0_307])).

cnf(refute_0_324, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208) |
     pi * 1707229/1391958 <= arctan(-1139208930514575/1108767509542208)),
    inference(arithmetic, [], [refute_0_323])).

cnf(refute_0_325, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     pi * 1707229/1391958 <= arctan(-1139208930514575/1108767509542208)),
    inference(resolve, [], [refute_0_324, refute_0_319])).

cnf(refute_0_326, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     -1 <= -1139208930514575/1108767509542208 |
     pi * 1707229/1391958 <=
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1),
    inference(resolve, [], [refute_0_325, refute_0_306])).

cnf(refute_0_327, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) |
     pi <= -210553904021621743782867/432197894051881636150000),
    inference(arithmetic, [], [refute_0_326])).

cnf(refute_0_328, plain,
    (-210553904021621743782867/432197894051881636150000 < pi),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_329, plain,
    (pi * -373833953443840000/3064655361465834062209 <
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209)),
    inference(resolve, [], [refute_0_327, refute_0_328])).

cnf(refute_0_330, plain,
    (pi * -20 <
     4673502317100267454362938592531700062477969974487978645856118128000/
     50506321767834152681255141268608666866792908661439445722000321 *
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     0 <=
     pi * -373833953443840000/3064655361465834062209 +
     ln(10) *
     (pi * -112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      224300372066304000/3064655361465834062209) |
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) <=
     4673502317100267454362938592531700062477969974487978645856118128000/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_331, plain,
    (~ lgen(0, X_000118, sqrt(X_000117)) | X_000118 <= sqrt(X_000117)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_332, plain,
    (X_000117 < 0 |
     ~ lgen(0, X_000118,
            368000 * X_000117 * (X_000117 + 8464000000) /
            (X_000117 ^ 2 + 50784000000 * X_000117 +
             71639296000000000000)) | lgen(0, X_000118, sqrt(X_000117))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_333, plain,
    (X_000117 < 0 |
     ~ lgen(0, X_000118,
            368000 * X_000117 * (X_000117 + 8464000000) /
            (X_000117 ^ 2 + 50784000000 * X_000117 +
             71639296000000000000)) | X_000118 <= sqrt(X_000117)),
    inference(resolve, [], [refute_0_332, refute_0_331])).

cnf(refute_0_334, plain,
    (X_000117 * (3114752000000000 + X_000117 * 368000) /
     (71639296000000000000 + X_000117 * (50784000000 + X_000117)) <
     X_000118 | X_000117 < 0 | X_000118 <= sqrt(X_000117)),
    inference(arithmetic, [], [refute_0_333])).

cnf(refute_0_335, plain,
    (2527162377580592078908020945889/295147905179352825856 < 0 |
     2527162377580592078908020945889/295147905179352825856 *
     (3114752000000000 +
      2527162377580592078908020945889/295147905179352825856 * 368000) /
     (71639296000000000000 +
      2527162377580592078908020945889/295147905179352825856 *
      (50784000000 +
       2527162377580592078908020945889/295147905179352825856)) <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) <=
     sqrt(2527162377580592078908020945889/295147905179352825856)),
    inference(subst, [], [refute_0_334])).

cnf(refute_0_336, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) *
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <
     pi * 20 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     (pi * -373833953443840000/3064655361465834062209) |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0 |
     (pi * 20 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      (pi * -373833953443840000/3064655361465834062209)) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <=
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_337, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) * (pi * -20) |
     0 <= pi * -20 |
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) <=
     (-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_338, plain,
    (~ lgen(0, X_000069, ln(X_000068)) | X_000069 <= ln(X_000068)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_339, plain,
    (~ lgen(0, X_000069, (X_000068 - 1) / X_000068) | X_000068 <= 0 |
     lgen(0, X_000069, ln(X_000068))),
    inference(subst, [], [ln_lower_bound_cf1])).

cnf(refute_0_340, plain,
    (~ lgen(0, X_000069, (X_000068 - 1) / X_000068) | X_000068 <= 0 |
     X_000069 <= ln(X_000068)),
    inference(resolve, [], [refute_0_339, refute_0_338])).

cnf(refute_0_341, plain,
    ((-1 + X_000068) / X_000068 < X_000069 | X_000068 <= 0 |
     X_000069 <= ln(X_000068)), inference(arithmetic, [], [refute_0_340])).

cnf(refute_0_342, plain,
    ((-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0 |
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209)),
    inference(subst, [], [refute_0_341])).

cnf(refute_0_343, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) * ln(10) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) |
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_344, plain,
    (ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) * pi <
     pi * -6 + arctan(-1139208930514575/1108767509542208) * 12 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) | pi <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_345, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) / ln(10) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_344])).

cnf(refute_0_346, plain, (0 < pi),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_347, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) / ln(10) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_345, refute_0_346])).

cnf(refute_0_348, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) * ln(10) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) |
     ln(10) <= 0), inference(resolve, [], [refute_0_343, refute_0_347])).

cnf(refute_0_349, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * -20) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_348])).

cnf(refute_0_350, plain,
    (ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(split, [], [refute_0_349])).

cnf(refute_0_351, plain,
    (ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_350])).

cnf(refute_0_352, plain,
    (ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_351, refute_0_346])).

cnf(refute_0_353, plain,
    ((-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_342, refute_0_352])).

cnf(refute_0_354, plain,
    ((-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_353])).

cnf(refute_0_355, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) * (pi * -20) |
     0 <= pi * -20 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_337, refute_0_354])).

cnf(refute_0_356, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (pi * 20 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      (pi * -373833953443840000/3064655361465834062209)) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) | pi <= 0),
    inference(arithmetic, [], [refute_0_355])).

cnf(refute_0_357, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (pi * 20 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      (pi * -373833953443840000/3064655361465834062209)) /
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_356, refute_0_346])).

cnf(refute_0_358, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) *
     (sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) <
     pi * 20 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     (pi * -373833953443840000/3064655361465834062209) |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_336, refute_0_357])).

cnf(refute_0_359, plain,
    (pi * -20 <
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (sqrt(2527162377580592078908020945889/295147905179352825856) <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     0 <=
     pi * -373833953443840000/3064655361465834062209 +
     ln(10) *
     (pi * -112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      224300372066304000/3064655361465834062209) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(split, [], [refute_0_359])).

cnf(refute_0_361, plain,
    (sqrt(2527162377580592078908020945889/295147905179352825856) <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_360])).

cnf(refute_0_362, plain,
    (sqrt(2527162377580592078908020945889/295147905179352825856) <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_361, refute_0_329])).

cnf(refute_0_363, plain,
    (2527162377580592078908020945889/295147905179352825856 < 0 |
     2527162377580592078908020945889/295147905179352825856 *
     (3114752000000000 +
      2527162377580592078908020945889/295147905179352825856 * 368000) /
     (71639296000000000000 +
      2527162377580592078908020945889/295147905179352825856 *
      (50784000000 +
       2527162377580592078908020945889/295147905179352825856)) <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_335, refute_0_362])).

cnf(refute_0_364, plain,
    (4673502317100267454362938592531700062477969974487978645856118128000/
     50506321767834152681255141268608666866792908661439445722000321 <
     pi * -20 /
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_363])).

cnf(refute_0_365, plain,
    (pi * -20 <
     4673502317100267454362938592531700062477969974487978645856118128000/
     50506321767834152681255141268608666866792908661439445722000321 *
     (pi * -373833953443840000/3064655361465834062209 +
      ln(10) *
      (pi * -112150186033152000/3064655361465834062209 +
       arctan(-1139208930514575/1108767509542208) *
       224300372066304000/3064655361465834062209)) |
     0 <=
     pi * -373833953443840000/3064655361465834062209 +
     ln(10) *
     (pi * -112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      224300372066304000/3064655361465834062209) |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_330, refute_0_364])).

cnf(refute_0_366, plain,
    (pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 <
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(arithmetic, [], [refute_0_365])).

cnf(refute_0_367, plain,
    (pi *
     -214332564415532725036822327013855303133856155929721705327908321/
     3526819038771018632070867785383718816067739622496674087593574400 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     arctan(-1139208930514575/1108767509542208) <=
     pi *
     -214332564415532725036822327013855303133856155929721705327908321/
     3526819038771018632070867785383718816067739622496674087593574400),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_368, plain,
    (pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 <
     231993/101125 *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     0 <=
     pi *
     -171025480019543519032027960264175370294629884320163037184000000/
     50506321767834152681255141268608666866792908661439445722000321 +
     arctan(-1139208930514575/1108767509542208) *
     342050960039087038064055920528350740589259768640326074368000000/
     50506321767834152681255141268608666866792908661439445722000321 |
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     231993/101125), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_369, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     10 <= 0 |
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     ln(10)), inference(subst, [], [refute_0_9])).

cnf(refute_0_370, plain,
    (pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 <
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    introduced(tautology, [assume])).

cnf(refute_0_371, plain,
    (ln(10) <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     0 <=
     pi *
     -171025480019543519032027960264175370294629884320163037184000000/
     50506321767834152681255141268608666866792908661439445722000321 +
     arctan(-1139208930514575/1108767509542208) *
     342050960039087038064055920528350740589259768640326074368000000/
     50506321767834152681255141268608666866792908661439445722000321 |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(split, [], [refute_0_370])).

cnf(refute_0_372, plain,
    (ln(10) <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208)),
    inference(arithmetic, [], [refute_0_371])).

cnf(refute_0_373, plain,
    (ln(10) <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_372, refute_0_319])).

cnf(refute_0_374, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     10 <= 0 |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_369, refute_0_373])).

cnf(refute_0_375, plain,
    (231993/101125 <
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 /
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 <
     231993/101125 *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) |
     0 <=
     pi *
     -171025480019543519032027960264175370294629884320163037184000000/
     50506321767834152681255141268608666866792908661439445722000321 +
     arctan(-1139208930514575/1108767509542208) *
     342050960039087038064055920528350740589259768640326074368000000/
     50506321767834152681255141268608666866792908661439445722000321 |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_368, refute_0_375])).

cnf(refute_0_377, plain,
    (pi *
     -214332564415532725036822327013855303133856155929721705327908321/
     3526819038771018632070867785383718816067739622496674087593574400 <
     arctan(-1139208930514575/1108767509542208) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321 |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208)),
    inference(arithmetic, [], [refute_0_376])).

cnf(refute_0_378, plain,
    (pi *
     -214332564415532725036822327013855303133856155929721705327908321/
     3526819038771018632070867785383718816067739622496674087593574400 <
     arctan(-1139208930514575/1108767509542208) |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_377, refute_0_319])).

cnf(refute_0_379, plain,
    (pi *
     -214332564415532725036822327013855303133856155929721705327908321/
     3526819038771018632070867785383718816067739622496674087593574400 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_367, refute_0_378])).

cnf(refute_0_380, plain,
    (4001120278365194815490688609229454332874755278917023850102741654684996460347392/
     406949285803775484005508876238978406635167517967859084269863806706024923797625 <
     pi |
     ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(arithmetic, [], [refute_0_379])).

cnf(refute_0_381, plain,
    (pi <=
     4001120278365194815490688609229454332874755278917023850102741654684996460347392/
     406949285803775484005508876238978406635167517967859084269863806706024923797625),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_382, plain,
    (ln(10) *
     (pi *
      -171025480019543519032027960264175370294629884320163037184000000/
      50506321767834152681255141268608666866792908661439445722000321 +
      arctan(-1139208930514575/1108767509542208) *
      342050960039087038064055920528350740589259768640326074368000000/
      50506321767834152681255141268608666866792908661439445722000321) <=
     pi *
     -440041501958204656851676291158255436353758558828245457160006420/
     50506321767834152681255141268608666866792908661439445722000321),
    inference(resolve, [], [refute_0_381, refute_0_380])).

cnf(refute_0_383, plain,
    (ln(10) *
     (pi * 112150186033152000/3064655361465834062209 +
      arctan(-1139208930514575/1108767509542208) *
      -224300372066304000/3064655361465834062209) <=
     pi * -373833953443840000/3064655361465834062209 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_382, refute_0_366])).

cnf(refute_0_384, plain,
    ((pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_383, refute_0_329])).

cnf(refute_0_385, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10)),
    inference(resolve, [], [refute_0_384, refute_0_301])).

cnf(refute_0_386, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * 20 / ln(10) * pi |
     pi <= 0), inference(resolve, [], [refute_0_24, refute_0_385])).

cnf(refute_0_387, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20) / ln(10) |
     pi <= 0), inference(arithmetic, [], [refute_0_386])).

cnf(refute_0_388, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20) / ln(10)),
    inference(resolve, [], [refute_0_387, refute_0_346])).

cnf(refute_0_389, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) * ln(10) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20) |
     ln(10) <= 0), inference(resolve, [], [refute_0_23, refute_0_388])).

cnf(refute_0_390, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) * (pi * 20)),
    inference(arithmetic, [], [refute_0_389])).

cnf(refute_0_391, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) | pi * 20 <= 0),
    inference(split, [], [refute_0_390])).

cnf(refute_0_392, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209) | pi <= 0),
    inference(arithmetic, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     ln(sqrt(2527162377580592078908020945889/295147905179352825856) *
        18691697672192000/3064655361465834062209)),
    inference(resolve, [], [refute_0_392, refute_0_346])).

cnf(refute_0_394, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     -1 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 |
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209 <= 0),
    inference(resolve, [], [refute_0_22, refute_0_393])).

cnf(refute_0_395, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 20) <
     -1 +
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     18691697672192000/3064655361465834062209),
    inference(arithmetic, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     (-1 +
      sqrt(2527162377580592078908020945889/295147905179352825856) *
      18691697672192000/3064655361465834062209) * (pi * 20) |
     pi * 20 <= 0), inference(resolve, [], [refute_0_17, refute_0_395])).

cnf(refute_0_397, plain,
    (pi * 20 +
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     (pi * 373833953443840000/3064655361465834062209) | pi <= 0),
    inference(arithmetic, [], [refute_0_396])).

cnf(refute_0_398, plain,
    (pi * 20 +
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     sqrt(2527162377580592078908020945889/295147905179352825856) *
     (pi * 373833953443840000/3064655361465834062209)),
    inference(resolve, [], [refute_0_397, refute_0_346])).

cnf(refute_0_399, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     sqrt(2527162377580592078908020945889/295147905179352825856) |
     pi * 373833953443840000/3064655361465834062209 <= 0),
    inference(split, [], [refute_0_398])).

cnf(refute_0_400, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     sqrt(2527162377580592078908020945889/295147905179352825856) |
     pi <= 0), inference(arithmetic, [], [refute_0_399])).

cnf(refute_0_401, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     sqrt(2527162377580592078908020945889/295147905179352825856)),
    inference(resolve, [], [refute_0_400, refute_0_346])).

cnf(refute_0_402, plain,
    (2527162377580592078908020945889/295147905179352825856 < 0 |
     (pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     (194672000000000 +
      2527162377580592078908020945889/295147905179352825856 *
      (138000 +
       2527162377580592078908020945889/295147905179352825856 * 1/368000)) /
     (8464000000 + 2527162377580592078908020945889/295147905179352825856)),
    inference(resolve, [], [refute_0_16, refute_0_401])).

cnf(refute_0_403, plain,
    ((pi * 20 +
      ln(10) *
      (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12)) /
     (pi * 373833953443840000/3064655361465834062209) <
     50506321767834152681255141268608666866792908661439445722000321/
     545819465729604363683111646639431975876487640999002112000),
    inference(arithmetic, [], [refute_0_402])).

cnf(refute_0_404, plain,
    (pi * 20 +
     ln(10) * (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) <
     50506321767834152681255141268608666866792908661439445722000321/
     545819465729604363683111646639431975876487640999002112000 *
     (pi * 373833953443840000/3064655361465834062209) |
     pi * 373833953443840000/3064655361465834062209 <= 0),
    inference(resolve, [], [refute_0_11, refute_0_403])).

cnf(refute_0_405, plain,
    (pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 <
     ln(10) * (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     pi <= 0), inference(arithmetic, [], [refute_0_404])).

cnf(refute_0_406, plain,
    (pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 <
     ln(10) * (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12)),
    inference(resolve, [], [refute_0_405, refute_0_346])).

cnf(refute_0_407, plain,
    (ln(10) <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     0 <= pi * -6 + arctan(-1139208930514575/1108767509542208) * 12),
    inference(split, [], [refute_0_406])).

cnf(refute_0_408, plain,
    (ln(10) <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208)),
    inference(arithmetic, [], [refute_0_407])).

cnf(refute_0_409, plain,
    (ln(10) <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12)),
    inference(resolve, [], [refute_0_408, refute_0_319])).

cnf(refute_0_410, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     10 <= 0), inference(resolve, [], [refute_0_10, refute_0_409])).

cnf(refute_0_411, plain,
    (231993/101125 <
     pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 /
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12)),
    inference(arithmetic, [], [refute_0_410])).

cnf(refute_0_412, plain,
    (pi *
     3313742377130635389177205235485892434332512014903596477846763755/
     380338986369826918354194612165881455900495462556297468210839552 <
     231993/101125 *
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) |
     0 <= pi * -6 + arctan(-1139208930514575/1108767509542208) * 12),
    inference(resolve, [], [refute_0_5, refute_0_411])).

cnf(refute_0_413, plain,
    (pi *
     96057565852967449038779211488856321334913015619562767230897531760799/
     117647976619860341692992894213599115464964858459764158056849733582848 <
     arctan(-1139208930514575/1108767509542208) |
     pi * 1/2 <= arctan(-1139208930514575/1108767509542208)),
    inference(arithmetic, [], [refute_0_412])).

cnf(refute_0_414, plain,
    (pi *
     96057565852967449038779211488856321334913015619562767230897531760799/
     117647976619860341692992894213599115464964858459764158056849733582848 <
     arctan(-1139208930514575/1108767509542208)),
    inference(resolve, [], [refute_0_413, refute_0_319])).

cnf(refute_0_415, plain,
    (pi *
     96057565852967449038779211488856321334913015619562767230897531760799/
     117647976619860341692992894213599115464964858459764158056849733582848 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208),
    inference(resolve, [], [refute_0_4, refute_0_414])).

cnf(refute_0_416, plain,
    (pi <
     -3336744101324726565737028008033996046186741284413654829900822123561557527224769314816/
     4559568202633016953878663296483063729315027904914510505724431293529164255870131059375),
    inference(arithmetic, [], [refute_0_415])).

cnf(refute_0_417, plain,
    (-3336744101324726565737028008033996046186741284413654829900822123561557527224769314816/
     4559568202633016953878663296483063729315027904914510505724431293529164255870131059375 <=
     pi),
    inference(decision, [], [refute_0_26, refute_0_298, refute_0_299])).

cnf(refute_0_418, plain, ($false),
    inference(resolve, [], [refute_0_417, refute_0_416])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-2-5.tptp
