% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-2-5-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(atan_upper_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, -157079/100000 - 1 / X, Y) |
     lgen(R, arctan(X), Y))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/30 *
            (6 * X ^ 4 + 481 * X ^ 3 + 1881 * X ^ 2 + 1281 * X + 131) *
            (X - 1) / (5 * X ^ 4 + 40 * X ^ 3 + 60 * X ^ 2 + 20 * X + 1),
            Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Fails_2_5', conjecture,
    (! [X] :
       (X = 57/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 43/10 *
                 ln(4250 ^ 2 *
                    (15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                     21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                     15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900) ^ 2))))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 57/128 =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 43/10 *
                 ln(4250 ^ 2 *
                    (15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                     21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                     15328072984 * X ^ 2 + 277555600) /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900) ^ 2)))),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Fails_2_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 57/128 =>
          interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   43/10 *
                   ln(4250 ^ 2 *
                      (15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                       21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                       15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900) ^ 2)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (~ interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   43/10 *
                   ln(4250 ^ 2 *
                      (15132100 * X ^ 10 + 1024027285 * X ^ 8 +
                       21404723599 * X ^ 6 + 129098541721 * X ^ 4 +
                       15328072984 * X ^ 2 + 277555600) /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900) ^ 2)) & X = 57/128)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)) &
     skoXC1 = 57/128), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 57/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (~ lgen(0, arctan(X_000061), X_000062) | arctan(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     lgen(0, arctan(X_000061), X_000062)),
    inference(subst, [], [atan_upper_bound_case_11])).

cnf(refute_0_2, plain,
    (~ lgen(0, -157079/100000 - 1 / X_000061, X_000062) | -1 <= X_000061 |
     arctan(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000062 < -157079/100000 + 1 / X_000061 * -1 | -1 <= X_000061 |
     arctan(X_000061) <= X_000062),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (pi *
     1405917034701595319456488473723434688950745903317325271517621063970616198931824721681393842298300905555955653/
     15616617180268515287954026041643763246410901513182614597276707723838918370650793188918868835482284032000000000 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     arctan(-1139208930514575/1108767509542208) <=
     pi *
     1405917034701595319456488473723434688950745903317325271517621063970616198931824721681393842298300905555955653/
     15616617180268515287954026041643763246410901513182614597276707723838918370650793188918868835482284032000000000),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     -446678480611581092408408689332449786110793361856324327473539497531314626957691060891491202937872635612375187/
     390415429506712882198850651041094081160272537829565364931917693095972959266269829722971720887057100800000000 *
     (pi * -43/10) | 0 <= pi * -43/10 |
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) <=
     -446678480611581092408408689332449786110793361856324327473539497531314626957691060891491202937872635612375187/
     390415429506712882198850651041094081160272537829565364931917693095972959266269829722971720887057100800000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000209, ln(X_000208)) | X_000209 <= ln(X_000208)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (~ lgen(0, X_000209,
            1/30 *
            (6 + 481 * X_000208 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 ^ 3 + 131 * X_000208 ^ 4) * (X_000208 - 1) /
            (X_000208 *
             (5 + 40 * X_000208 + 60 * X_000208 ^ 2 + 20 * X_000208 ^ 3 +
              X_000208 ^ 4))) | X_000208 <= 0 |
     lgen(0, X_000209, ln(X_000208))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000209,
            1/30 *
            (6 + 481 * X_000208 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 ^ 3 + 131 * X_000208 ^ 4) * (X_000208 - 1) /
            (X_000208 *
             (5 + 40 * X_000208 + 60 * X_000208 ^ 2 + 20 * X_000208 ^ 3 +
              X_000208 ^ 4))) | X_000208 <= 0 | X_000209 <= ln(X_000208)),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    ((-1/5 +
      X_000208 *
      (-95/6 +
       X_000208 *
       (-140/3 +
        X_000208 * (20 + X_000208 * (115/3 + X_000208 * 131/30))))) /
     (X_000208 *
      (5 +
       X_000208 * (40 + X_000208 * (60 + X_000208 * (20 + X_000208))))) <
     X_000209 | X_000208 <= 0 | X_000209 <= ln(X_000208)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    ((-1/5 +
      976133675290624000000/3064655361465834062209 *
      (-95/6 +
       976133675290624000000/3064655361465834062209 *
       (-140/3 +
        976133675290624000000/3064655361465834062209 *
        (20 +
         976133675290624000000/3064655361465834062209 *
         (115/3 +
          976133675290624000000/3064655361465834062209 * 131/30))))) /
     (976133675290624000000/3064655361465834062209 *
      (5 +
       976133675290624000000/3064655361465834062209 *
       (40 +
        976133675290624000000/3064655361465834062209 *
        (60 +
         976133675290624000000/3064655361465834062209 *
         (20 + 976133675290624000000/3064655361465834062209))))) <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) | 976133675290624000000/3064655361465834062209 <= 0 |
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) <= ln(976133675290624000000/3064655361465834062209)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (ln(976133675290624000000/3064655361465834062209) * 43/10 * pi <
     pi * -6 + arctan(-1139208930514575/1108767509542208) * 12 |
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi <=
     ln(976133675290624000000/3064655361465834062209) * 43/10 | pi <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_12, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_13, plain, (skoXC1 = 57/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_15, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 57/128 |
     skoXC1 ^ 2 = 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_16, plain, (skoXC1 != 57/128 | skoXC1 ^ 2 = 57/128 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_15])).

cnf(refute_0_17, plain, (skoXC1 ^ 2 = 57/128 ^ 2),
    inference(resolve, [], [refute_0_13, refute_0_16])).

cnf(refute_0_18, plain, (406627 * skoXC1 ^ 2 = 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_19, plain,
    (406627 * skoXC1 ^ 2 != 406627 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 | 406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_20, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 | 406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain, (406627 * skoXC1 ^ 2 = 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_17, refute_0_20])).

cnf(refute_0_22, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_23, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 57/128 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_24, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 57/128 ^ 2 |
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_21, refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_27, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 57/128 |
     skoXC1 ^ 4 = 57/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_28, plain, (skoXC1 != 57/128 | skoXC1 ^ 4 = 57/128 ^ 4),
    inference(resolve, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain, (skoXC1 ^ 4 = 57/128 ^ 4),
    inference(resolve, [], [refute_0_13, refute_0_28])).

cnf(refute_0_30, plain, (14325 * skoXC1 ^ 4 = 14325 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_31, plain,
    (14325 * skoXC1 ^ 4 != 14325 * skoXC1 ^ 4 | skoXC1 ^ 4 != 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_32, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 | 14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain, (14325 * skoXC1 ^ 4 = 14325 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_29, refute_0_32])).

cnf(refute_0_34, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_35, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_36, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 57/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_33, refute_0_36])).

cnf(refute_0_38, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_39, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_40, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_42, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (14325 * 57/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_37, refute_0_43])).

cnf(refute_0_45, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_25, refute_0_44])).

cnf(refute_0_46, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660),
    introduced(tautology, [refl])).

cnf(refute_0_47, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_45, refute_0_48])).

cnf(refute_0_50, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_51, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660 |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_52, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660 |
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_50, refute_0_51])).

cnf(refute_0_53, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_49, refute_0_52])).

cnf(refute_0_54, plain, (169932 * skoXC1 = 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_55, plain,
    (169932 * skoXC1 != 169932 * skoXC1 | skoXC1 != 57/128 |
     169932 * skoXC1 = 169932 * 57/128),
    introduced(tautology, [equality])).

cnf(refute_0_56, plain,
    (skoXC1 != 57/128 | 169932 * skoXC1 = 169932 * 57/128),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain, (169932 * skoXC1 = 169932 * 57/128),
    inference(resolve, [], [refute_0_13, refute_0_56])).

cnf(refute_0_58, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_59, plain,
    (169932 * skoXC1 != 169932 * 57/128 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    introduced(tautology, [equality])).

cnf(refute_0_60, plain,
    (169932 * skoXC1 != 169932 * 57/128 |
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_57, refute_0_60])).

cnf(refute_0_62, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_63, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 57/128 |
     skoXC1 ^ 5 = 57/128 ^ 5), introduced(tautology, [equality])).

cnf(refute_0_64, plain, (skoXC1 != 57/128 | skoXC1 ^ 5 = 57/128 ^ 5),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain, (skoXC1 ^ 5 = 57/128 ^ 5),
    inference(resolve, [], [refute_0_13, refute_0_64])).

cnf(refute_0_66, plain, (3890 * skoXC1 ^ 5 = 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_67, plain,
    (3890 * skoXC1 ^ 5 != 3890 * skoXC1 ^ 5 | skoXC1 ^ 5 != 57/128 ^ 5 |
     3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_68, plain,
    (skoXC1 ^ 5 != 57/128 ^ 5 | 3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain, (3890 * skoXC1 ^ 5 = 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_65, refute_0_68])).

cnf(refute_0_70, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_71, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 57/128 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_72, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 57/128 ^ 5 |
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_75, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 57/128 |
     skoXC1 ^ 3 = 57/128 ^ 3), introduced(tautology, [equality])).

cnf(refute_0_76, plain, (skoXC1 != 57/128 | skoXC1 ^ 3 = 57/128 ^ 3),
    inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain, (skoXC1 ^ 3 = 57/128 ^ 3),
    inference(resolve, [], [refute_0_13, refute_0_76])).

cnf(refute_0_78, plain, (105247 * skoXC1 ^ 3 = 105247 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_79, plain,
    (105247 * skoXC1 ^ 3 != 105247 * skoXC1 ^ 3 |
     skoXC1 ^ 3 != 57/128 ^ 3 | 105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_80, plain,
    (skoXC1 ^ 3 != 57/128 ^ 3 | 105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain, (105247 * skoXC1 ^ 3 = 105247 * 57/128 ^ 3),
    inference(resolve, [], [refute_0_77, refute_0_80])).

cnf(refute_0_82, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_83, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 57/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_84, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 57/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_82, refute_0_83])).

cnf(refute_0_85, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_81, refute_0_84])).

cnf(refute_0_86, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_87, plain,
    (105247 * 57/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5),
    inference(resolve, [], [refute_0_73, refute_0_87])).

cnf(refute_0_89, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_94, plain,
    (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128),
    inference(resolve, [], [refute_0_61, refute_0_94])).

cnf(refute_0_96, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_97, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_98, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_96, refute_0_97])).

cnf(refute_0_99, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_95, refute_0_98])).

cnf(refute_0_100, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_101, plain,
    ((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_53, refute_0_101])).

cnf(refute_0_103, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
     (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_102, refute_0_105])).

cnf(refute_0_107, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_106, refute_0_109])).

cnf(refute_0_111, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_110, refute_0_113])).

cnf(refute_0_115, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) - 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_120, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_122, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_123, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
            (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_121, refute_0_122])).

cnf(refute_0_124, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_110, refute_0_123])).

cnf(refute_0_125, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18),
    introduced(tautology, [refl])).

cnf(refute_0_126, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    introduced(tautology, [equality])).

cnf(refute_0_127, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_124, refute_0_127])).

cnf(refute_0_129, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 - 169932 * 57/128) /
             (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 - 16660))) + 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_130, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                    21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                    15328072984 * skoXC1 ^ 2 + 277555600) /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_17, refute_0_133])).

cnf(refute_0_135, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 57/128 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 57/128 ^ 2 |
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain, (217230815 * skoXC1 ^ 4 = 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 57/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain, (217230815 * skoXC1 ^ 4 = 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_29, refute_0_141])).

cnf(refute_0_143, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_144, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 57/128 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_145, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 57/128 ^ 4 |
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_142, refute_0_145])).

cnf(refute_0_147, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_148, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 57/128 |
     skoXC1 ^ 6 = 57/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_149, plain, (skoXC1 != 57/128 | skoXC1 ^ 6 = 57/128 ^ 6),
    inference(resolve, [], [refute_0_147, refute_0_148])).

cnf(refute_0_150, plain, (skoXC1 ^ 6 = 57/128 ^ 6),
    inference(resolve, [], [refute_0_13, refute_0_149])).

cnf(refute_0_151, plain, (3783025 * skoXC1 ^ 6 = 3783025 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_152, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_153, plain,
    (skoXC1 ^ 6 != 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain, (3783025 * skoXC1 ^ 6 = 3783025 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_150, refute_0_153])).

cnf(refute_0_155, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_156, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_157, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 57/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_155, refute_0_156])).

cnf(refute_0_158, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_154, refute_0_157])).

cnf(refute_0_159, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_160, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_146, refute_0_160])).

cnf(refute_0_162, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_163, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_164, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_161, refute_0_164])).

cnf(refute_0_166, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2), inference(subst, [], [refute_0_42])).

cnf(refute_0_167, plain,
    (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_138, refute_0_167])).

cnf(refute_0_169, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900), introduced(tautology, [refl])).

cnf(refute_0_170, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    introduced(tautology, [equality])).

cnf(refute_0_171, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_168, refute_0_171])).

cnf(refute_0_173, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_174, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_175, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
     3120782021 * 57/128 ^ 2 + 69388900 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_173, refute_0_174])).

cnf(refute_0_176, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_175])).

cnf(refute_0_177, plain,
    (4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_178, plain,
    (4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_179, plain,
    ((3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_176, refute_0_179])).

cnf(refute_0_181, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_182, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 57/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_183, plain,
    (skoXC1 ^ 2 != 57/128 ^ 2 |
     15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_181, refute_0_182])).

cnf(refute_0_184, plain,
    (15328072984 * skoXC1 ^ 2 = 15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_17, refute_0_183])).

cnf(refute_0_185, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_186, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 57/128 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_187, plain,
    (15328072984 * skoXC1 ^ 2 != 15328072984 * 57/128 ^ 2 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_185, refute_0_186])).

cnf(refute_0_188, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_184, refute_0_187])).

cnf(refute_0_189, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_190, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 57/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_191, plain,
    (skoXC1 ^ 4 != 57/128 ^ 4 |
     129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_189, refute_0_190])).

cnf(refute_0_192, plain,
    (129098541721 * skoXC1 ^ 4 = 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_29, refute_0_191])).

cnf(refute_0_193, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_194, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 57/128 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_195, plain,
    (129098541721 * skoXC1 ^ 4 != 129098541721 * 57/128 ^ 4 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_193, refute_0_194])).

cnf(refute_0_196, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_192, refute_0_195])).

cnf(refute_0_197, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_198, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 57/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_199, plain,
    (skoXC1 ^ 6 != 57/128 ^ 6 |
     21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_197, refute_0_198])).

cnf(refute_0_200, plain,
    (21404723599 * skoXC1 ^ 6 = 21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_150, refute_0_199])).

cnf(refute_0_201, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_202, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 57/128 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_203, plain,
    (21404723599 * skoXC1 ^ 6 != 21404723599 * 57/128 ^ 6 |
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_201, refute_0_202])).

cnf(refute_0_204, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_200, refute_0_203])).

cnf(refute_0_205, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_206, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 57/128 |
     skoXC1 ^ 8 = 57/128 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_207, plain, (skoXC1 != 57/128 | skoXC1 ^ 8 = 57/128 ^ 8),
    inference(resolve, [], [refute_0_205, refute_0_206])).

cnf(refute_0_208, plain, (skoXC1 ^ 8 = 57/128 ^ 8),
    inference(resolve, [], [refute_0_13, refute_0_207])).

cnf(refute_0_209, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_210, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 57/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_211, plain,
    (skoXC1 ^ 8 != 57/128 ^ 8 |
     1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_209, refute_0_210])).

cnf(refute_0_212, plain,
    (1024027285 * skoXC1 ^ 8 = 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_208, refute_0_211])).

cnf(refute_0_213, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_214, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 57/128 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_215, plain,
    (1024027285 * skoXC1 ^ 8 != 1024027285 * 57/128 ^ 8 |
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_213, refute_0_214])).

cnf(refute_0_216, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_212, refute_0_215])).

cnf(refute_0_217, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_218, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 57/128 |
     skoXC1 ^ 10 = 57/128 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_219, plain, (skoXC1 != 57/128 | skoXC1 ^ 10 = 57/128 ^ 10),
    inference(resolve, [], [refute_0_217, refute_0_218])).

cnf(refute_0_220, plain, (skoXC1 ^ 10 = 57/128 ^ 10),
    inference(resolve, [], [refute_0_13, refute_0_219])).

cnf(refute_0_221, plain, (15132100 * skoXC1 ^ 10 = 15132100 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_222, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * skoXC1 ^ 10 |
     skoXC1 ^ 10 != 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_223, plain,
    (skoXC1 ^ 10 != 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    inference(resolve, [], [refute_0_221, refute_0_222])).

cnf(refute_0_224, plain, (15132100 * skoXC1 ^ 10 = 15132100 * 57/128 ^ 10),
    inference(resolve, [], [refute_0_220, refute_0_223])).

cnf(refute_0_225, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_226, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_227, plain,
    (15132100 * skoXC1 ^ 10 != 15132100 * 57/128 ^ 10 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_225, refute_0_226])).

cnf(refute_0_228, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_224, refute_0_227])).

cnf(refute_0_229, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_230, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8),
    inference(resolve, [], [refute_0_216, refute_0_230])).

cnf(refute_0_232, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_233, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_234, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_232, refute_0_233])).

cnf(refute_0_235, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_231, refute_0_234])).

cnf(refute_0_236, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6), inference(subst, [], [refute_0_42])).

cnf(refute_0_237, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_235, refute_0_236])).

cnf(refute_0_238, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6),
    inference(resolve, [], [refute_0_204, refute_0_237])).

cnf(refute_0_239, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_240, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_241, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_239, refute_0_240])).

cnf(refute_0_242, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_238, refute_0_241])).

cnf(refute_0_243, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_244, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_242, refute_0_243])).

cnf(refute_0_245, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4),
    inference(resolve, [], [refute_0_196, refute_0_244])).

cnf(refute_0_246, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_247, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_248, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_246, refute_0_247])).

cnf(refute_0_249, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_245, refute_0_248])).

cnf(refute_0_250, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2), inference(subst, [], [refute_0_42])).

cnf(refute_0_251, plain,
    (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_249, refute_0_250])).

cnf(refute_0_252, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2),
    inference(resolve, [], [refute_0_188, refute_0_251])).

cnf(refute_0_253, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600), introduced(tautology, [refl])).

cnf(refute_0_254, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    introduced(tautology, [equality])).

cnf(refute_0_255, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_253, refute_0_254])).

cnf(refute_0_256, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 =
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600),
    inference(resolve, [], [refute_0_252, refute_0_255])).

cnf(refute_0_257, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600)),
    introduced(tautology, [refl])).

cnf(refute_0_258, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) |
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600)),
    introduced(tautology, [equality])).

cnf(refute_0_259, plain,
    (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 !=
     15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
     21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
     15328072984 * 57/128 ^ 2 + 277555600 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_257, refute_0_258])).

cnf(refute_0_260, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600)),
    inference(resolve, [], [refute_0_256, refute_0_259])).

cnf(refute_0_261, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_262, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_263, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_261, refute_0_262])).

cnf(refute_0_264, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_260, refute_0_263])).

cnf(refute_0_265, plain,
    (4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_266, plain,
    (4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_264, refute_0_265])).

cnf(refute_0_267, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 =
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2),
    inference(resolve, [], [refute_0_180, refute_0_266])).

cnf(refute_0_268, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_269, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_270, plain,
    (4250 ^ 2 *
     (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
      21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
      15328072984 * skoXC1 ^ 2 + 277555600) /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) ^ 2 !=
     4250 ^ 2 *
     (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
      21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
      15328072984 * 57/128 ^ 2 + 277555600) /
     (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
      3120782021 * 57/128 ^ 2 + 69388900) ^ 2 |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_268, refute_0_269])).

cnf(refute_0_271, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_267, refute_0_270])).

cnf(refute_0_272, plain,
    (43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_273, plain,
    (43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) |
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2) |
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_274, plain,
    (ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2) |
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_272, refute_0_273])).

cnf(refute_0_275, plain,
    (43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) =
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2)),
    inference(resolve, [], [refute_0_271, refute_0_274])).

cnf(refute_0_276, plain,
    (43/10 *
     ln(4250 ^ 2 *
        (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
         21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
         15328072984 * skoXC1 ^ 2 + 277555600) /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900) ^ 2) !=
     43/10 *
     ln(4250 ^ 2 *
        (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
         21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
         15328072984 * 57/128 ^ 2 + 277555600) /
        (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
         3120782021 * 57/128 ^ 2 + 69388900) ^ 2) |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                    21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
                    15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_277, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                    21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
                    15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_275, refute_0_276])).

cnf(refute_0_278, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                    21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
                    15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                       169932 * 57/128) /
                      (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_277, refute_0_130])).

cnf(refute_0_279, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                    21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
                    15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900) ^ 2)) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              43/10 *
              ln(4250 ^ 2 *
                 (15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
                  21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
                  15328072984 * skoXC1 ^ 2 + 277555600) /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_278, refute_0_120])).

cnf(refute_0_280, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 57/128 ^ 3 + 3890 * 57/128 ^ 5 -
                         169932 * 57/128) /
                        (14325 * 57/128 ^ 4 + 406627 * 57/128 ^ 2 -
                         16660))) + 18,
                43/10 *
                ln(4250 ^ 2 *
                   (15132100 * 57/128 ^ 10 + 1024027285 * 57/128 ^ 8 +
                    21404723599 * 57/128 ^ 6 + 129098541721 * 57/128 ^ 4 +
                    15328072984 * 57/128 ^ 2 + 277555600) /
                   (3783025 * 57/128 ^ 6 + 217230815 * 57/128 ^ 4 +
                    3120782021 * 57/128 ^ 2 + 69388900) ^ 2))),
    inference(resolve, [], [refute_0_279, refute_0_12])).

cnf(refute_0_281, plain,
    (~ interval(0,
                (pi * -6 +
                 arctan(-1139208930514575/1108767509542208) * 12) / pi, 0,
                (pi * 6 +
                 arctan(-1139208930514575/1108767509542208) * -12) / pi,
                ln(976133675290624000000/3064655361465834062209) * 43/10) |
     pi = 0), inference(arithmetic, [], [refute_0_280])).

cnf(refute_0_282, plain,
    (~ lgen(0, ln(976133675290624000000/3064655361465834062209) * 43/10,
            (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
            pi, ln(976133675290624000000/3064655361465834062209) * 43/10) |
     interval(0,
              (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
              pi, 0,
              (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
              pi,
              ln(976133675290624000000/3064655361465834062209) * 43/10)),
    inference(subst, [], [interval_intro])).

cnf(refute_0_283, plain,
    (~ lgen(0, ln(976133675290624000000/3064655361465834062209) * 43/10,
            (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) /
            pi, ln(976133675290624000000/3064655361465834062209) * 43/10) |
     pi = 0), inference(resolve, [], [refute_0_282, refute_0_281])).

cnf(refute_0_284, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi <
     ln(976133675290624000000/3064655361465834062209) * 43/10 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi |
     pi = 0), inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_286, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_287, plain, (pi != 0),
    inference(decision, [], [refute_0_13, refute_0_285, refute_0_286])).

cnf(refute_0_288, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi <
     ln(976133675290624000000/3064655361465834062209) * 43/10 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi),
    inference(resolve, [], [refute_0_284, refute_0_287])).

cnf(refute_0_289, plain,
    (pi *
     161825883892233463664021870527495035282165469864321383487325107930007899893231796168051198435103247309194921657/
     177838537739284539037924345390336521515414256429668334994224970866745391213387724379009352268924996462850224400 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     arctan(-1139208930514575/1108767509542208) <=
     pi *
     161825883892233463664021870527495035282165469864321383487325107930007899893231796168051198435103247309194921657/
     177838537739284539037924345390336521515414256429668334994224970866745391213387724379009352268924996462850224400),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_290, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     -5086508024831943777562304499929774966822674998801433673735764360230363089758460510131152718649354586821149497/
     4445963443482113475948108634758413037885356410741708374855624271668634780334693109475233806723124911571255610 *
     (pi * 43/10) |
     -5086508024831943777562304499929774966822674998801433673735764360230363089758460510131152718649354586821149497/
     4445963443482113475948108634758413037885356410741708374855624271668634780334693109475233806723124911571255610 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) | pi * 43/10 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_291, plain,
    (~ lgen(0, ln(X_000201), X_000202) | ln(X_000201) <= X_000202),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_292, plain,
    (~ lgen(0,
            1/30 *
            (6 * X_000201 ^ 4 + 481 * X_000201 ^ 3 + 1881 * X_000201 ^ 2 +
             1281 * X_000201 + 131) * (X_000201 - 1) /
            (5 * X_000201 ^ 4 + 40 * X_000201 ^ 3 + 60 * X_000201 ^ 2 +
             20 * X_000201 + 1), X_000202) | X_000201 <= 0 |
     lgen(0, ln(X_000201), X_000202)),
    inference(subst, [], [ln_upper_bound_cf9])).

cnf(refute_0_293, plain,
    (~ lgen(0,
            1/30 *
            (6 * X_000201 ^ 4 + 481 * X_000201 ^ 3 + 1881 * X_000201 ^ 2 +
             1281 * X_000201 + 131) * (X_000201 - 1) /
            (5 * X_000201 ^ 4 + 40 * X_000201 ^ 3 + 60 * X_000201 ^ 2 +
             20 * X_000201 + 1), X_000202) | X_000201 <= 0 |
     ln(X_000201) <= X_000202),
    inference(resolve, [], [refute_0_292, refute_0_291])).

cnf(refute_0_294, plain,
    (X_000202 <
     (-131/30 +
      X_000201 *
      (-115/3 +
       X_000201 *
       (-20 + X_000201 * (140/3 + X_000201 * (95/6 + X_000201 * 1/5))))) /
     (1 +
      X_000201 * (20 + X_000201 * (60 + X_000201 * (40 + X_000201 * 5)))) |
     X_000201 <= 0 | ln(X_000201) <= X_000202),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) <
     (-131/30 +
      976133675290624000000/3064655361465834062209 *
      (-115/3 +
       976133675290624000000/3064655361465834062209 *
       (-20 +
        976133675290624000000/3064655361465834062209 *
        (140/3 +
         976133675290624000000/3064655361465834062209 *
         (95/6 + 976133675290624000000/3064655361465834062209 * 1/5))))) /
     (1 +
      976133675290624000000/3064655361465834062209 *
      (20 +
       976133675290624000000/3064655361465834062209 *
       (60 +
        976133675290624000000/3064655361465834062209 *
        (40 + 976133675290624000000/3064655361465834062209 * 5)))) |
     976133675290624000000/3064655361465834062209 <= 0 |
     ln(976133675290624000000/3064655361465834062209) <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10)), inference(subst, [], [refute_0_294])).

cnf(refute_0_296, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(976133675290624000000/3064655361465834062209) * 43/10 * pi |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_297, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(976133675290624000000/3064655361465834062209) * (pi * 43/10) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi <= 0), inference(arithmetic, [], [refute_0_296])).

cnf(refute_0_298, plain, (0 < pi),
    inference(decision, [], [refute_0_13, refute_0_285, refute_0_286])).

cnf(refute_0_299, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(976133675290624000000/3064655361465834062209) * (pi * 43/10) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_297, refute_0_298])).

cnf(refute_0_300, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) < ln(976133675290624000000/3064655361465834062209) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi * 43/10 <= 0), inference(split, [], [refute_0_299])).

cnf(refute_0_301, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) < ln(976133675290624000000/3064655361465834062209) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi <= 0), inference(arithmetic, [], [refute_0_300])).

cnf(refute_0_302, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) < ln(976133675290624000000/3064655361465834062209) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_301, refute_0_298])).

cnf(refute_0_303, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) <
     (-131/30 +
      976133675290624000000/3064655361465834062209 *
      (-115/3 +
       976133675290624000000/3064655361465834062209 *
       (-20 +
        976133675290624000000/3064655361465834062209 *
        (140/3 +
         976133675290624000000/3064655361465834062209 *
         (95/6 + 976133675290624000000/3064655361465834062209 * 1/5))))) /
     (1 +
      976133675290624000000/3064655361465834062209 *
      (20 +
       976133675290624000000/3064655361465834062209 *
       (60 +
        976133675290624000000/3064655361465834062209 *
        (40 + 976133675290624000000/3064655361465834062209 * 5)))) |
     976133675290624000000/3064655361465834062209 <= 0 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_295, refute_0_302])).

cnf(refute_0_304, plain,
    ((pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * 43/10) <
     -5086508024831943777562304499929774966822674998801433673735764360230363089758460510131152718649354586821149497/
     4445963443482113475948108634758413037885356410741708374855624271668634780334693109475233806723124911571255610 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(arithmetic, [], [refute_0_303])).

cnf(refute_0_305, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     -5086508024831943777562304499929774966822674998801433673735764360230363089758460510131152718649354586821149497/
     4445963443482113475948108634758413037885356410741708374855624271668634780334693109475233806723124911571255610 *
     (pi * 43/10) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi * 43/10 <= 0),
    inference(resolve, [], [refute_0_290, refute_0_304])).

cnf(refute_0_306, plain,
    (pi *
     161825883892233463664021870527495035282165469864321383487325107930007899893231796168051198435103247309194921657/
     177838537739284539037924345390336521515414256429668334994224970866745391213387724379009352268924996462850224400 <
     arctan(-1139208930514575/1108767509542208) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi |
     pi <= 0), inference(arithmetic, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (pi *
     161825883892233463664021870527495035282165469864321383487325107930007899893231796168051198435103247309194921657/
     177838537739284539037924345390336521515414256429668334994224970866745391213387724379009352268924996462850224400 <
     arctan(-1139208930514575/1108767509542208) |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_306, refute_0_298])).

cnf(refute_0_308, plain,
    (pi *
     161825883892233463664021870527495035282165469864321383487325107930007899893231796168051198435103247309194921657/
     177838537739284539037924345390336521515414256429668334994224970866745391213387724379009352268924996462850224400 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_289, refute_0_307])).

cnf(refute_0_309, plain,
    (pi <
     -1210530007581140589123181803218363112830762479566429004489793358196832612735295638116120523841098599501376661921099539814281977/
     1843534921184470736547302691851402087507937307874011051141484806750691319419995238421434574676465902413765793464337738216507750 |
     ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(arithmetic, [], [refute_0_308])).

cnf(refute_0_310, plain,
    (-1210530007581140589123181803218363112830762479566429004489793358196832612735295638116120523841098599501376661921099539814281977/
     1843534921184470736547302691851402087507937307874011051141484806750691319419995238421434574676465902413765793464337738216507750 <=
     pi),
    inference(decision, [], [refute_0_13, refute_0_285, refute_0_286])).

cnf(refute_0_311, plain,
    (ln(976133675290624000000/3064655361465834062209) * 43/10 <=
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) / pi),
    inference(resolve, [], [refute_0_310, refute_0_309])).

cnf(refute_0_312, plain,
    (ln(976133675290624000000/3064655361465834062209) * 43/10 <
     (pi * -6 + arctan(-1139208930514575/1108767509542208) * 12) / pi),
    inference(resolve, [], [refute_0_311, refute_0_288])).

cnf(refute_0_313, plain,
    (ln(976133675290624000000/3064655361465834062209) * 43/10 * pi <
     pi * -6 + arctan(-1139208930514575/1108767509542208) * 12 | pi <= 0),
    inference(resolve, [], [refute_0_11, refute_0_312])).

cnf(refute_0_314, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(976133675290624000000/3064655361465834062209) * (pi * -43/10) |
     pi <= 0), inference(arithmetic, [], [refute_0_313])).

cnf(refute_0_315, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     ln(976133675290624000000/3064655361465834062209) * (pi * -43/10)),
    inference(resolve, [], [refute_0_314, refute_0_298])).

cnf(refute_0_316, plain,
    (ln(976133675290624000000/3064655361465834062209) <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) | 0 <= pi * -43/10),
    inference(split, [], [refute_0_315])).

cnf(refute_0_317, plain,
    (ln(976133675290624000000/3064655361465834062209) <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) | pi <= 0), inference(arithmetic, [], [refute_0_316])).

cnf(refute_0_318, plain,
    (ln(976133675290624000000/3064655361465834062209) <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10)), inference(resolve, [], [refute_0_317, refute_0_298])).

cnf(refute_0_319, plain,
    ((-1/5 +
      976133675290624000000/3064655361465834062209 *
      (-95/6 +
       976133675290624000000/3064655361465834062209 *
       (-140/3 +
        976133675290624000000/3064655361465834062209 *
        (20 +
         976133675290624000000/3064655361465834062209 *
         (115/3 +
          976133675290624000000/3064655361465834062209 * 131/30))))) /
     (976133675290624000000/3064655361465834062209 *
      (5 +
       976133675290624000000/3064655361465834062209 *
       (40 +
        976133675290624000000/3064655361465834062209 *
        (60 +
         976133675290624000000/3064655361465834062209 *
         (20 + 976133675290624000000/3064655361465834062209))))) <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10) | 976133675290624000000/3064655361465834062209 <= 0),
    inference(resolve, [], [refute_0_10, refute_0_318])).

cnf(refute_0_320, plain,
    (-446678480611581092408408689332449786110793361856324327473539497531314626957691060891491202937872635612375187/
     390415429506712882198850651041094081160272537829565364931917693095972959266269829722971720887057100800000000 <
     (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12) /
     (pi * -43/10)), inference(arithmetic, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (pi * 6 + arctan(-1139208930514575/1108767509542208) * -12 <
     -446678480611581092408408689332449786110793361856324327473539497531314626957691060891491202937872635612375187/
     390415429506712882198850651041094081160272537829565364931917693095972959266269829722971720887057100800000000 *
     (pi * -43/10) | 0 <= pi * -43/10),
    inference(resolve, [], [refute_0_5, refute_0_320])).

cnf(refute_0_322, plain,
    (pi *
     1405917034701595319456488473723434688950745903317325271517621063970616198931824721681393842298300905555955653/
     15616617180268515287954026041643763246410901513182614597276707723838918370650793188918868835482284032000000000 <
     arctan(-1139208930514575/1108767509542208) | pi <= 0),
    inference(arithmetic, [], [refute_0_321])).

cnf(refute_0_323, plain,
    (pi *
     1405917034701595319456488473723434688950745903317325271517621063970616198931824721681393842298300905555955653/
     15616617180268515287954026041643763246410901513182614597276707723838918370650793188918868835482284032000000000 <
     arctan(-1139208930514575/1108767509542208)),
    inference(resolve, [], [refute_0_322, refute_0_298])).

cnf(refute_0_324, plain,
    (pi *
     1405917034701595319456488473723434688950745903317325271517621063970616198931824721681393842298300905555955653/
     15616617180268515287954026041643763246410901513182614597276707723838918370650793188918868835482284032000000000 <
     -157079/100000 + 1 / -1139208930514575/1108767509542208 * -1 |
     -1 <= -1139208930514575/1108767509542208),
    inference(resolve, [], [refute_0_4, refute_0_323])).

cnf(refute_0_325, plain,
    (pi <
     -425203309787364208091100477971155892669946406146354910026807802735298592581891859496226619819730427256251211418501898240000/
     64065329659785081252104313088685188171627938731473745849082448800904283616738840587982087446097574073483179322321486805699),
    inference(arithmetic, [], [refute_0_324])).

cnf(refute_0_326, plain,
    (-425203309787364208091100477971155892669946406146354910026807802735298592581891859496226619819730427256251211418501898240000/
     64065329659785081252104313088685188171627938731473745849082448800904283616738840587982087446097574073483179322321486805699 <=
     pi),
    inference(decision, [], [refute_0_286, refute_0_285, refute_0_13])).

cnf(refute_0_327, plain, ($false),
    inference(resolve, [], [refute_0_326, refute_0_325])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-2-5-weak.tptp
