% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-6.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_33, axiom,
    (X < 0 |
     ~ lgen(R,
            1/15 * (64 * X ^ 4 + 735 * X ^ 2 + 945) * X /
            (15 * X ^ 4 + 70 * X ^ 2 + 63), Y) | lgen(R, arctan(X), Y))).

cnf(ln_lower_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, Y, 1/2 * (1 + 5 * X) * (X - 1) / (X * (2 + X))) |
     lgen(R, Y, ln(X)))).

cnf(ln_lower_bound_cf9, axiom,
    (X <= 0 |
     ~ lgen(R, Y,
            1/30 *
            (6 + 481 * X + 1881 * X ^ 2 + 1281 * X ^ 3 + 131 * X ^ 4) *
            (X - 1) /
            (X * (5 + 40 * X + 60 * X ^ 2 + 20 * X ^ 3 + X ^ 4))) |
     lgen(R, Y, ln(X)))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Fails_1_6', conjecture,
    (! [X, Y] :
       ((X = 265/128 & 0 <= Y &
         Y ^ 2 =
         15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
         129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 * Y /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900)) + 3/2)))).

fof(subgoal_0, plain,
    (! [X, Y] :
       ((X = 265/128 & 0 <= Y &
         Y ^ 2 =
         15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
         129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) =>
        interval(0,
                 -12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                 0,
                 12 / pi *
                 (-pi -
                  arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                         (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                 20 / ln(10) *
                 ln(4250 * Y /
                    (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                     3120782021 * X ^ 2 + 69388900)) + 3/2))),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Fails_1_6'])).

fof(negate_0_0, plain,
    (~ ! [X, Y] :
         ((X = 265/128 & 0 <= Y &
           Y ^ 2 =
           15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
           129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600) =>
          interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 * Y /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900)) + 3/2))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X, Y] :
       (~ interval(0,
                   -12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) - 18,
                   0,
                   12 / pi *
                   (-pi -
                    arctan((105247 * X ^ 3 + 3890 * X ^ 5 - 169932 * X) /
                           (14325 * X ^ 4 + 406627 * X ^ 2 - 16660))) + 18,
                   20 / ln(10) *
                   ln(4250 * Y /
                      (3783025 * X ^ 6 + 217230815 * X ^ 4 +
                       3120782021 * X ^ 2 + 69388900)) + 3/2) & 0 <= Y &
        X = 265/128 &
        Y ^ 2 =
        15132100 * X ^ 10 + 1024027285 * X ^ 8 + 21404723599 * X ^ 6 +
        129098541721 * X ^ 4 + 15328072984 * X ^ 2 + 277555600)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) &
     0 <= skoYC1 &
     skoYC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600 & skoXC1 = 265/128),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 265/128),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain,
    (skoYC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_6, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_7, plain, (0 <= skoYC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (~ lgen(0, arctan(X_000040), X_000041) | arctan(X_000040) <= X_000041),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (X_000040 < 0 | ~ lgen(0, X_000040, X_000041) |
     lgen(0, arctan(X_000040), X_000041)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_2, plain,
    (X_000040 < 0 | ~ lgen(0, X_000040, X_000041) |
     arctan(X_000040) <= X_000041),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000041 < X_000040 | X_000040 < 0 | arctan(X_000040) <= X_000041),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (836177642653679/2278501083899584 < 0 |
     pi * 6132937/5567832 +
     skoYC1 * (pi * -378039585420083200000/2171363234305355192028683907) <
     836177642653679/2278501083899584 |
     arctan(836177642653679/2278501083899584) <=
     pi * 6132937/5567832 +
     skoYC1 * (pi * -378039585420083200000/2171363234305355192028683907)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi * 20 + skoYC1 * (pi * -14953358137753600/3119868895908289175433) <
     231993/101125 *
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12 |
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) <=
     231993/101125), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000209, ln(X_000208)) | X_000209 <= ln(X_000208)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (~ lgen(0, X_000209,
            1/30 *
            (6 + 481 * X_000208 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 ^ 3 + 131 * X_000208 ^ 4) * (X_000208 - 1) /
            (X_000208 *
             (5 + 40 * X_000208 + 60 * X_000208 ^ 2 + 20 * X_000208 ^ 3 +
              X_000208 ^ 4))) | X_000208 <= 0 |
     lgen(0, X_000209, ln(X_000208))),
    inference(subst, [], [ln_lower_bound_cf9])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000209,
            1/30 *
            (6 + 481 * X_000208 + 1881 * X_000208 ^ 2 +
             1281 * X_000208 ^ 3 + 131 * X_000208 ^ 4) * (X_000208 - 1) /
            (X_000208 *
             (5 + 40 * X_000208 + 60 * X_000208 ^ 2 + 20 * X_000208 ^ 3 +
              X_000208 ^ 4))) | X_000208 <= 0 | X_000209 <= ln(X_000208)),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    ((-1/5 +
      X_000208 *
      (-95/6 +
       X_000208 *
       (-140/3 +
        X_000208 * (20 + X_000208 * (115/3 + X_000208 * 131/30))))) /
     (X_000208 *
      (5 +
       X_000208 * (40 + X_000208 * (60 + X_000208 * (20 + X_000208))))) <
     X_000209 | X_000208 <= 0 | X_000209 <= ln(X_000208)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     10 <= 0 |
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) <=
     ln(10)), inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1 + skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20) |
     pi * 20 <= 0 |
     -1 + skoYC1 * 747667906887680/3119868895908289175433 <=
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain,
    (~ lgen(0, ln(X_000047), X_000048) | ln(X_000047) <= X_000048),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_13, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     lgen(0, ln(X_000047), X_000048)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_14, plain,
    (~ lgen(0, X_000047 - 1, X_000048) | X_000047 <= 0 |
     ln(X_000047) <= X_000048),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    (X_000048 < -1 + X_000047 | X_000047 <= 0 | ln(X_000047) <= X_000048),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < -1 + skoYC1 * 747667906887680/3119868895908289175433 |
     skoYC1 * 747667906887680/3119868895908289175433 <= 0 |
     ln(skoYC1 * 747667906887680/3119868895908289175433) <=
     ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20)), inference(subst, [], [refute_0_15])).

cnf(refute_0_17, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * 3/2) +
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20) |
     (ln(10) * (pi * 3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20)) /
     ln(10) <= pi * 6 + arctan(836177642653679/2278501083899584) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_18, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) *
     pi |
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) <=
     (pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi |
     pi <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_19, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_20, plain, (skoXC1 = 265/128),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_21, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_22, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 265/128 |
     skoXC1 ^ 2 = 265/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_23, plain, (skoXC1 != 265/128 | skoXC1 ^ 2 = 265/128 ^ 2),
    inference(resolve, [], [refute_0_21, refute_0_22])).

cnf(refute_0_24, plain, (skoXC1 ^ 2 = 265/128 ^ 2),
    inference(resolve, [], [refute_0_20, refute_0_23])).

cnf(refute_0_25, plain, (406627 * skoXC1 ^ 2 = 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_26, plain,
    (406627 * skoXC1 ^ 2 != 406627 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 265/128 ^ 2 |
     406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_27, plain,
    (skoXC1 ^ 2 != 265/128 ^ 2 |
     406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_25, refute_0_26])).

cnf(refute_0_28, plain, (406627 * skoXC1 ^ 2 = 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_24, refute_0_27])).

cnf(refute_0_29, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_30, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 265/128 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_31, plain,
    (406627 * skoXC1 ^ 2 != 406627 * 265/128 ^ 2 |
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_28, refute_0_31])).

cnf(refute_0_33, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_34, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 265/128 |
     skoXC1 ^ 4 = 265/128 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_35, plain, (skoXC1 != 265/128 | skoXC1 ^ 4 = 265/128 ^ 4),
    inference(resolve, [], [refute_0_33, refute_0_34])).

cnf(refute_0_36, plain, (skoXC1 ^ 4 = 265/128 ^ 4),
    inference(resolve, [], [refute_0_20, refute_0_35])).

cnf(refute_0_37, plain, (14325 * skoXC1 ^ 4 = 14325 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_38, plain,
    (14325 * skoXC1 ^ 4 != 14325 * skoXC1 ^ 4 | skoXC1 ^ 4 != 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_39, plain,
    (skoXC1 ^ 4 != 265/128 ^ 4 | 14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain, (14325 * skoXC1 ^ 4 = 14325 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_36, refute_0_39])).

cnf(refute_0_41, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_42, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_43, plain,
    (14325 * skoXC1 ^ 4 != 14325 * 265/128 ^ 4 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_41, refute_0_42])).

cnf(refute_0_44, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_40, refute_0_43])).

cnf(refute_0_45, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_46, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_47, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_49, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (14325 * 265/128 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_44, refute_0_50])).

cnf(refute_0_52, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_32, refute_0_51])).

cnf(refute_0_53, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660),
    introduced(tautology, [refl])).

cnf(refute_0_54, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    introduced(tautology, [equality])).

cnf(refute_0_55, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 |
     14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 =
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660),
    inference(resolve, [], [refute_0_52, refute_0_55])).

cnf(refute_0_57, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_58, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660 |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_59, plain,
    (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660 !=
     14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660 |
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_56, refute_0_59])).

cnf(refute_0_61, plain, (169932 * skoXC1 = 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_62, plain,
    (169932 * skoXC1 != 169932 * skoXC1 | skoXC1 != 265/128 |
     169932 * skoXC1 = 169932 * 265/128),
    introduced(tautology, [equality])).

cnf(refute_0_63, plain,
    (skoXC1 != 265/128 | 169932 * skoXC1 = 169932 * 265/128),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain, (169932 * skoXC1 = 169932 * 265/128),
    inference(resolve, [], [refute_0_20, refute_0_63])).

cnf(refute_0_65, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_66, plain,
    (169932 * skoXC1 != 169932 * 265/128 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    introduced(tautology, [equality])).

cnf(refute_0_67, plain,
    (169932 * skoXC1 != 169932 * 265/128 |
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_64, refute_0_67])).

cnf(refute_0_69, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_70, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 265/128 |
     skoXC1 ^ 5 = 265/128 ^ 5), introduced(tautology, [equality])).

cnf(refute_0_71, plain, (skoXC1 != 265/128 | skoXC1 ^ 5 = 265/128 ^ 5),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain, (skoXC1 ^ 5 = 265/128 ^ 5),
    inference(resolve, [], [refute_0_20, refute_0_71])).

cnf(refute_0_73, plain, (3890 * skoXC1 ^ 5 = 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_74, plain,
    (3890 * skoXC1 ^ 5 != 3890 * skoXC1 ^ 5 | skoXC1 ^ 5 != 265/128 ^ 5 |
     3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_75, plain,
    (skoXC1 ^ 5 != 265/128 ^ 5 | 3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain, (3890 * skoXC1 ^ 5 = 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_72, refute_0_75])).

cnf(refute_0_77, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_78, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 265/128 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_79, plain,
    (3890 * skoXC1 ^ 5 != 3890 * 265/128 ^ 5 |
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_76, refute_0_79])).

cnf(refute_0_81, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_82, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 265/128 |
     skoXC1 ^ 3 = 265/128 ^ 3), introduced(tautology, [equality])).

cnf(refute_0_83, plain, (skoXC1 != 265/128 | skoXC1 ^ 3 = 265/128 ^ 3),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain, (skoXC1 ^ 3 = 265/128 ^ 3),
    inference(resolve, [], [refute_0_20, refute_0_83])).

cnf(refute_0_85, plain, (105247 * skoXC1 ^ 3 = 105247 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (105247 * skoXC1 ^ 3 != 105247 * skoXC1 ^ 3 |
     skoXC1 ^ 3 != 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (skoXC1 ^ 3 != 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain, (105247 * skoXC1 ^ 3 = 105247 * 265/128 ^ 3),
    inference(resolve, [], [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (105247 * skoXC1 ^ 3 != 105247 * 265/128 ^ 3 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_94, plain,
    (105247 * 265/128 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5),
    inference(resolve, [], [refute_0_80, refute_0_94])).

cnf(refute_0_96, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_97, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_98, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_96, refute_0_97])).

cnf(refute_0_99, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1),
    inference(resolve, [], [refute_0_95, refute_0_98])).

cnf(refute_0_100, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_101, plain,
    (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 =
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128),
    inference(resolve, [], [refute_0_68, refute_0_101])).

cnf(refute_0_103, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1 !=
     105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128 |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_102, refute_0_105])).

cnf(refute_0_107, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_108, plain,
    ((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     (105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) =
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)),
    inference(resolve, [], [refute_0_60, refute_0_108])).

cnf(refute_0_110, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_111, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_112, plain,
    ((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
     (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660) !=
     (105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 - 169932 * 265/128) /
     (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_109, refute_0_112])).

cnf(refute_0_114, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))),
    introduced(tautology, [refl])).

cnf(refute_0_115, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) |
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    introduced(tautology, [equality])).

cnf(refute_0_116, plain,
    (arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) =
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))),
    inference(resolve, [], [refute_0_113, refute_0_116])).

cnf(refute_0_118, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_119, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_120, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_117, refute_0_120])).

cnf(refute_0_122, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_123, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_124, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     -12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 =
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18),
    inference(resolve, [], [refute_0_121, refute_0_124])).

cnf(refute_0_126, plain,
    (-12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) - 18 !=
     -12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) - 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_127, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                         169932 * skoXC1) /
                        (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)))),
    introduced(tautology, [refl])).

cnf(refute_0_129, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) |
     -pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    introduced(tautology, [equality])).

cnf(refute_0_130, plain,
    (-pi -
     arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
            (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660)) !=
     -pi -
     arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
             169932 * 265/128) /
            (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660)))),
    inference(resolve, [], [refute_0_117, refute_0_130])).

cnf(refute_0_132, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18),
    introduced(tautology, [refl])).

cnf(refute_0_133, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    introduced(tautology, [equality])).

cnf(refute_0_134, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) |
     12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_132, refute_0_133])).

cnf(refute_0_135, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 =
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18),
    inference(resolve, [], [refute_0_131, refute_0_134])).

cnf(refute_0_136, plain,
    (12 / pi *
     (-pi -
      arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 - 169932 * skoXC1) /
             (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 - 16660))) + 18 !=
     12 / pi *
     (-pi -
      arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
              169932 * 265/128) /
             (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 - 16660))) + 18 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                    3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 265/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (skoXC1 ^ 2 != 265/128 ^ 2 |
     3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (3120782021 * skoXC1 ^ 2 = 3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_24, refute_0_140])).

cnf(refute_0_142, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_143, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 265/128 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_144, plain,
    (3120782021 * skoXC1 ^ 2 != 3120782021 * 265/128 ^ 2 |
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_141, refute_0_144])).

cnf(refute_0_146, plain, (217230815 * skoXC1 ^ 4 = 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_147, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 265/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_148, plain,
    (skoXC1 ^ 4 != 265/128 ^ 4 |
     217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain,
    (217230815 * skoXC1 ^ 4 = 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_36, refute_0_148])).

cnf(refute_0_150, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_151, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 265/128 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_152, plain,
    (217230815 * skoXC1 ^ 4 != 217230815 * 265/128 ^ 4 |
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_150, refute_0_151])).

cnf(refute_0_153, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_149, refute_0_152])).

cnf(refute_0_154, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_155, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 265/128 |
     skoXC1 ^ 6 = 265/128 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_156, plain, (skoXC1 != 265/128 | skoXC1 ^ 6 = 265/128 ^ 6),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain, (skoXC1 ^ 6 = 265/128 ^ 6),
    inference(resolve, [], [refute_0_20, refute_0_156])).

cnf(refute_0_158, plain, (3783025 * skoXC1 ^ 6 = 3783025 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (skoXC1 ^ 6 != 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain, (3783025 * skoXC1 ^ 6 = 3783025 * 265/128 ^ 6),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_163, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_164, plain,
    (3783025 * skoXC1 ^ 6 != 3783025 * 265/128 ^ 6 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_161, refute_0_164])).

cnf(refute_0_166, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_167, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4),
    inference(resolve, [], [refute_0_153, refute_0_167])).

cnf(refute_0_169, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_170, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_171, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_168, refute_0_171])).

cnf(refute_0_173, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2), inference(subst, [], [refute_0_49])).

cnf(refute_0_174, plain,
    (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2),
    inference(resolve, [], [refute_0_145, refute_0_174])).

cnf(refute_0_176, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900), introduced(tautology, [refl])).

cnf(refute_0_177, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    introduced(tautology, [equality])).

cnf(refute_0_178, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 |
     3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_176, refute_0_177])).

cnf(refute_0_179, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 =
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900),
    inference(resolve, [], [refute_0_175, refute_0_178])).

cnf(refute_0_180, plain,
    (4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900)), introduced(tautology, [refl])).

cnf(refute_0_181, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900 |
     4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) |
     4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 * skoYC1 /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    introduced(tautology, [equality])).

cnf(refute_0_182, plain,
    (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
     3120782021 * skoXC1 ^ 2 + 69388900 !=
     3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
     3120782021 * 265/128 ^ 2 + 69388900 |
     4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 * skoYC1 /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_180, refute_0_181])).

cnf(refute_0_183, plain,
    (4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) =
     4250 * skoYC1 /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900)),
    inference(resolve, [], [refute_0_179, refute_0_182])).

cnf(refute_0_184, plain,
    (ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_185, plain,
    (4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 * skoYC1 /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_186, plain,
    (4250 * skoYC1 /
     (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
      3120782021 * skoXC1 ^ 2 + 69388900) !=
     4250 * skoYC1 /
     (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
      3120782021 * 265/128 ^ 2 + 69388900) |
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_184, refute_0_185])).

cnf(refute_0_187, plain,
    (ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_183, refute_0_186])).

cnf(refute_0_188, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900))),
    introduced(tautology, [refl])).

cnf(refute_0_189, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) |
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    introduced(tautology, [equality])).

cnf(refute_0_190, plain,
    (ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_188, refute_0_189])).

cnf(refute_0_191, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900))),
    inference(resolve, [], [refute_0_187, refute_0_190])).

cnf(refute_0_192, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2),
    introduced(tautology, [refl])).

cnf(refute_0_193, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 !=
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 |
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    introduced(tautology, [equality])).

cnf(refute_0_194, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) !=
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) |
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    inference(resolve, [], [refute_0_192, refute_0_193])).

cnf(refute_0_195, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 =
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2),
    inference(resolve, [], [refute_0_191, refute_0_194])).

cnf(refute_0_196, plain,
    (20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
         3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2 !=
     20 / ln(10) *
     ln(4250 * skoYC1 /
        (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
         3120782021 * 265/128 ^ 2 + 69388900)) + 3/2 |
     ~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    introduced(tautology, [equality])).

cnf(refute_0_197, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_195, refute_0_196])).

cnf(refute_0_198, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                       169932 * 265/128) /
                      (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_197, refute_0_137])).

cnf(refute_0_199, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2) |
     interval(0,
              -12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) - 18, 0,
              12 / pi *
              (-pi -
               arctan((105247 * skoXC1 ^ 3 + 3890 * skoXC1 ^ 5 -
                       169932 * skoXC1) /
                      (14325 * skoXC1 ^ 4 + 406627 * skoXC1 ^ 2 -
                       16660))) + 18,
              20 / ln(10) *
              ln(4250 * skoYC1 /
                 (3783025 * skoXC1 ^ 6 + 217230815 * skoXC1 ^ 4 +
                  3120782021 * skoXC1 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_198, refute_0_127])).

cnf(refute_0_200, plain,
    (~ interval(0,
                -12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) - 18, 0,
                12 / pi *
                (-pi -
                 arctan((105247 * 265/128 ^ 3 + 3890 * 265/128 ^ 5 -
                         169932 * 265/128) /
                        (14325 * 265/128 ^ 4 + 406627 * 265/128 ^ 2 -
                         16660))) + 18,
                20 / ln(10) *
                ln(4250 * skoYC1 /
                   (3783025 * 265/128 ^ 6 + 217230815 * 265/128 ^ 4 +
                    3120782021 * 265/128 ^ 2 + 69388900)) + 3/2)),
    inference(resolve, [], [refute_0_199, refute_0_19])).

cnf(refute_0_201, plain,
    (~ interval(0,
                (pi * -6 + arctan(836177642653679/2278501083899584) * 12) /
                pi, 0,
                (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
                pi,
                (ln(10) * 3/2 +
                 ln(skoYC1 * 747667906887680/3119868895908289175433) *
                 20) / ln(10)) | pi = 0),
    inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (~ lgen(0,
            (ln(10) * 3/2 +
             ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) /
            ln(10),
            (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi,
            (ln(10) * 3/2 +
             ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) /
            ln(10)) |
     interval(0,
              (pi * -6 + arctan(836177642653679/2278501083899584) * 12) /
              pi, 0,
              (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
              pi,
              (ln(10) * 3/2 +
               ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) /
              ln(10))), inference(subst, [], [interval_intro])).

cnf(refute_0_203, plain,
    (~ lgen(0,
            (ln(10) * 3/2 +
             ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) /
            ln(10),
            (pi * 6 + arctan(836177642653679/2278501083899584) * -12) /
            pi) |
     ~ lgen(0,
            (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi,
            (ln(10) * 3/2 +
             ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) /
            ln(10)) | pi = 0),
    inference(resolve, [], [refute_0_202, refute_0_201])).

cnf(refute_0_204, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) <
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi |
     pi = 0), inference(arithmetic, [], [refute_0_203])).

cnf(refute_0_205, plain,
    (skoYC1 ^ 2 =
     15132100 * skoXC1 ^ 10 + 1024027285 * skoXC1 ^ 8 +
     21404723599 * skoXC1 ^ 6 + 129098541721 * skoXC1 ^ 4 +
     15328072984 * skoXC1 ^ 2 + 277555600),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_206, plain,
    (skoYC1 * skoYC1 =
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain, (skoXC1 * 15132100 = skoXC1 * 15132100),
    introduced(tautology, [refl])).

cnf(refute_0_208, plain,
    (skoXC1 * 15132100 != skoXC1 * 15132100 | skoXC1 != 265/128 |
     skoXC1 * 15132100 = 265/128 * 15132100),
    introduced(tautology, [equality])).

cnf(refute_0_209, plain,
    (skoXC1 != 265/128 | skoXC1 * 15132100 = 265/128 * 15132100),
    inference(resolve, [], [refute_0_207, refute_0_208])).

cnf(refute_0_210, plain, (skoXC1 * 15132100 = 265/128 * 15132100),
    inference(resolve, [], [refute_0_20, refute_0_209])).

cnf(refute_0_211, plain,
    (265/128 * (skoXC1 * 15132100) = 265/128 * (skoXC1 * 15132100)),
    introduced(tautology, [refl])).

cnf(refute_0_212, plain,
    (265/128 * (skoXC1 * 15132100) != 265/128 * (skoXC1 * 15132100) |
     skoXC1 * 15132100 != 265/128 * 15132100 |
     265/128 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    introduced(tautology, [equality])).

cnf(refute_0_213, plain,
    (skoXC1 * 15132100 != 265/128 * 15132100 |
     265/128 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_211, refute_0_212])).

cnf(refute_0_214, plain,
    (265/128 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_210, refute_0_213])).

cnf(refute_0_215, plain,
    (skoXC1 * (skoXC1 * 15132100) = skoXC1 * (skoXC1 * 15132100)),
    introduced(tautology, [refl])).

cnf(refute_0_216, plain,
    (skoXC1 * (skoXC1 * 15132100) != skoXC1 * (skoXC1 * 15132100) |
     skoXC1 != 265/128 |
     skoXC1 * (skoXC1 * 15132100) = 265/128 * (skoXC1 * 15132100)),
    introduced(tautology, [equality])).

cnf(refute_0_217, plain,
    (skoXC1 != 265/128 |
     skoXC1 * (skoXC1 * 15132100) = 265/128 * (skoXC1 * 15132100)),
    inference(resolve, [], [refute_0_215, refute_0_216])).

cnf(refute_0_218, plain,
    (skoXC1 * (skoXC1 * 15132100) = 265/128 * (skoXC1 * 15132100)),
    inference(resolve, [], [refute_0_20, refute_0_217])).

cnf(refute_0_219, plain,
    (265/128 * (skoXC1 * 15132100) != 265/128 * (265/128 * 15132100) |
     skoXC1 * (skoXC1 * 15132100) != 265/128 * (skoXC1 * 15132100) |
     skoXC1 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_220, plain,
    (265/128 * (skoXC1 * 15132100) != 265/128 * (265/128 * 15132100) |
     skoXC1 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_218, refute_0_219])).

cnf(refute_0_221, plain,
    (skoXC1 * (skoXC1 * 15132100) = 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_214, refute_0_220])).

cnf(refute_0_222, plain,
    (1024027285 + skoXC1 * (skoXC1 * 15132100) =
     1024027285 + skoXC1 * (skoXC1 * 15132100)),
    introduced(tautology, [refl])).

cnf(refute_0_223, plain,
    (skoXC1 * (skoXC1 * 15132100) != 265/128 * (265/128 * 15132100) |
     1024027285 + skoXC1 * (skoXC1 * 15132100) !=
     1024027285 + skoXC1 * (skoXC1 * 15132100) |
     1024027285 + skoXC1 * (skoXC1 * 15132100) =
     1024027285 + 265/128 * (265/128 * 15132100)),
    introduced(tautology, [equality])).

cnf(refute_0_224, plain,
    (skoXC1 * (skoXC1 * 15132100) != 265/128 * (265/128 * 15132100) |
     1024027285 + skoXC1 * (skoXC1 * 15132100) =
     1024027285 + 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    (1024027285 + skoXC1 * (skoXC1 * 15132100) =
     1024027285 + 265/128 * (265/128 * 15132100)),
    inference(resolve, [], [refute_0_221, refute_0_224])).

cnf(refute_0_226, plain,
    (265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100))),
    introduced(tautology, [refl])).

cnf(refute_0_227, plain,
    (265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) |
     1024027285 + skoXC1 * (skoXC1 * 15132100) !=
     1024027285 + 265/128 * (265/128 * 15132100) |
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    introduced(tautology, [equality])).

cnf(refute_0_228, plain,
    (1024027285 + skoXC1 * (skoXC1 * 15132100) !=
     1024027285 + 265/128 * (265/128 * 15132100) |
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain,
    (265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    inference(resolve, [], [refute_0_225, refute_0_228])).

cnf(refute_0_230, plain,
    (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) |
     skoXC1 != 265/128 |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100))),
    introduced(tautology, [equality])).

cnf(refute_0_232, plain,
    (skoXC1 != 265/128 |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100))),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100))),
    inference(resolve, [], [refute_0_20, refute_0_232])).

cnf(refute_0_234, plain,
    (265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100)) |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_235, plain,
    (265/128 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100)) |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    inference(resolve, [], [refute_0_233, refute_0_234])).

cnf(refute_0_236, plain,
    (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) =
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100))),
    inference(resolve, [], [refute_0_229, refute_0_235])).

cnf(refute_0_237, plain,
    (265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    introduced(tautology, [refl])).

cnf(refute_0_238, plain,
    (265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) |
     skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100)) |
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    introduced(tautology, [equality])).

cnf(refute_0_239, plain,
    (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)) !=
     265/128 * (1024027285 + 265/128 * (265/128 * 15132100)) |
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_236, refute_0_239])).

cnf(refute_0_241, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    introduced(tautology, [refl])).

cnf(refute_0_242, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) |
     skoXC1 != 265/128 |
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    introduced(tautology, [equality])).

cnf(refute_0_243, plain,
    (skoXC1 != 265/128 |
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    inference(resolve, [], [refute_0_241, refute_0_242])).

cnf(refute_0_244, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    inference(resolve, [], [refute_0_20, refute_0_243])).

cnf(refute_0_245, plain,
    (265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) |
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_246, plain,
    (265/128 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_240, refute_0_246])).

cnf(refute_0_248, plain,
    (21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))),
    introduced(tautology, [refl])).

cnf(refute_0_249, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) |
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     21404723599 +
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    introduced(tautology, [equality])).

cnf(refute_0_250, plain,
    (skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     21404723599 +
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    (21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) =
     21404723599 +
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))),
    inference(resolve, [], [refute_0_247, refute_0_250])).

cnf(refute_0_252, plain,
    (265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))),
    introduced(tautology, [refl])).

cnf(refute_0_253, plain,
    (265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) |
     21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     21404723599 +
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    introduced(tautology, [equality])).

cnf(refute_0_254, plain,
    (21404723599 +
     skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))) !=
     21404723599 +
     265/128 * (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))) |
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    inference(resolve, [], [refute_0_252, refute_0_253])).

cnf(refute_0_255, plain,
    (265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    inference(resolve, [], [refute_0_251, refute_0_254])).

cnf(refute_0_256, plain,
    (skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))),
    introduced(tautology, [refl])).

cnf(refute_0_257, plain,
    (skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))),
    introduced(tautology, [equality])).

cnf(refute_0_258, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))),
    inference(resolve, [], [refute_0_256, refute_0_257])).

cnf(refute_0_259, plain,
    (skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))),
    inference(resolve, [], [refute_0_20, refute_0_258])).

cnf(refute_0_260, plain,
    (265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))) |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_261, plain,
    (265/128 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))) |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    inference(resolve, [], [refute_0_259, refute_0_260])).

cnf(refute_0_262, plain,
    (skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) =
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))),
    inference(resolve, [], [refute_0_255, refute_0_261])).

cnf(refute_0_263, plain,
    (265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    introduced(tautology, [refl])).

cnf(refute_0_264, plain,
    (265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) |
     skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))) |
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    introduced(tautology, [equality])).

cnf(refute_0_265, plain,
    (skoXC1 *
     (21404723599 +
      skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))) !=
     265/128 *
     (21404723599 +
      265/128 *
      (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))) |
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_263, refute_0_264])).

cnf(refute_0_266, plain,
    (265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_262, refute_0_265])).

cnf(refute_0_267, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    introduced(tautology, [refl])).

cnf(refute_0_268, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    introduced(tautology, [equality])).

cnf(refute_0_269, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    inference(resolve, [], [refute_0_267, refute_0_268])).

cnf(refute_0_270, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    inference(resolve, [], [refute_0_20, refute_0_269])).

cnf(refute_0_271, plain,
    (265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) |
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_272, plain,
    (265/128 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_270, refute_0_271])).

cnf(refute_0_273, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_266, refute_0_272])).

cnf(refute_0_274, plain,
    (129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))),
    introduced(tautology, [refl])).

cnf(refute_0_275, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) |
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     129098541721 +
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    introduced(tautology, [equality])).

cnf(refute_0_276, plain,
    (skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     129098541721 +
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_274, refute_0_275])).

cnf(refute_0_277, plain,
    (129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) =
     129098541721 +
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))),
    inference(resolve, [], [refute_0_273, refute_0_276])).

cnf(refute_0_278, plain,
    (265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))),
    introduced(tautology, [refl])).

cnf(refute_0_279, plain,
    (265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) |
     129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     129098541721 +
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    introduced(tautology, [equality])).

cnf(refute_0_280, plain,
    (129098541721 +
     skoXC1 *
     (skoXC1 *
      (21404723599 +
       skoXC1 * (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))) !=
     129098541721 +
     265/128 *
     (265/128 *
      (21404723599 +
       265/128 *
       (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))) |
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    inference(resolve, [], [refute_0_278, refute_0_279])).

cnf(refute_0_281, plain,
    (265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    inference(resolve, [], [refute_0_277, refute_0_280])).

cnf(refute_0_282, plain,
    (skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))),
    introduced(tautology, [refl])).

cnf(refute_0_283, plain,
    (skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))),
    introduced(tautology, [equality])).

cnf(refute_0_284, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))),
    inference(resolve, [], [refute_0_282, refute_0_283])).

cnf(refute_0_285, plain,
    (skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))),
    inference(resolve, [], [refute_0_20, refute_0_284])).

cnf(refute_0_286, plain,
    (265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))) |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_287, plain,
    (265/128 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))) |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    inference(resolve, [], [refute_0_285, refute_0_286])).

cnf(refute_0_288, plain,
    (skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) =
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))),
    inference(resolve, [], [refute_0_281, refute_0_287])).

cnf(refute_0_289, plain,
    (265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    introduced(tautology, [refl])).

cnf(refute_0_290, plain,
    (265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) |
     skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))) |
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    introduced(tautology, [equality])).

cnf(refute_0_291, plain,
    (skoXC1 *
     (129098541721 +
      skoXC1 *
      (skoXC1 *
       (21404723599 +
        skoXC1 *
        (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))) !=
     265/128 *
     (129098541721 +
      265/128 *
      (265/128 *
       (21404723599 +
        265/128 *
        (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))) |
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_289, refute_0_290])).

cnf(refute_0_292, plain,
    (265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_288, refute_0_291])).

cnf(refute_0_293, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    introduced(tautology, [refl])).

cnf(refute_0_294, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    introduced(tautology, [equality])).

cnf(refute_0_295, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    inference(resolve, [], [refute_0_293, refute_0_294])).

cnf(refute_0_296, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    inference(resolve, [], [refute_0_20, refute_0_295])).

cnf(refute_0_297, plain,
    (265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) |
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_298, plain,
    (265/128 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_296, refute_0_297])).

cnf(refute_0_299, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_292, refute_0_298])).

cnf(refute_0_300, plain,
    (15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))),
    introduced(tautology, [refl])).

cnf(refute_0_301, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) |
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     15328072984 +
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    introduced(tautology, [equality])).

cnf(refute_0_302, plain,
    (skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     15328072984 +
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_300, refute_0_301])).

cnf(refute_0_303, plain,
    (15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) =
     15328072984 +
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))),
    inference(resolve, [], [refute_0_299, refute_0_302])).

cnf(refute_0_304, plain,
    (265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))),
    introduced(tautology, [refl])).

cnf(refute_0_305, plain,
    (265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) |
     15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     15328072984 +
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_306, plain,
    (15328072984 +
     skoXC1 *
     (skoXC1 *
      (129098541721 +
       skoXC1 *
       (skoXC1 *
        (21404723599 +
         skoXC1 *
         (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))) !=
     15328072984 +
     265/128 *
     (265/128 *
      (129098541721 +
       265/128 *
       (265/128 *
        (21404723599 +
         265/128 *
         (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))) |
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    inference(resolve, [], [refute_0_304, refute_0_305])).

cnf(refute_0_307, plain,
    (265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    inference(resolve, [], [refute_0_303, refute_0_306])).

cnf(refute_0_308, plain,
    (skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))),
    introduced(tautology, [refl])).

cnf(refute_0_309, plain,
    (skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_310, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))),
    inference(resolve, [], [refute_0_308, refute_0_309])).

cnf(refute_0_311, plain,
    (skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))),
    inference(resolve, [], [refute_0_20, refute_0_310])).

cnf(refute_0_312, plain,
    (265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))) |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_313, plain,
    (265/128 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))) |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    inference(resolve, [], [refute_0_311, refute_0_312])).

cnf(refute_0_314, plain,
    (skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) =
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100))))))))),
    inference(resolve, [], [refute_0_307, refute_0_313])).

cnf(refute_0_315, plain,
    (265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    introduced(tautology, [refl])).

cnf(refute_0_316, plain,
    (265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))) |
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_317, plain,
    (skoXC1 *
     (15328072984 +
      skoXC1 *
      (skoXC1 *
       (129098541721 +
        skoXC1 *
        (skoXC1 *
         (21404723599 +
          skoXC1 *
          (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))) !=
     265/128 *
     (15328072984 +
      265/128 *
      (265/128 *
       (129098541721 +
        265/128 *
        (265/128 *
         (21404723599 +
          265/128 *
          (265/128 * (1024027285 + 265/128 * (265/128 * 15132100)))))))) |
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_315, refute_0_316])).

cnf(refute_0_318, plain,
    (265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_314, refute_0_317])).

cnf(refute_0_319, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    introduced(tautology, [refl])).

cnf(refute_0_320, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_321, plain,
    (skoXC1 != 265/128 |
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    inference(resolve, [], [refute_0_319, refute_0_320])).

cnf(refute_0_322, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    inference(resolve, [], [refute_0_20, refute_0_321])).

cnf(refute_0_323, plain,
    (265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100))))))))) |
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(subst, [], [refute_0_49])).

cnf(refute_0_324, plain,
    (265/128 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100))))))))) |
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_322, refute_0_323])).

cnf(refute_0_325, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_318, refute_0_324])).

cnf(refute_0_326, plain,
    (277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100)))))))))),
    introduced(tautology, [refl])).

cnf(refute_0_327, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100))))))))) |
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_328, plain,
    (skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100))))))))) |
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_326, refute_0_327])).

cnf(refute_0_329, plain,
    (277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_325, refute_0_328])).

cnf(refute_0_330, plain,
    (skoYC1 * skoYC1 !=
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) !=
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100))))))))) |
     skoYC1 * skoYC1 =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_331, plain,
    (skoYC1 * skoYC1 !=
     277555600 +
     skoXC1 *
     (skoXC1 *
      (15328072984 +
       skoXC1 *
       (skoXC1 *
        (129098541721 +
         skoXC1 *
         (skoXC1 *
          (21404723599 +
           skoXC1 *
           (skoXC1 * (1024027285 + skoXC1 * (skoXC1 * 15132100))))))))) |
     skoYC1 * skoYC1 =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_329, refute_0_330])).

cnf(refute_0_332, plain,
    (skoYC1 * skoYC1 =
     277555600 +
     265/128 *
     (265/128 *
      (15328072984 +
       265/128 *
       (265/128 *
        (129098541721 +
         265/128 *
         (265/128 *
          (21404723599 +
           265/128 *
           (265/128 *
            (1024027285 + 265/128 * (265/128 * 15132100)))))))))),
    inference(resolve, [], [refute_0_206, refute_0_331])).

cnf(refute_0_333, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_334, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_335, plain, (0 <= skoYC1),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_336, plain, (pi != 0),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_337, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) <
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi),
    inference(resolve, [], [refute_0_204, refute_0_336])).

cnf(refute_0_338, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (pi *
       30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) <
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 *
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 <=
     (pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_339, plain,
    (~ lgen(0, arctan(X_000131), X_000132) | arctan(X_000131) <= X_000132),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_340, plain,
    (X_000131 < 0 |
     ~ lgen(0,
            1/15 * (64 * X_000131 ^ 4 + 735 * X_000131 ^ 2 + 945) *
            X_000131 / (15 * X_000131 ^ 4 + 70 * X_000131 ^ 2 + 63),
            X_000132) | lgen(0, arctan(X_000131), X_000132)),
    inference(subst, [], [atan_upper_bound_case_33])).

cnf(refute_0_341, plain,
    (X_000131 < 0 |
     ~ lgen(0,
            1/15 * (64 * X_000131 ^ 4 + 735 * X_000131 ^ 2 + 945) *
            X_000131 / (15 * X_000131 ^ 4 + 70 * X_000131 ^ 2 + 63),
            X_000132) | arctan(X_000131) <= X_000132),
    inference(resolve, [], [refute_0_340, refute_0_339])).

cnf(refute_0_342, plain,
    (X_000131 < 0 |
     X_000132 <
     X_000131 *
     (63 + X_000131 * (X_000131 * (49 + X_000131 * (X_000131 * 64/15)))) /
     (63 + X_000131 * (X_000131 * (70 + X_000131 * (X_000131 * 15)))) |
     arctan(X_000131) <= X_000132),
    inference(arithmetic, [], [refute_0_341])).

cnf(refute_0_343, plain,
    (836177642653679/2278501083899584 < 0 |
     (pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <
     836177642653679/2278501083899584 *
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (49 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 64/15)))) /
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (70 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 15)))) |
     arctan(836177642653679/2278501083899584) <=
     (pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815))),
    inference(subst, [], [refute_0_342])).

cnf(refute_0_344, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -29906716275507200/3119868895908289175433 +
      skoYC1 *
      (pi *
       27950364949490226519206789120000/
       9733581927656007318022218939278759050737489)) <
     231993/101125 *
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     0 <=
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) +
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (5981343255101440/1039956298636096391811 +
       skoYC1 *
       2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) <= 231993/101125),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_345, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) | 10 <= 0 |
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) <= ln(10)),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_346, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1/2 +
      skoYC1 *
      (-1495335813775360/3119868895908289175433 +
       skoYC1 *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) * (pi * -20) |
     0 <= pi * -20 |
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) <=
     (-1/2 +
      skoYC1 *
      (-1495335813775360/3119868895908289175433 +
       skoYC1 *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_347, plain,
    (~ lgen(0, X_000111, ln(X_000110)) | X_000111 <= ln(X_000110)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_348, plain,
    (~ lgen(0, X_000111,
            1/2 * (1 + 5 * X_000110) * (X_000110 - 1) /
            (X_000110 * (2 + X_000110))) | X_000110 <= 0 |
     lgen(0, X_000111, ln(X_000110))),
    inference(subst, [], [ln_lower_bound_cf3])).

cnf(refute_0_349, plain,
    (~ lgen(0, X_000111,
            1/2 * (1 + 5 * X_000110) * (X_000110 - 1) /
            (X_000110 * (2 + X_000110))) | X_000110 <= 0 |
     X_000111 <= ln(X_000110)),
    inference(resolve, [], [refute_0_348, refute_0_347])).

cnf(refute_0_350, plain,
    ((-1/2 + X_000110 * (-2 + X_000110 * 5/2)) /
     (X_000110 * (2 + X_000110)) < X_000111 | X_000110 <= 0 |
     X_000111 <= ln(X_000110)), inference(arithmetic, [], [refute_0_349])).

cnf(refute_0_351, plain,
    ((-1/2 +
      skoYC1 * 747667906887680/3119868895908289175433 *
      (-2 + skoYC1 * 747667906887680/3119868895908289175433 * 5/2)) /
     (skoYC1 * 747667906887680/3119868895908289175433 *
      (2 + skoYC1 * 747667906887680/3119868895908289175433)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) | skoYC1 * 747667906887680/3119868895908289175433 <= 0 |
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) <= ln(skoYC1 * 747667906887680/3119868895908289175433)),
    inference(subst, [], [refute_0_350])).

cnf(refute_0_352, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * -3/2) +
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20) |
     (ln(10) * (pi * -3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20)) /
     ln(10) <= pi * 6 + arctan(836177642653679/2278501083899584) * -12 |
     ln(10) <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_353, plain,
    ((ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) *
     pi < pi * -6 + arctan(836177642653679/2278501083899584) * 12 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     pi <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_354, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * -3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20)) /
     ln(10) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     pi <= 0), inference(arithmetic, [], [refute_0_353])).

cnf(refute_0_355, plain, (0 < pi),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_356, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * -3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20)) /
     ln(10) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_354, refute_0_355])).

cnf(refute_0_357, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * -3/2) +
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     ln(10) <= 0), inference(resolve, [], [refute_0_352, refute_0_356])).

cnf(refute_0_358, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_357])).

cnf(refute_0_359, plain,
    (ln(skoYC1 * 747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (ln(skoYC1 * 747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     pi <= 0), inference(arithmetic, [], [refute_0_359])).

cnf(refute_0_361, plain,
    (ln(skoYC1 * 747667906887680/3119868895908289175433) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_360, refute_0_355])).

cnf(refute_0_362, plain,
    ((-1/2 +
      skoYC1 * 747667906887680/3119868895908289175433 *
      (-2 + skoYC1 * 747667906887680/3119868895908289175433 * 5/2)) /
     (skoYC1 * 747667906887680/3119868895908289175433 *
      (2 + skoYC1 * 747667906887680/3119868895908289175433)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) | skoYC1 * 747667906887680/3119868895908289175433 <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_351, refute_0_361])).

cnf(refute_0_363, plain,
    ((-1/2 +
      skoYC1 *
      (-1495335813775360/3119868895908289175433 +
       skoYC1 *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     skoYC1 <= 0), inference(arithmetic, [], [refute_0_362])).

cnf(refute_0_364, plain, (0 < skoYC1),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_365, plain,
    ((-1/2 +
      skoYC1 *
      (-1495335813775360/3119868895908289175433 +
       skoYC1 *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) <
     ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * -20) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_363, refute_0_364])).

cnf(refute_0_366, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1/2 +
      skoYC1 *
      (-1495335813775360/3119868895908289175433 +
       skoYC1 *
       1397518247474511325960339456000/
       9733581927656007318022218939278759050737489)) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) * (pi * -20) |
     0 <= pi * -20 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_346, refute_0_365])).

cnf(refute_0_367, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (pi * 10 +
      skoYC1 *
      (pi * 29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        -27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) |
     pi <= 0), inference(arithmetic, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (ln(10) *
     (pi * 15/2 + arctan(836177642653679/2278501083899584) * -12) <
     (pi * 10 +
      skoYC1 *
      (pi * 29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        -27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (1495335813775360/3119868895908289175433 +
       skoYC1 *
       559007298989804530384135782400/
       9733581927656007318022218939278759050737489)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_367, refute_0_355])).

cnf(refute_0_369, plain,
    (ln(10) <
     -((pi * 10 +
       skoYC1 *
       (pi * 29906716275507200/3119868895908289175433 +
        skoYC1 *
        (pi *
         -27950364949490226519206789120000/
         9733581927656007318022218939278759050737489))) /
      (skoYC1 *
       (1495335813775360/3119868895908289175433 +
        skoYC1 *
        559007298989804530384135782400/
        9733581927656007318022218939278759050737489))) /
     (pi * -15/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -15/2 + arctan(836177642653679/2278501083899584) * 12 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_368])).

cnf(refute_0_370, plain,
    (ln(10) <
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     pi * 5/8 <= arctan(836177642653679/2278501083899584) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (arctan(X_000040) < X_000041 | ~ lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_372, plain,
    (X_000040 < 0 | ~ lgen(1, X_000040, X_000041) |
     lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_373, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 |
     ~ lgen(1, X_000040, X_000041)),
    inference(resolve, [], [refute_0_372, refute_0_371])).

cnf(refute_0_374, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 | X_000041 <= X_000040),
    inference(arithmetic, [], [refute_0_373])).

cnf(refute_0_375, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) < pi * 5/8 |
     pi * 5/8 <= 836177642653679/2278501083899584),
    inference(subst, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (arctan(836177642653679/2278501083899584) < pi * 5/8 |
     pi <= 836177642653679/1424063177437240),
    inference(arithmetic, [], [refute_0_375])).

cnf(refute_0_377, plain, (836177642653679/1424063177437240 < pi),
    inference(decision, [],
              [refute_0_335, refute_0_334, refute_0_333, refute_0_20,
               refute_0_206, refute_0_332])).

cnf(refute_0_378, plain,
    (arctan(836177642653679/2278501083899584) < pi * 5/8),
    inference(resolve, [], [refute_0_376, refute_0_377])).

cnf(refute_0_379, plain,
    (ln(10) <
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_370, refute_0_378])).

cnf(refute_0_380, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) | 10 <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_345, refute_0_379])).

cnf(refute_0_381, plain,
    (231993/101125 <
     (pi * -10 +
      skoYC1 *
      (pi * -29906716275507200/3119868895908289175433 +
       skoYC1 *
       (pi *
        27950364949490226519206789120000/
        9733581927656007318022218939278759050737489))) /
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_380])).

cnf(refute_0_382, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -29906716275507200/3119868895908289175433 +
      skoYC1 *
      (pi *
       27950364949490226519206789120000/
       9733581927656007318022218939278759050737489)) <
     231993/101125 *
     (skoYC1 *
      (pi * -3738339534438400/1039956298636096391811 +
       skoYC1 *
       (pi *
        -1397518247474511325960339456000/
        3244527309218669106007406313092919683579163)) +
      arctan(836177642653679/2278501083899584) *
      (skoYC1 *
       (5981343255101440/1039956298636096391811 +
        skoYC1 *
        2236029195959218121536543129600/
        3244527309218669106007406313092919683579163))) |
     0 <=
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) +
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (5981343255101440/1039956298636096391811 +
       skoYC1 *
       2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_344, refute_0_381])).

cnf(refute_0_383, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (pi *
       30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) <=
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_382])).

cnf(refute_0_384, plain,
    (368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 <
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) |
     0 <=
     -5981343255101440/1039956298636096391811 +
     skoYC1 *
     -2236029195959218121536543129600/
     3244527309218669106007406313092919683579163 |
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 *
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     pi * -3738339534438400/1039956298636096391811 +
     skoYC1 *
     (pi *
      -1397518247474511325960339456000/
      3244527309218669106007406313092919683579163)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_385, plain,
    (arctan(X_000131) < X_000132 | ~ lgen(1, arctan(X_000131), X_000132)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_386, plain,
    (X_000131 < 0 |
     ~ lgen(1,
            1/15 * (64 * X_000131 ^ 4 + 735 * X_000131 ^ 2 + 945) *
            X_000131 / (15 * X_000131 ^ 4 + 70 * X_000131 ^ 2 + 63),
            X_000132) | lgen(1, arctan(X_000131), X_000132)),
    inference(subst, [], [atan_upper_bound_case_33])).

cnf(refute_0_387, plain,
    (arctan(X_000131) < X_000132 | X_000131 < 0 |
     ~ lgen(1,
            1/15 * (64 * X_000131 ^ 4 + 735 * X_000131 ^ 2 + 945) *
            X_000131 / (15 * X_000131 ^ 4 + 70 * X_000131 ^ 2 + 63),
            X_000132)),
    inference(resolve, [], [refute_0_386, refute_0_385])).

cnf(refute_0_388, plain,
    (arctan(X_000131) < X_000132 | X_000131 < 0 |
     X_000132 <=
     X_000131 *
     (63 + X_000131 * (X_000131 * (49 + X_000131 * (X_000131 * 64/15)))) /
     (63 + X_000131 * (X_000131 * (70 + X_000131 * (X_000131 * 15))))),
    inference(arithmetic, [], [refute_0_387])).

cnf(refute_0_389, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) <
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     836177642653679/2278501083899584 *
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (49 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 64/15)))) /
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (70 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 15))))),
    inference(subst, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) <=
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163))),
    introduced(tautology, [assume])).

cnf(refute_0_391, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     0 <=
     skoYC1 *
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) |
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) <=
     arctan(836177642653679/2278501083899584)),
    inference(split, [], [refute_0_390])).

cnf(refute_0_392, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     skoYC1 *
     (5981343255101440/1039956298636096391811 +
      skoYC1 *
      2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <= 0 |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     arctan(836177642653679/2278501083899584) | skoYC1 = 0),
    inference(arithmetic, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (0 <
     skoYC1 *
     (5981343255101440/1039956298636096391811 +
      skoYC1 *
      2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) | skoYC1 = 0),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_394, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     arctan(836177642653679/2278501083899584) | skoYC1 = 0),
    inference(resolve, [], [refute_0_392, refute_0_393])).

cnf(refute_0_395, plain, (skoYC1 != 0),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_396, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     arctan(836177642653679/2278501083899584)),
    inference(resolve, [], [refute_0_394, refute_0_395])).

cnf(refute_0_397, plain,
    (836177642653679/2278501083899584 < 0 |
     skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     836177642653679/2278501083899584 *
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (49 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 64/15)))) /
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (70 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 15))))),
    inference(resolve, [], [refute_0_396, refute_0_389])).

cnf(refute_0_398, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) /
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395),
    inference(arithmetic, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     0 <=
     -5981343255101440/1039956298636096391811 +
     skoYC1 *
     -2236029195959218121536543129600/
     3244527309218669106007406313092919683579163 |
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 *
     (-5981343255101440/1039956298636096391811 +
      skoYC1 *
      -2236029195959218121536543129600/
      3244527309218669106007406313092919683579163) <=
     pi * -3738339534438400/1039956298636096391811 +
     skoYC1 *
     (pi *
      -1397518247474511325960339456000/
      3244527309218669106007406313092919683579163)),
    inference(resolve, [], [refute_0_398, refute_0_384])).

cnf(refute_0_400, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     skoYC1 *
     (-164589252385093257664549687274043248448918080349845742913564183905973730536146811911796073714674078245191680/
      679001603193107537148258008613572780336353970041499542139332542429816786289273351628414411727044090376390367941484610277 +
      pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163) <=
     440273685332381206066174553100671235179957448454590824690229761648916231870174273759820644352/
     217637864233211447464004791195867052481735309247785568184813140752837775935355565404063780351423869 +
     pi * -3738339534438400/1039956298636096391811 |
     skoYC1 <= -3119868895908289175433/373833953443840),
    inference(arithmetic, [], [refute_0_399])).

cnf(refute_0_401, plain,
    (440273685332381206066174553100671235179957448454590824690229761648916231870174273759820644352/
     217637864233211447464004791195867052481735309247785568184813140752837775935355565404063780351423869 +
     pi * -3738339534438400/1039956298636096391811 <
     skoYC1 *
     (-164589252385093257664549687274043248448918080349845742913564183905973730536146811911796073714674078245191680/
      679001603193107537148258008613572780336353970041499542139332542429816786289273351628414411727044090376390367941484610277 +
      pi *
      1397518247474511325960339456000/
      3244527309218669106007406313092919683579163) |
     skoYC1 <= -3119868895908289175433/373833953443840),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_402, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163)) |
     skoYC1 <= -3119868895908289175433/373833953443840),
    inference(resolve, [], [refute_0_400, refute_0_401])).

cnf(refute_0_403, plain,
    (-3119868895908289175433/373833953443840 < skoYC1),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_404, plain,
    (skoYC1 *
     (pi * -3738339534438400/1039956298636096391811 +
      skoYC1 *
      (pi *
       -1397518247474511325960339456000/
       3244527309218669106007406313092919683579163)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (-5981343255101440/1039956298636096391811 +
       skoYC1 *
       -2236029195959218121536543129600/
       3244527309218669106007406313092919683579163))),
    inference(resolve, [], [refute_0_402, refute_0_403])).

cnf(refute_0_405, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (pi *
       30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) <
     arctan(836177642653679/2278501083899584) *
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_383, refute_0_404])).

cnf(refute_0_406, plain,
    ((pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <
     arctan(836177642653679/2278501083899584) |
     skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(split, [], [refute_0_405])).

cnf(refute_0_407, plain,
    (0 <
     skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815)),
    inference(decision, [],
              [refute_0_335, refute_0_334, refute_0_333, refute_0_20,
               refute_0_206, refute_0_332])).

cnf(refute_0_408, plain,
    ((pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <
     arctan(836177642653679/2278501083899584) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_406, refute_0_407])).

cnf(refute_0_409, plain,
    (836177642653679/2278501083899584 < 0 |
     (pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <
     836177642653679/2278501083899584 *
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (49 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 64/15)))) /
     (63 +
      836177642653679/2278501083899584 *
      (836177642653679/2278501083899584 *
       (70 +
        836177642653679/2278501083899584 *
        (836177642653679/2278501083899584 * 15)))) |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_343, refute_0_408])).

cnf(refute_0_410, plain,
    ((pi * -10 +
      skoYC1 *
      (pi * -16900434900870496256/12619869683949029714626485 +
       skoYC1 *
       (pi *
        30392992063010096599082701152059392/
        7874467779473709920279975121876516072046628601))) /
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) <
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_409])).

cnf(refute_0_411, plain,
    (pi * -10 +
     skoYC1 *
     (pi * -16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (pi *
       30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) <
     368039139834413689736665989221257086666426891663403086128487571507498674743279/
     1046379855185471252343748962202841595263528265300827365806286471378052327014395 *
     (skoYC1 *
      (92508651052049891328/7011038713305016508125825 +
       skoYC1 *
       2305524983369630620753898890067968/
       1458234773976612948199995392940095568897523815)) |
     skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_338, refute_0_410])).

cnf(refute_0_412, plain,
    (pi * -10 <
     skoYC1 *
     (11348934786812790348767781455276002429233763148813987688040052566024113708917482254706896749461504/
      2445403224509278624977498278853561742468387016409146175854358762070080010162536839053994420893082083625 +
      pi * 16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (848523431746109780563819457772602563053552271435594743016588793708856970406051274130073478428630742985261187072/
       1525867491620066659869279802689945498033639893787703137752000074515893833633394837409409053019940747540277243512947360316925 +
       pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) |
     skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(arithmetic, [], [refute_0_411])).

cnf(refute_0_413, plain,
    (skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0 |
     skoYC1 *
     (11348934786812790348767781455276002429233763148813987688040052566024113708917482254706896749461504/
      2445403224509278624977498278853561742468387016409146175854358762070080010162536839053994420893082083625 +
      pi * 16900434900870496256/12619869683949029714626485 +
      skoYC1 *
      (848523431746109780563819457772602563053552271435594743016588793708856970406051274130073478428630742985261187072/
       1525867491620066659869279802689945498033639893787703137752000074515893833633394837409409053019940747540277243512947360316925 +
       pi *
       -30392992063010096599082701152059392/
       7874467779473709920279975121876516072046628601)) <= pi * -10),
    inference(decision, [],
              [refute_0_335, refute_0_334, refute_0_333, refute_0_20,
               refute_0_206, refute_0_332])).

cnf(refute_0_414, plain,
    (skoYC1 *
     (92508651052049891328/7011038713305016508125825 +
      skoYC1 *
      2305524983369630620753898890067968/
      1458234773976612948199995392940095568897523815) <= 0 |
     (pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_413, refute_0_412])).

cnf(refute_0_415, plain,
    ((pi * -6 + arctan(836177642653679/2278501083899584) * 12) / pi <=
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_414, refute_0_407])).

cnf(refute_0_416, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) / pi <
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10)),
    inference(resolve, [], [refute_0_415, refute_0_337])).

cnf(refute_0_417, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * 3/2 +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * 20) / ln(10) *
     pi | pi <= 0), inference(resolve, [], [refute_0_18, refute_0_416])).

cnf(refute_0_418, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * 3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20)) /
     ln(10) | pi <= 0), inference(arithmetic, [], [refute_0_417])).

cnf(refute_0_419, plain,
    (pi * 6 + arctan(836177642653679/2278501083899584) * -12 <
     (ln(10) * (pi * 3/2) +
      ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20)) /
     ln(10)), inference(resolve, [], [refute_0_418, refute_0_355])).

cnf(refute_0_420, plain,
    ((pi * 6 + arctan(836177642653679/2278501083899584) * -12) * ln(10) <
     ln(10) * (pi * 3/2) +
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20) |
     ln(10) <= 0), inference(resolve, [], [refute_0_17, refute_0_419])).

cnf(refute_0_421, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     ln(skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20)),
    inference(arithmetic, [], [refute_0_420])).

cnf(refute_0_422, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < ln(skoYC1 * 747667906887680/3119868895908289175433) |
     pi * 20 <= 0), inference(split, [], [refute_0_421])).

cnf(refute_0_423, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < ln(skoYC1 * 747667906887680/3119868895908289175433) |
     pi <= 0), inference(arithmetic, [], [refute_0_422])).

cnf(refute_0_424, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < ln(skoYC1 * 747667906887680/3119868895908289175433)),
    inference(resolve, [], [refute_0_423, refute_0_355])).

cnf(refute_0_425, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < -1 + skoYC1 * 747667906887680/3119868895908289175433 |
     skoYC1 * 747667906887680/3119868895908289175433 <= 0),
    inference(resolve, [], [refute_0_16, refute_0_424])).

cnf(refute_0_426, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < -1 + skoYC1 * 747667906887680/3119868895908289175433 |
     skoYC1 <= 0), inference(arithmetic, [], [refute_0_425])).

cnf(refute_0_427, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) /
     (pi * 20) < -1 + skoYC1 * 747667906887680/3119868895908289175433),
    inference(resolve, [], [refute_0_426, refute_0_364])).

cnf(refute_0_428, plain,
    (ln(10) * (pi * 9/2 + arctan(836177642653679/2278501083899584) * -12) <
     (-1 + skoYC1 * 747667906887680/3119868895908289175433) * (pi * 20) |
     pi * 20 <= 0), inference(resolve, [], [refute_0_11, refute_0_427])).

cnf(refute_0_429, plain,
    (pi * 20 + skoYC1 * (pi * -14953358137753600/3119868895908289175433) <
     ln(10) * (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     pi <= 0), inference(arithmetic, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (pi * 20 + skoYC1 * (pi * -14953358137753600/3119868895908289175433) <
     ln(10) * (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(resolve, [], [refute_0_429, refute_0_355])).

cnf(refute_0_431, plain,
    (ln(10) <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12),
    inference(split, [], [refute_0_430])).

cnf(refute_0_432, plain,
    (ln(10) <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     pi * 3/8 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_431])).

cnf(refute_0_433, plain,
    (836177642653679/2278501083899584 < 0 |
     arctan(836177642653679/2278501083899584) < pi * 3/8 |
     pi * 3/8 <= 836177642653679/2278501083899584),
    inference(subst, [], [refute_0_374])).

cnf(refute_0_434, plain,
    (arctan(836177642653679/2278501083899584) < pi * 3/8 |
     pi <= 836177642653679/854437906462344),
    inference(arithmetic, [], [refute_0_433])).

cnf(refute_0_435, plain, (836177642653679/854437906462344 < pi),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_436, plain,
    (arctan(836177642653679/2278501083899584) < pi * 3/8),
    inference(resolve, [], [refute_0_434, refute_0_435])).

cnf(refute_0_437, plain,
    (ln(10) <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(resolve, [], [refute_0_432, refute_0_436])).

cnf(refute_0_438, plain,
    ((-1/5 +
      10 *
      (-95/6 + 10 * (-140/3 + 10 * (20 + 10 * (115/3 + 10 * 131/30))))) /
     (10 * (5 + 10 * (40 + 10 * (60 + 10 * (20 + 10))))) <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     10 <= 0), inference(resolve, [], [refute_0_10, refute_0_437])).

cnf(refute_0_439, plain,
    (231993/101125 <
     (pi * 20 +
      skoYC1 * (pi * -14953358137753600/3119868895908289175433)) /
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12)),
    inference(arithmetic, [], [refute_0_438])).

cnf(refute_0_440, plain,
    (pi * 20 + skoYC1 * (pi * -14953358137753600/3119868895908289175433) <
     231993/101125 *
     (pi * -9/2 + arctan(836177642653679/2278501083899584) * 12) |
     0 <= pi * -9/2 + arctan(836177642653679/2278501083899584) * 12),
    inference(resolve, [], [refute_0_5, refute_0_439])).

cnf(refute_0_441, plain,
    (pi * 6132937/5567832 +
     skoYC1 * (pi * -378039585420083200000/2171363234305355192028683907) <
     arctan(836177642653679/2278501083899584) |
     pi * 3/8 <= arctan(836177642653679/2278501083899584)),
    inference(arithmetic, [], [refute_0_440])).

cnf(refute_0_442, plain,
    (pi * 6132937/5567832 +
     skoYC1 * (pi * -378039585420083200000/2171363234305355192028683907) <
     arctan(836177642653679/2278501083899584)),
    inference(resolve, [], [refute_0_441, refute_0_436])).

cnf(refute_0_443, plain,
    (836177642653679/2278501083899584 < 0 |
     pi * 6132937/5567832 +
     skoYC1 * (pi * -378039585420083200000/2171363234305355192028683907) <
     836177642653679/2278501083899584),
    inference(resolve, [], [refute_0_4, refute_0_442])).

cnf(refute_0_444, plain,
    (-836177642653679/2278501083899584 + pi * 6132937/5567832 <
     skoYC1 * (pi * 378039585420083200000/2171363234305355192028683907)),
    inference(arithmetic, [], [refute_0_443])).

cnf(refute_0_445, plain,
    (skoYC1 * (pi * 378039585420083200000/2171363234305355192028683907) <=
     -836177642653679/2278501083899584 + pi * 6132937/5567832),
    inference(decision, [],
              [refute_0_332, refute_0_206, refute_0_20, refute_0_333,
               refute_0_334, refute_0_335])).

cnf(refute_0_446, plain, ($false),
    inference(resolve, [], [refute_0_445, refute_0_444])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Fails-1-6.tptp
