% SZS output start CNFRefutation for Nichols-Plot-Inverted-Pendulum-Meets-2-5-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf5, axiom,
    (X <= 0 |
     ~ lgen(R,
            1/3 * (X ^ 2 + 19 * X + 10) * (X - 1) /
            (3 * X ^ 2 + 6 * X + 1), Y) | lgen(R, ln(X), Y))).

cnf(log_def, axiom, (log(X) = 868589/2000000 * ln(X))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Inverted_Pendulum_Meets_2_5', conjecture,
    (! [X] :
       (X = 629/512 =>
        10 *
        log(25 ^ 2 *
            (484 * X ^ 10 + 35161 * X ^ 8 + 781414 * X ^ 6 +
             4871449 * X ^ 4 + 569821 * X ^ 2 + 9604) /
            (121 * X ^ 6 + 7550 * X ^ 4 + 117845 * X ^ 2 + 2401) ^ 2) <
        12 / pi *
        (-pi -
         arctan((22 * X ^ 5 + 650 * X ^ 3 - 1029 * X) /
                (81 * X ^ 4 + 2495 * X ^ 2 - 98))) + 6))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 629/512 =>
        10 *
        log(25 ^ 2 *
            (484 * X ^ 10 + 35161 * X ^ 8 + 781414 * X ^ 6 +
             4871449 * X ^ 4 + 569821 * X ^ 2 + 9604) /
            (121 * X ^ 6 + 7550 * X ^ 4 + 117845 * X ^ 2 + 2401) ^ 2) <
        12 / pi *
        (-pi -
         arctan((22 * X ^ 5 + 650 * X ^ 3 - 1029 * X) /
                (81 * X ^ 4 + 2495 * X ^ 2 - 98))) + 6)),
    inference(strip, [], ['Nichols_Plot_Inverted_Pendulum_Meets_2_5'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 629/512 =>
          10 *
          log(25 ^ 2 *
              (484 * X ^ 10 + 35161 * X ^ 8 + 781414 * X ^ 6 +
               4871449 * X ^ 4 + 569821 * X ^ 2 + 9604) /
              (121 * X ^ 6 + 7550 * X ^ 4 + 117845 * X ^ 2 + 2401) ^ 2) <
          12 / pi *
          (-pi -
           arctan((22 * X ^ 5 + 650 * X ^ 3 - 1029 * X) /
                  (81 * X ^ 4 + 2495 * X ^ 2 - 98))) + 6)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (12 / pi *
        (-pi -
         arctan((22 * X ^ 5 + 650 * X ^ 3 - 1029 * X) /
                (81 * X ^ 4 + 2495 * X ^ 2 - 98))) + 6 <=
        10 *
        log(25 ^ 2 *
            (484 * X ^ 10 + 35161 * X ^ 8 + 781414 * X ^ 6 +
             4871449 * X ^ 4 + 569821 * X ^ 2 + 9604) /
            (121 * X ^ 6 + 7550 * X ^ 4 + 117845 * X ^ 2 + 2401) ^ 2) &
        X = 629/512)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 <=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) & skoXC1 = 629/512),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 <=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2)), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 629/512),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (arctan(X_000041) < X_000042 | ~ lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (X_000041 < 0 | ~ lgen(1, X_000041, X_000042) |
     lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_2, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 |
     ~ lgen(1, X_000041, X_000042)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 | X_000042 <= X_000041),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (2187307897531/3226980916903168 < 0 |
     arctan(2187307897531/3226980916903168) <
     pi *
     11832891332215638782211167414341154826944761742949985256368726988561/
     675544967691940330551304227679429166256359013499838293840629803280000 |
     pi *
     11832891332215638782211167414341154826944761742949985256368726988561/
     675544967691940330551304227679429166256359013499838293840629803280000 <=
     2187307897531/3226980916903168), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (-4024980458861277359693287403525208561875919088232399122907193490/
     2814770698716418043963767615330954859401495889582659557669290847 <
     (pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) | 0 <= pi * -868589/200000 |
     -4024980458861277359693287403525208561875919088232399122907193490/
     2814770698716418043963767615330954859401495889582659557669290847 *
     (pi * -868589/200000) <=
     pi * 6 + arctan(2187307897531/3226980916903168) * 12),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_6, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 <=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2)), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain, (skoXC1 = 629/512),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_8, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_9, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 629/512 |
     skoXC1 ^ 2 = 629/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_10, plain, (skoXC1 != 629/512 | skoXC1 ^ 2 = 629/512 ^ 2),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain, (skoXC1 ^ 2 = 629/512 ^ 2),
    inference(resolve, [], [refute_0_7, refute_0_10])).

cnf(refute_0_12, plain, (2495 * skoXC1 ^ 2 = 2495 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_13, plain,
    (2495 * skoXC1 ^ 2 != 2495 * skoXC1 ^ 2 | skoXC1 ^ 2 != 629/512 ^ 2 |
     2495 * skoXC1 ^ 2 = 2495 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_14, plain,
    (skoXC1 ^ 2 != 629/512 ^ 2 | 2495 * skoXC1 ^ 2 = 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_12, refute_0_13])).

cnf(refute_0_15, plain, (2495 * skoXC1 ^ 2 = 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_14])).

cnf(refute_0_16, plain,
    (81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_17, plain,
    (2495 * skoXC1 ^ 2 != 2495 * 629/512 ^ 2 |
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 |
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_18, plain,
    (2495 * skoXC1 ^ 2 != 2495 * 629/512 ^ 2 |
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_21, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 629/512 |
     skoXC1 ^ 4 = 629/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_22, plain, (skoXC1 != 629/512 | skoXC1 ^ 4 = 629/512 ^ 4),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 ^ 4 = 629/512 ^ 4),
    inference(resolve, [], [refute_0_7, refute_0_22])).

cnf(refute_0_24, plain, (81 * skoXC1 ^ 4 = 81 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_25, plain,
    (81 * skoXC1 ^ 4 != 81 * skoXC1 ^ 4 | skoXC1 ^ 4 != 629/512 ^ 4 |
     81 * skoXC1 ^ 4 = 81 * 629/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_26, plain,
    (skoXC1 ^ 4 != 629/512 ^ 4 | 81 * skoXC1 ^ 4 = 81 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain, (81 * skoXC1 ^ 4 = 81 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_26])).

cnf(refute_0_28, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2), introduced(tautology, [refl])).

cnf(refute_0_29, plain,
    (81 * skoXC1 ^ 4 != 81 * 629/512 ^ 4 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_30, plain,
    (81 * skoXC1 ^ 4 != 81 * 629/512 ^ 4 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_28, refute_0_29])).

cnf(refute_0_31, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_27, refute_0_30])).

cnf(refute_0_32, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_33, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_34, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_36, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (81 * 629/512 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_31, refute_0_37])).

cnf(refute_0_39, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_19, refute_0_38])).

cnf(refute_0_40, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 =
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98),
    introduced(tautology, [refl])).

cnf(refute_0_41, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 !=
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98),
    introduced(tautology, [equality])).

cnf(refute_0_42, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 |
     81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 =
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98),
    inference(resolve, [], [refute_0_39, refute_0_42])).

cnf(refute_0_44, plain,
    ((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)),
    introduced(tautology, [refl])).

cnf(refute_0_45, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98 |
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) |
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    introduced(tautology, [equality])).

cnf(refute_0_46, plain,
    (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98 !=
     81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98 |
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    ((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    inference(resolve, [], [refute_0_43, refute_0_46])).

cnf(refute_0_48, plain, (1029 * skoXC1 = 1029 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_49, plain,
    (1029 * skoXC1 != 1029 * skoXC1 | skoXC1 != 629/512 |
     1029 * skoXC1 = 1029 * 629/512), introduced(tautology, [equality])).

cnf(refute_0_50, plain,
    (skoXC1 != 629/512 | 1029 * skoXC1 = 1029 * 629/512),
    inference(resolve, [], [refute_0_48, refute_0_49])).

cnf(refute_0_51, plain, (1029 * skoXC1 = 1029 * 629/512),
    inference(resolve, [], [refute_0_7, refute_0_50])).

cnf(refute_0_52, plain,
    (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_53, plain,
    (1029 * skoXC1 != 1029 * 629/512 |
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 |
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    introduced(tautology, [equality])).

cnf(refute_0_54, plain,
    (1029 * skoXC1 != 1029 * 629/512 |
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    inference(resolve, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain,
    (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    inference(resolve, [], [refute_0_51, refute_0_54])).

cnf(refute_0_56, plain, (skoXC1 ^ 3 = skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_57, plain,
    (skoXC1 ^ 3 != skoXC1 ^ 3 | skoXC1 != 629/512 |
     skoXC1 ^ 3 = 629/512 ^ 3), introduced(tautology, [equality])).

cnf(refute_0_58, plain, (skoXC1 != 629/512 | skoXC1 ^ 3 = 629/512 ^ 3),
    inference(resolve, [], [refute_0_56, refute_0_57])).

cnf(refute_0_59, plain, (skoXC1 ^ 3 = 629/512 ^ 3),
    inference(resolve, [], [refute_0_7, refute_0_58])).

cnf(refute_0_60, plain, (650 * skoXC1 ^ 3 = 650 * skoXC1 ^ 3),
    introduced(tautology, [refl])).

cnf(refute_0_61, plain,
    (650 * skoXC1 ^ 3 != 650 * skoXC1 ^ 3 | skoXC1 ^ 3 != 629/512 ^ 3 |
     650 * skoXC1 ^ 3 = 650 * 629/512 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_62, plain,
    (skoXC1 ^ 3 != 629/512 ^ 3 | 650 * skoXC1 ^ 3 = 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain, (650 * skoXC1 ^ 3 = 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_59, refute_0_62])).

cnf(refute_0_64, plain,
    (22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3), introduced(tautology, [refl])).

cnf(refute_0_65, plain,
    (650 * skoXC1 ^ 3 != 650 * 629/512 ^ 3 |
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 |
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_66, plain,
    (650 * skoXC1 ^ 3 != 650 * 629/512 ^ 3 |
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_64, refute_0_65])).

cnf(refute_0_67, plain,
    (22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_63, refute_0_66])).

cnf(refute_0_68, plain, (skoXC1 ^ 5 = skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_69, plain,
    (skoXC1 ^ 5 != skoXC1 ^ 5 | skoXC1 != 629/512 |
     skoXC1 ^ 5 = 629/512 ^ 5), introduced(tautology, [equality])).

cnf(refute_0_70, plain, (skoXC1 != 629/512 | skoXC1 ^ 5 = 629/512 ^ 5),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain, (skoXC1 ^ 5 = 629/512 ^ 5),
    inference(resolve, [], [refute_0_7, refute_0_70])).

cnf(refute_0_72, plain, (22 * skoXC1 ^ 5 = 22 * skoXC1 ^ 5),
    introduced(tautology, [refl])).

cnf(refute_0_73, plain,
    (22 * skoXC1 ^ 5 != 22 * skoXC1 ^ 5 | skoXC1 ^ 5 != 629/512 ^ 5 |
     22 * skoXC1 ^ 5 = 22 * 629/512 ^ 5),
    introduced(tautology, [equality])).

cnf(refute_0_74, plain,
    (skoXC1 ^ 5 != 629/512 ^ 5 | 22 * skoXC1 ^ 5 = 22 * 629/512 ^ 5),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain, (22 * skoXC1 ^ 5 = 22 * 629/512 ^ 5),
    inference(resolve, [], [refute_0_71, refute_0_74])).

cnf(refute_0_76, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3), introduced(tautology, [refl])).

cnf(refute_0_77, plain,
    (22 * skoXC1 ^ 5 != 22 * 629/512 ^ 5 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3),
    introduced(tautology, [equality])).

cnf(refute_0_78, plain,
    (22 * skoXC1 ^ 5 != 22 * 629/512 ^ 5 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3),
    inference(resolve, [], [refute_0_75, refute_0_78])).

cnf(refute_0_80, plain,
    (22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_81, plain,
    (22 * 629/512 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3),
    inference(resolve, [], [refute_0_67, refute_0_81])).

cnf(refute_0_83, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1),
    introduced(tautology, [refl])).

cnf(refute_0_84, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 !=
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1),
    introduced(tautology, [equality])).

cnf(refute_0_85, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1),
    inference(resolve, [], [refute_0_82, refute_0_85])).

cnf(refute_0_87, plain,
    (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_88, plain,
    (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512 |
     22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    inference(resolve, [], [refute_0_86, refute_0_87])).

cnf(refute_0_89, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 =
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512),
    inference(resolve, [], [refute_0_55, refute_0_88])).

cnf(refute_0_90, plain,
    ((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)),
    introduced(tautology, [refl])).

cnf(refute_0_91, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512 |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)),
    introduced(tautology, [equality])).

cnf(refute_0_92, plain,
    (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1 !=
     22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512 |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)),
    inference(resolve, [], [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain,
    ((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)),
    inference(resolve, [], [refute_0_89, refute_0_92])).

cnf(refute_0_94, plain,
    ((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98) |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_95, plain,
    ((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98) |
     (22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    ((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) =
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)),
    inference(resolve, [], [refute_0_47, refute_0_95])).

cnf(refute_0_97, plain,
    (arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    ((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98) |
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) |
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    ((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
     (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98) !=
     (22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
     (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98) |
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain,
    (-pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     -pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))),
    introduced(tautology, [refl])).

cnf(refute_0_102, plain,
    (-pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     -pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) |
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)) |
     -pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     -pi -
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    introduced(tautology, [equality])).

cnf(refute_0_103, plain,
    (arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)) |
     -pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     -pi -
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    inference(resolve, [], [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain,
    (-pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) =
     -pi -
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))),
    inference(resolve, [], [refute_0_100, refute_0_103])).

cnf(refute_0_105, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) =
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)))),
    introduced(tautology, [refl])).

cnf(refute_0_106, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) !=
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) |
     -pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     -pi -
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)) |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)))),
    introduced(tautology, [equality])).

cnf(refute_0_107, plain,
    (-pi -
     arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
            (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98)) !=
     -pi -
     arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
            (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)) |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)))),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98)))),
    inference(resolve, [], [refute_0_104, refute_0_107])).

cnf(refute_0_109, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 =
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6),
    introduced(tautology, [refl])).

cnf(refute_0_110, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) !=
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 !=
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6),
    introduced(tautology, [equality])).

cnf(refute_0_111, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) !=
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    (12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 =
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6),
    inference(resolve, [], [refute_0_108, refute_0_111])).

cnf(refute_0_113, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) <
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 !=
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2)), introduced(tautology, [equality])).

cnf(refute_0_114, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) <
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2)),
    inference(resolve, [], [refute_0_112, refute_0_113])).

cnf(refute_0_115, plain, (117845 * skoXC1 ^ 2 = 117845 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (117845 * skoXC1 ^ 2 != 117845 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 629/512 ^ 2 |
     117845 * skoXC1 ^ 2 = 117845 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (skoXC1 ^ 2 != 629/512 ^ 2 |
     117845 * skoXC1 ^ 2 = 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain, (117845 * skoXC1 ^ 2 = 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_117])).

cnf(refute_0_119, plain,
    (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (117845 * skoXC1 ^ 2 != 117845 * 629/512 ^ 2 |
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 |
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (117845 * skoXC1 ^ 2 != 117845 * 629/512 ^ 2 |
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain,
    (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_118, refute_0_121])).

cnf(refute_0_123, plain, (7550 * skoXC1 ^ 4 = 7550 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (7550 * skoXC1 ^ 4 != 7550 * skoXC1 ^ 4 | skoXC1 ^ 4 != 629/512 ^ 4 |
     7550 * skoXC1 ^ 4 = 7550 * 629/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (skoXC1 ^ 4 != 629/512 ^ 4 | 7550 * skoXC1 ^ 4 = 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain, (7550 * skoXC1 ^ 4 = 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_125])).

cnf(refute_0_127, plain,
    (121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_128, plain,
    (7550 * skoXC1 ^ 4 != 7550 * 629/512 ^ 4 |
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 |
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_129, plain,
    (7550 * skoXC1 ^ 4 != 7550 * 629/512 ^ 4 |
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_126, refute_0_129])).

cnf(refute_0_131, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 629/512 |
     skoXC1 ^ 6 = 629/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_133, plain, (skoXC1 != 629/512 | skoXC1 ^ 6 = 629/512 ^ 6),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 ^ 6 = 629/512 ^ 6),
    inference(resolve, [], [refute_0_7, refute_0_133])).

cnf(refute_0_135, plain, (121 * skoXC1 ^ 6 = 121 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (121 * skoXC1 ^ 6 != 121 * skoXC1 ^ 6 | skoXC1 ^ 6 != 629/512 ^ 6 |
     121 * skoXC1 ^ 6 = 121 * 629/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (skoXC1 ^ 6 != 629/512 ^ 6 | 121 * skoXC1 ^ 6 = 121 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain, (121 * skoXC1 ^ 6 = 121 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4), introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (121 * skoXC1 ^ 6 != 121 * 629/512 ^ 6 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (121 * skoXC1 ^ 6 != 121 * 629/512 ^ 6 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_138, refute_0_141])).

cnf(refute_0_143, plain,
    (121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_144, plain,
    (121 * 629/512 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_130, refute_0_144])).

cnf(refute_0_146, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_147, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_148, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_145, refute_0_148])).

cnf(refute_0_150, plain,
    (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_151, plain,
    (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_122, refute_0_151])).

cnf(refute_0_153, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 =
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401),
    introduced(tautology, [refl])).

cnf(refute_0_154, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 !=
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 + 2401),
    introduced(tautology, [equality])).

cnf(refute_0_155, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 |
     121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 + 2401),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 =
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 + 2401),
    inference(resolve, [], [refute_0_152, refute_0_155])).

cnf(refute_0_157, plain,
    ((121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), introduced(tautology, [refl])).

cnf(refute_0_158, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 + 2401 |
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 |
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), introduced(tautology, [equality])).

cnf(refute_0_159, plain,
    (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401 !=
     121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 + 2401 |
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_157, refute_0_158])).

cnf(refute_0_160, plain,
    ((121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_156, refute_0_159])).

cnf(refute_0_161, plain,
    (25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), introduced(tautology, [refl])).

cnf(refute_0_162, plain,
    (25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 |
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), introduced(tautology, [equality])).

cnf(refute_0_163, plain,
    ((121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_160, refute_0_163])).

cnf(refute_0_165, plain, (569821 * skoXC1 ^ 2 = 569821 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (569821 * skoXC1 ^ 2 != 569821 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 629/512 ^ 2 |
     569821 * skoXC1 ^ 2 = 569821 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (skoXC1 ^ 2 != 629/512 ^ 2 |
     569821 * skoXC1 ^ 2 = 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain, (569821 * skoXC1 ^ 2 = 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_11, refute_0_167])).

cnf(refute_0_169, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_170, plain,
    (569821 * skoXC1 ^ 2 != 569821 * 629/512 ^ 2 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_171, plain,
    (569821 * skoXC1 ^ 2 != 569821 * 629/512 ^ 2 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_168, refute_0_171])).

cnf(refute_0_173, plain, (4871449 * skoXC1 ^ 4 = 4871449 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_174, plain,
    (4871449 * skoXC1 ^ 4 != 4871449 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 629/512 ^ 4 |
     4871449 * skoXC1 ^ 4 = 4871449 * 629/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_175, plain,
    (skoXC1 ^ 4 != 629/512 ^ 4 |
     4871449 * skoXC1 ^ 4 = 4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_173, refute_0_174])).

cnf(refute_0_176, plain, (4871449 * skoXC1 ^ 4 = 4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_23, refute_0_175])).

cnf(refute_0_177, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4), introduced(tautology, [refl])).

cnf(refute_0_178, plain,
    (4871449 * skoXC1 ^ 4 != 4871449 * 629/512 ^ 4 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_179, plain,
    (4871449 * skoXC1 ^ 4 != 4871449 * 629/512 ^ 4 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_176, refute_0_179])).

cnf(refute_0_181, plain, (781414 * skoXC1 ^ 6 = 781414 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_182, plain,
    (781414 * skoXC1 ^ 6 != 781414 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 629/512 ^ 6 |
     781414 * skoXC1 ^ 6 = 781414 * 629/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_183, plain,
    (skoXC1 ^ 6 != 629/512 ^ 6 |
     781414 * skoXC1 ^ 6 = 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_181, refute_0_182])).

cnf(refute_0_184, plain, (781414 * skoXC1 ^ 6 = 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_134, refute_0_183])).

cnf(refute_0_185, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_186, plain,
    (781414 * skoXC1 ^ 6 != 781414 * 629/512 ^ 6 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_187, plain,
    (781414 * skoXC1 ^ 6 != 781414 * 629/512 ^ 6 |
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_185, refute_0_186])).

cnf(refute_0_188, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_184, refute_0_187])).

cnf(refute_0_189, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_190, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 629/512 |
     skoXC1 ^ 8 = 629/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_191, plain, (skoXC1 != 629/512 | skoXC1 ^ 8 = 629/512 ^ 8),
    inference(resolve, [], [refute_0_189, refute_0_190])).

cnf(refute_0_192, plain, (skoXC1 ^ 8 = 629/512 ^ 8),
    inference(resolve, [], [refute_0_7, refute_0_191])).

cnf(refute_0_193, plain, (35161 * skoXC1 ^ 8 = 35161 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_194, plain,
    (35161 * skoXC1 ^ 8 != 35161 * skoXC1 ^ 8 | skoXC1 ^ 8 != 629/512 ^ 8 |
     35161 * skoXC1 ^ 8 = 35161 * 629/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_195, plain,
    (skoXC1 ^ 8 != 629/512 ^ 8 | 35161 * skoXC1 ^ 8 = 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_193, refute_0_194])).

cnf(refute_0_196, plain, (35161 * skoXC1 ^ 8 = 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_192, refute_0_195])).

cnf(refute_0_197, plain,
    (484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_198, plain,
    (35161 * skoXC1 ^ 8 != 35161 * 629/512 ^ 8 |
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 |
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_199, plain,
    (35161 * skoXC1 ^ 8 != 35161 * 629/512 ^ 8 |
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_197, refute_0_198])).

cnf(refute_0_200, plain,
    (484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_196, refute_0_199])).

cnf(refute_0_201, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_202, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 629/512 |
     skoXC1 ^ 10 = 629/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_203, plain, (skoXC1 != 629/512 | skoXC1 ^ 10 = 629/512 ^ 10),
    inference(resolve, [], [refute_0_201, refute_0_202])).

cnf(refute_0_204, plain, (skoXC1 ^ 10 = 629/512 ^ 10),
    inference(resolve, [], [refute_0_7, refute_0_203])).

cnf(refute_0_205, plain, (484 * skoXC1 ^ 10 = 484 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_206, plain,
    (484 * skoXC1 ^ 10 != 484 * skoXC1 ^ 10 | skoXC1 ^ 10 != 629/512 ^ 10 |
     484 * skoXC1 ^ 10 = 484 * 629/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_207, plain,
    (skoXC1 ^ 10 != 629/512 ^ 10 | 484 * skoXC1 ^ 10 = 484 * 629/512 ^ 10),
    inference(resolve, [], [refute_0_205, refute_0_206])).

cnf(refute_0_208, plain, (484 * skoXC1 ^ 10 = 484 * 629/512 ^ 10),
    inference(resolve, [], [refute_0_204, refute_0_207])).

cnf(refute_0_209, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_210, plain,
    (484 * skoXC1 ^ 10 != 484 * 629/512 ^ 10 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_211, plain,
    (484 * skoXC1 ^ 10 != 484 * 629/512 ^ 10 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_209, refute_0_210])).

cnf(refute_0_212, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_208, refute_0_211])).

cnf(refute_0_213, plain,
    (484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_214, plain,
    (484 * 629/512 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_212, refute_0_213])).

cnf(refute_0_215, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8),
    inference(resolve, [], [refute_0_200, refute_0_214])).

cnf(refute_0_216, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_217, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_218, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_216, refute_0_217])).

cnf(refute_0_219, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_215, refute_0_218])).

cnf(refute_0_220, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_221, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_219, refute_0_220])).

cnf(refute_0_222, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6),
    inference(resolve, [], [refute_0_188, refute_0_221])).

cnf(refute_0_223, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4), introduced(tautology, [refl])).

cnf(refute_0_224, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_225, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_222, refute_0_225])).

cnf(refute_0_227, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4), inference(subst, [], [refute_0_36])).

cnf(refute_0_228, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_226, refute_0_227])).

cnf(refute_0_229, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4),
    inference(resolve, [], [refute_0_180, refute_0_228])).

cnf(refute_0_230, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_231, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_232, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_230, refute_0_231])).

cnf(refute_0_233, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_229, refute_0_232])).

cnf(refute_0_234, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    inference(subst, [], [refute_0_36])).

cnf(refute_0_235, plain,
    (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_233, refute_0_234])).

cnf(refute_0_236, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_235])).

cnf(refute_0_237, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 =
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604),
    introduced(tautology, [refl])).

cnf(refute_0_238, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 !=
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604),
    introduced(tautology, [equality])).

cnf(refute_0_239, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 =
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604),
    inference(resolve, [], [refute_0_236, refute_0_239])).

cnf(refute_0_241, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) =
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604)),
    introduced(tautology, [refl])).

cnf(refute_0_242, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) !=
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) |
     484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604)),
    introduced(tautology, [equality])).

cnf(refute_0_243, plain,
    (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
     4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604 !=
     484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
     4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604)),
    inference(resolve, [], [refute_0_241, refute_0_242])).

cnf(refute_0_244, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604)),
    inference(resolve, [], [refute_0_240, refute_0_243])).

cnf(refute_0_245, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), introduced(tautology, [refl])).

cnf(refute_0_246, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), introduced(tautology, [equality])).

cnf(refute_0_247, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), inference(resolve, [], [refute_0_245, refute_0_246])).

cnf(refute_0_248, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2), inference(resolve, [], [refute_0_244, refute_0_247])).

cnf(refute_0_249, plain,
    (25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(subst, [], [refute_0_36])).

cnf(refute_0_250, plain,
    (25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_248, refute_0_249])).

cnf(refute_0_251, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 =
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2), inference(resolve, [], [refute_0_164, refute_0_250])).

cnf(refute_0_252, plain,
    (ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2)), introduced(tautology, [refl])).

cnf(refute_0_253, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) |
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)), introduced(tautology, [equality])).

cnf(refute_0_254, plain,
    (25 ^ 2 *
     (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
      4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
     (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 + 2401) ^
     2 !=
     25 ^ 2 *
     (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
      4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
     (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
      2401) ^ 2 |
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_252, refute_0_253])).

cnf(refute_0_255, plain,
    (ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_251, refute_0_254])).

cnf(refute_0_256, plain,
    (868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2)), introduced(tautology, [refl])).

cnf(refute_0_257, plain,
    (868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) |
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)), introduced(tautology, [equality])).

cnf(refute_0_258, plain,
    (ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_256, refute_0_257])).

cnf(refute_0_259, plain,
    (868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_255, refute_0_258])).

cnf(refute_0_260, plain,
    (log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2)), inference(subst, [], [log_def])).

cnf(refute_0_261, plain,
    (868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) |
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)), inference(subst, [], [refute_0_36])).

cnf(refute_0_262, plain,
    (868589/2000000 *
     ln(25 ^ 2 *
        (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
         4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
        (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
         2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_260, refute_0_261])).

cnf(refute_0_263, plain,
    (log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2)),
    inference(resolve, [], [refute_0_259, refute_0_262])).

cnf(refute_0_264, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2)), introduced(tautology, [refl])).

cnf(refute_0_265, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) !=
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) |
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))), introduced(tautology, [equality])).

cnf(refute_0_266, plain,
    (log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) !=
     868589/2000000 *
     ln(25 ^ 2 *
        (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
         4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
        (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
         2401) ^ 2) |
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))),
    inference(resolve, [], [refute_0_264, refute_0_265])).

cnf(refute_0_267, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) =
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))),
    inference(resolve, [], [refute_0_263, refute_0_266])).

cnf(refute_0_268, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) <
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 |
     10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) !=
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2)) |
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))), introduced(tautology, [equality])).

cnf(refute_0_269, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) <
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))),
    inference(resolve, [], [refute_0_267, refute_0_268])).

cnf(refute_0_270, plain,
    (10 *
     log(25 ^ 2 *
         (484 * skoXC1 ^ 10 + 35161 * skoXC1 ^ 8 + 781414 * skoXC1 ^ 6 +
          4871449 * skoXC1 ^ 4 + 569821 * skoXC1 ^ 2 + 9604) /
         (121 * skoXC1 ^ 6 + 7550 * skoXC1 ^ 4 + 117845 * skoXC1 ^ 2 +
          2401) ^ 2) <
     12 / pi *
     (-pi -
      arctan((22 * skoXC1 ^ 5 + 650 * skoXC1 ^ 3 - 1029 * skoXC1) /
             (81 * skoXC1 ^ 4 + 2495 * skoXC1 ^ 2 - 98))) + 6 |
     12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))),
    inference(resolve, [], [refute_0_114, refute_0_269])).

cnf(refute_0_271, plain,
    (12 / pi *
     (-pi -
      arctan((22 * 629/512 ^ 5 + 650 * 629/512 ^ 3 - 1029 * 629/512) /
             (81 * 629/512 ^ 4 + 2495 * 629/512 ^ 2 - 98))) + 6 <=
     10 *
     (868589/2000000 *
      ln(25 ^ 2 *
         (484 * 629/512 ^ 10 + 35161 * 629/512 ^ 8 + 781414 * 629/512 ^ 6 +
          4871449 * 629/512 ^ 4 + 569821 * 629/512 ^ 2 + 9604) /
         (121 * 629/512 ^ 6 + 7550 * 629/512 ^ 4 + 117845 * 629/512 ^ 2 +
          2401) ^ 2))),
    inference(resolve, [], [refute_0_6, refute_0_270])).

cnf(refute_0_272, plain,
    ((pi * -12 + arctan(2187307897531/3226980916903168) * -12) / pi <=
     -6 + ln(168863878248267776000/712910064512710111149) * 868589/200000),
    inference(arithmetic, [], [refute_0_271])).

cnf(refute_0_273, plain,
    (-6 + ln(168863878248267776000/712910064512710111149) * 868589/200000 <
     (pi * -12 + arctan(2187307897531/3226980916903168) * -12) / pi |
     pi * -12 + arctan(2187307897531/3226980916903168) * -12 <=
     (-6 +
      ln(168863878248267776000/712910064512710111149) * 868589/200000) *
     pi | pi <= 0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_274, plain,
    (pi * -12 + arctan(2187307897531/3226980916903168) * -12 <=
     (-6 +
      ln(168863878248267776000/712910064512710111149) * 868589/200000) *
     pi | pi <= 0), inference(resolve, [], [refute_0_272, refute_0_273])).

cnf(refute_0_275, plain,
    (ln(168863878248267776000/712910064512710111149) *
     (pi * -868589/200000) <=
     pi * 6 + arctan(2187307897531/3226980916903168) * 12 | pi <= 0),
    inference(arithmetic, [], [refute_0_274])).

cnf(refute_0_276, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_277, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_278, plain, (0 < pi),
    inference(decision, [], [refute_0_7, refute_0_276, refute_0_277])).

cnf(refute_0_279, plain,
    (ln(168863878248267776000/712910064512710111149) *
     (pi * -868589/200000) <=
     pi * 6 + arctan(2187307897531/3226980916903168) * 12),
    inference(resolve, [], [refute_0_275, refute_0_278])).

cnf(refute_0_280, plain,
    (0 <= pi * -868589/200000 |
     (pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     ln(168863878248267776000/712910064512710111149)),
    inference(split, [], [refute_0_279])).

cnf(refute_0_281, plain,
    ((pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     ln(168863878248267776000/712910064512710111149) | pi <= 0),
    inference(arithmetic, [], [refute_0_280])).

cnf(refute_0_282, plain,
    ((pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     ln(168863878248267776000/712910064512710111149)),
    inference(resolve, [], [refute_0_281, refute_0_278])).

cnf(refute_0_283, plain,
    (ln(X_000118) < X_000119 | ~ lgen(1, ln(X_000118), X_000119)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_284, plain,
    (~ lgen(1,
            1/3 * (X_000118 ^ 2 + 19 * X_000118 + 10) * (X_000118 - 1) /
            (3 * X_000118 ^ 2 + 6 * X_000118 + 1), X_000119) |
     X_000118 <= 0 | lgen(1, ln(X_000118), X_000119)),
    inference(subst, [], [ln_upper_bound_cf5])).

cnf(refute_0_285, plain,
    (ln(X_000118) < X_000119 |
     ~ lgen(1,
            1/3 * (X_000118 ^ 2 + 19 * X_000118 + 10) * (X_000118 - 1) /
            (3 * X_000118 ^ 2 + 6 * X_000118 + 1), X_000119) |
     X_000118 <= 0), inference(resolve, [], [refute_0_284, refute_0_283])).

cnf(refute_0_286, plain,
    (ln(X_000118) < X_000119 | X_000118 <= 0 |
     X_000119 <=
     (-10/3 + X_000118 * (-3 + X_000118 * (6 + X_000118 * 1/3))) /
     (1 + X_000118 * (6 + X_000118 * 3))),
    inference(arithmetic, [], [refute_0_285])).

cnf(refute_0_287, plain,
    (ln(168863878248267776000/712910064512710111149) <
     (pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) |
     168863878248267776000/712910064512710111149 <= 0 |
     (pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     (-10/3 +
      168863878248267776000/712910064512710111149 *
      (-3 +
       168863878248267776000/712910064512710111149 *
       (6 + 168863878248267776000/712910064512710111149 * 1/3))) /
     (1 +
      168863878248267776000/712910064512710111149 *
      (6 + 168863878248267776000/712910064512710111149 * 3))),
    inference(subst, [], [refute_0_286])).

cnf(refute_0_288, plain,
    (168863878248267776000/712910064512710111149 <= 0 |
     (pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     (-10/3 +
      168863878248267776000/712910064512710111149 *
      (-3 +
       168863878248267776000/712910064512710111149 *
       (6 + 168863878248267776000/712910064512710111149 * 1/3))) /
     (1 +
      168863878248267776000/712910064512710111149 *
      (6 + 168863878248267776000/712910064512710111149 * 3))),
    inference(resolve, [], [refute_0_282, refute_0_287])).

cnf(refute_0_289, plain,
    ((pi * 6 + arctan(2187307897531/3226980916903168) * 12) /
     (pi * -868589/200000) <=
     -4024980458861277359693287403525208561875919088232399122907193490/
     2814770698716418043963767615330954859401495889582659557669290847),
    inference(arithmetic, [], [refute_0_288])).

cnf(refute_0_290, plain,
    (0 <= pi * -868589/200000 |
     -4024980458861277359693287403525208561875919088232399122907193490/
     2814770698716418043963767615330954859401495889582659557669290847 *
     (pi * -868589/200000) <=
     pi * 6 + arctan(2187307897531/3226980916903168) * 12),
    inference(resolve, [], [refute_0_289, refute_0_5])).

cnf(refute_0_291, plain,
    (pi *
     11832891332215638782211167414341154826944761742949985256368726988561/
     675544967691940330551304227679429166256359013499838293840629803280000 <=
     arctan(2187307897531/3226980916903168) | pi <= 0),
    inference(arithmetic, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (pi *
     11832891332215638782211167414341154826944761742949985256368726988561/
     675544967691940330551304227679429166256359013499838293840629803280000 <=
     arctan(2187307897531/3226980916903168)),
    inference(resolve, [], [refute_0_291, refute_0_278])).

cnf(refute_0_293, plain,
    (2187307897531/3226980916903168 < 0 |
     pi *
     11832891332215638782211167414341154826944761742949985256368726988561/
     675544967691940330551304227679429166256359013499838293840629803280000 <=
     2187307897531/3226980916903168),
    inference(resolve, [], [refute_0_292, refute_0_4])).

cnf(refute_0_294, plain,
    (pi <=
     11543944085702385360250612627933155817987788156073848003976798142166210372669375/
     298316519694131024575620715436031773368966412917254269833124232307222872802192666),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (11543944085702385360250612627933155817987788156073848003976798142166210372669375/
     298316519694131024575620715436031773368966412917254269833124232307222872802192666 <
     pi),
    inference(decision, [], [refute_0_7, refute_0_276, refute_0_277])).

cnf(refute_0_296, plain, ($false),
    inference(resolve, [], [refute_0_294, refute_0_295])).
% SZS output end CNFRefutation for Nichols-Plot-Inverted-Pendulum-Meets-2-5-weak.tptp
