% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-6-scaled.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_14, axiom,
    (X <= 1 | ~ lgen(R, Y, 157079/100000 - 1 / X) |
     lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

cnf(sqrt_lower_bound_0, axiom,
    (X < 0 | ~ lgen(R, Y, 2 * X / (X + 1)) | lgen(R, Y, sqrt(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_1_6', conjecture,
    (! [X] :
       (X = 354991/512 =>
        20 / ln(10) *
        ln(33600000000000 /
           (10 ^ 15 *
            sqrt((9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2) / 10 ^ 30))) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 354991/512 =>
        20 / ln(10) *
        ln(33600000000000 /
           (10 ^ 15 *
            sqrt((9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2) / 10 ^ 30))) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_6'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 354991/512 =>
          20 / ln(10) *
          ln(33600000000000 /
             (10 ^ 15 *
              sqrt((9 * X ^ 10 + 309000000 * X ^ 8 +
                    10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                    4000000000000000000000000 * X ^ 2) / 10 ^ 30))) <
          -12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18 <=
        20 / ln(10) *
        ln(33600000000000 /
           (10 ^ 15 *
            sqrt((9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2) / 10 ^ 30))) &
        X = 354991/512)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) &
     skoXC1 = 354991/512), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 354991/512),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (X_000056 < arctan(X_000055) | ~ lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1 |
     lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [atan_lower_bound_case_14])).

cnf(refute_0_2, plain,
    (X_000056 < arctan(X_000055) |
     ~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000056 < arctan(X_000055) | X_000055 <= 1 |
     157079/100000 + 1 / X_000055 * -1 <= X_000056),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (pi *
     -9680784138530540771532970536716373365672387749736451068072657/
     987341467154252308784874460722953477617652829467085267514961250 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     -9680784138530540771532970536716373365672387749736451068072657/
     987341467154252308784874460722953477617652829467085267514961250),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <
     8600841224988153445414906413408839182802664647802164997019218/
     4301658552533362102982352330928419591401332323901082498509609 |
     8600841224988153445414906413408839182802664647802164997019218/
     4301658552533362102982352330928419591401332323901082498509609 *
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <= pi * 672000000000000 |
     pi * 5000000000000000/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) *
     270000000000000000/7 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (X_000084 < sqrt(X_000083) | ~ lgen(1, X_000084, sqrt(X_000083))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_7, plain,
    (X_000083 < 0 | ~ lgen(1, X_000084, 2 * X_000083 / (X_000083 + 1)) |
     lgen(1, X_000084, sqrt(X_000083))),
    inference(subst, [], [sqrt_lower_bound_0])).

cnf(refute_0_8, plain,
    (X_000083 < 0 | X_000084 < sqrt(X_000083) |
     ~ lgen(1, X_000084, 2 * X_000083 / (X_000083 + 1))),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    (X_000083 < 0 | X_000084 < sqrt(X_000083) |
     X_000083 * 2 / (1 + X_000083) <= X_000084),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224000000000000000000000000000000 < 0 |
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) |
     4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224000000000000000000000000000000 * 2 /
     (1 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224000000000000000000000000000000) <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)), inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (45/14 <
     (pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) |
     0 <=
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * -6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      12000000000000000) |
     45/14 *
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <=
     pi * 672000000000000 +
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * -20000000000000000)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_12, plain,
    ((33600000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      -1000000000000000) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) | 0 <= pi * -20 |
     (33600000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      -1000000000000000) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_13, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain, (skoXC1 = 354991/512),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_15, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 354991/512 |
     skoXC1 ^ 2 = 354991/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain, (skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain, (29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_21])).

cnf(refute_0_23, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_24, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_29, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_14, refute_0_33])).

cnf(refute_0_35, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_36, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_39, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_34, refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_30, refute_0_41])).

cnf(refute_0_43, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 354991/512 |
     skoXC1 ^ 4 = 354991/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_14, refute_0_49])).

cnf(refute_0_51, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 354991/512 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_56, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_57, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_54, refute_0_57])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_61])).

cnf(refute_0_63, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_68, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_58, refute_0_68])).

cnf(refute_0_70, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_71, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_72, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_75, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_76, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_73, refute_0_76])).

cnf(refute_0_78, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_79, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_80, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_77, refute_0_80])).

cnf(refute_0_82, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_83, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_46, refute_0_83])).

cnf(refute_0_85, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_94, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_95, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_92, refute_0_95])).

cnf(refute_0_97, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [equality])).

cnf(refute_0_102, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_105])).

cnf(refute_0_107, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_106, refute_0_109])).

cnf(refute_0_111, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_113])).

cnf(refute_0_115, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 354991/512 |
     skoXC1 ^ 6 = 354991/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain, (skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_14, refute_0_121])).

cnf(refute_0_123, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_128, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_129, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_126, refute_0_129])).

cnf(refute_0_131, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 354991/512 |
     skoXC1 ^ 8 = 354991/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_14, refute_0_133])).

cnf(refute_0_135, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_138, refute_0_141])).

cnf(refute_0_143, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_144, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 354991/512 |
     skoXC1 ^ 10 = 354991/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_145, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain, (skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_14, refute_0_145])).

cnf(refute_0_147, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_149, plain,
    (skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_147, refute_0_148])).

cnf(refute_0_150, plain, (9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_146, refute_0_149])).

cnf(refute_0_151, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_150, refute_0_153])).

cnf(refute_0_155, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_156, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_142, refute_0_156])).

cnf(refute_0_158, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_163, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_130, refute_0_163])).

cnf(refute_0_165, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_164, refute_0_167])).

cnf(refute_0_169, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_170, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_118, refute_0_170])).

cnf(refute_0_172, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_173, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_174, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_171, refute_0_174])).

cnf(refute_0_176, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_177, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_110, refute_0_177])).

cnf(refute_0_179, plain,
    ((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 =
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30),
    introduced(tautology, [refl])).

cnf(refute_0_180, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 !=
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 |
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 =
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30),
    introduced(tautology, [equality])).

cnf(refute_0_181, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 =
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    ((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 =
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30),
    inference(resolve, [], [refute_0_178, refute_0_181])).

cnf(refute_0_183, plain,
    (sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)),
    introduced(tautology, [refl])).

cnf(refute_0_184, plain,
    ((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 !=
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30 |
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) |
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    introduced(tautology, [equality])).

cnf(refute_0_185, plain,
    ((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30 !=
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30 |
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    inference(resolve, [], [refute_0_183, refute_0_184])).

cnf(refute_0_186, plain,
    (sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    inference(resolve, [], [refute_0_182, refute_0_185])).

cnf(refute_0_187, plain,
    (10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)),
    introduced(tautology, [refl])).

cnf(refute_0_188, plain,
    (10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) |
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30) |
     10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     10 ^ 15 *
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30) |
     10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     10 ^ 15 *
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) =
     10 ^ 15 *
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)),
    inference(resolve, [], [refute_0_186, refute_0_189])).

cnf(refute_0_191, plain,
    (33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) =
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))),
    introduced(tautology, [refl])).

cnf(refute_0_192, plain,
    (10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     10 ^ 15 *
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30) |
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) !=
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) |
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) =
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
            10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
            4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))),
    introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (10 ^ 15 *
     sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
           10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
           4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30) !=
     10 ^ 15 *
     sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
           10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
           4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30) |
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) =
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
            10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
            4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) =
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
            10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
            4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))),
    inference(resolve, [], [refute_0_190, refute_0_193])).

cnf(refute_0_195, plain,
    (ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [refl])).

cnf(refute_0_196, plain,
    (33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) !=
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
            10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
            4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)) |
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) !=
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) |
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [equality])).

cnf(refute_0_197, plain,
    (33600000000000 /
     (10 ^ 15 *
      sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
            10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
            4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)) !=
     33600000000000 /
     (10 ^ 15 *
      sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
            10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
            4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)) |
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_195, refute_0_196])).

cnf(refute_0_198, plain,
    (ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_194, refute_0_197])).

cnf(refute_0_199, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [refl])).

cnf(refute_0_200, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) !=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) |
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) !=
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))) |
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [equality])).

cnf(refute_0_201, plain,
    (ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) !=
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))) |
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) =
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_198, refute_0_201])).

cnf(refute_0_203, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) !=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    introduced(tautology, [equality])).

cnf(refute_0_204, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain,
    (20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
               10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
               4000000000000000000000000 * skoXC1 ^ 2) / 10 ^ 30))) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_102, refute_0_204])).

cnf(refute_0_206, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     20 / ln(10) *
     ln(33600000000000 /
        (10 ^ 15 *
         sqrt((9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
               10300000000000000 * 354991/512 ^ 6 +
               10 ^ 22 * 354991/512 ^ 4 +
               4000000000000000000000000 * 354991/512 ^ 2) / 10 ^ 30)))),
    inference(resolve, [], [refute_0_13, refute_0_205])).

cnf(refute_0_207, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * -20 / ln(10) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     pi | pi = 0), inference(arithmetic, [], [refute_0_206])).

cnf(refute_0_208, plain,
    ((pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     pi <
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * -20 / ln(10) |
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * -20 / ln(10) * pi <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_209, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * -20 / ln(10) * pi <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0 | pi = 0),
    inference(resolve, [], [refute_0_207, refute_0_208])).

cnf(refute_0_210, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0 | pi = 0), inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_212, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_213, plain, (0 < pi | pi = 0),
    inference(decision, [], [refute_0_14, refute_0_211, refute_0_212])).

cnf(refute_0_214, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi = 0), inference(resolve, [], [refute_0_210, refute_0_213])).

cnf(refute_0_215, plain, (pi != 0),
    inference(decision, [], [refute_0_14, refute_0_211, refute_0_212])).

cnf(refute_0_216, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) / ln(10) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12),
    inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 <
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) / ln(10) |
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) *
     ln(10) | ln(10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_218, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) <=
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) *
     ln(10) | ln(10) <= 0),
    inference(resolve, [], [refute_0_216, refute_0_217])).

cnf(refute_0_219, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (0 <= pi * -20 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000))), inference(split, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) | pi <= 0),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain, (0 < pi),
    inference(decision, [], [refute_0_14, refute_0_211, refute_0_212])).

cnf(refute_0_223, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000))),
    inference(resolve, [], [refute_0_221, refute_0_222])).

cnf(refute_0_224, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_225, plain,
    (~ lgen(1, X_000048 - 1, X_000049) | X_000048 <= 0 |
     lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_226, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, X_000048 - 1, X_000049) |
     X_000048 <= 0), inference(resolve, [], [refute_0_225, refute_0_224])).

cnf(refute_0_227, plain,
    (ln(X_000048) < X_000049 | X_000048 <= 0 | X_000049 <= -1 + X_000048),
    inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain,
    (ln(33600000000000 /
        (sqrt(4300420612494076722707453206704419591401332323901082498509609/
              1237940039285380274899124224000000000000000000000000000000) *
         1000000000000000)) <
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) |
     33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     -1 +
     33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000)), inference(subst, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     -1 +
     33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000)),
    inference(resolve, [], [refute_0_223, refute_0_228])).

cnf(refute_0_230, plain,
    (33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <= 0 |
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     (33600000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      -1000000000000000) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000)), inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (0 <
     33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) |
     33600000000000 <=
     0 *
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     1000000000000000 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_232, plain,
    (0 <
     33600000000000 /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000)), inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    (ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -20) <=
     (33600000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      -1000000000000000) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000)),
    inference(resolve, [], [refute_0_230, refute_0_232])).

cnf(refute_0_234, plain,
    (0 <= pi * -20 |
     (33600000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      -1000000000000000) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) * (pi * -20) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(resolve, [], [refute_0_233, refute_0_12])).

cnf(refute_0_235, plain,
    ((pi * -672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * 20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) |
     pi <= 0), inference(arithmetic, [], [refute_0_234])).

cnf(refute_0_236, plain,
    ((pi * -672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * 20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      1000000000000000) <=
     ln(10) *
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12)),
    inference(resolve, [], [refute_0_235, refute_0_222])).

cnf(refute_0_237, plain,
    (0 <=
     pi * -6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 12 |
     -((pi * -672000000000000 +
       sqrt(4300420612494076722707453206704419591401332323901082498509609/
            1237940039285380274899124224000000000000000000000000000000) *
       (pi * 20000000000000000)) /
      (sqrt(4300420612494076722707453206704419591401332323901082498509609/
            1237940039285380274899124224000000000000000000000000000000) *
       1000000000000000)) /
     (pi * -6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) <=
     ln(10)), inference(split, [], [refute_0_236])).

cnf(refute_0_238, plain,
    (pi * 1/2 <=
     arctan(4131275018871243593760317/4131263554720769638778880) |
     (pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <= ln(10)),
    inference(arithmetic, [], [refute_0_237])).

cnf(refute_0_239, plain,
    (arctan(X_000041) < X_000042 | ~ lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_240, plain,
    (X_000041 < 0 | ~ lgen(1, X_000041, X_000042) |
     lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_241, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 |
     ~ lgen(1, X_000041, X_000042)),
    inference(resolve, [], [refute_0_240, refute_0_239])).

cnf(refute_0_242, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 | X_000042 <= X_000041),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (4131275018871243593760317/4131263554720769638778880 < 0 |
     arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2 |
     pi * 1/2 <= 4131275018871243593760317/4131263554720769638778880),
    inference(subst, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2 | pi <= 4131275018871243593760317/2065631777360384819389440),
    inference(arithmetic, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (4131275018871243593760317/2065631777360384819389440 < pi),
    inference(decision, [], [refute_0_212, refute_0_211, refute_0_14])).

cnf(refute_0_246, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <
     pi * 1/2), inference(resolve, [], [refute_0_244, refute_0_245])).

cnf(refute_0_247, plain,
    ((pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <= ln(10)),
    inference(resolve, [], [refute_0_238, refute_0_246])).

cnf(refute_0_248, plain,
    (ln(X_000118) < X_000119 | ~ lgen(1, ln(X_000118), X_000119)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_249, plain,
    (~ lgen(1, 1/2 * (X_000118 + 5) * (X_000118 - 1) / (2 * X_000118 + 1),
            X_000119) | X_000118 <= 0 | lgen(1, ln(X_000118), X_000119)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_250, plain,
    (ln(X_000118) < X_000119 |
     ~ lgen(1, 1/2 * (X_000118 + 5) * (X_000118 - 1) / (2 * X_000118 + 1),
            X_000119) | X_000118 <= 0),
    inference(resolve, [], [refute_0_249, refute_0_248])).

cnf(refute_0_251, plain,
    (ln(X_000118) < X_000119 | X_000118 <= 0 |
     X_000119 <=
     (-5/2 + X_000118 * (2 + X_000118 * 1/2)) / (1 + X_000118 * 2)),
    inference(arithmetic, [], [refute_0_250])).

cnf(refute_0_252, plain,
    (ln(10) <
     (pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) | 10 <= 0 |
     (pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <= (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(subst, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (10 <= 0 |
     (pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <= (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(resolve, [], [refute_0_247, refute_0_252])).

cnf(refute_0_254, plain,
    ((pi * 672000000000000 +
      sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -20000000000000000)) /
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <= 45/14),
    inference(arithmetic, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (0 <=
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * -6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      12000000000000000) |
     45/14 *
     (sqrt(4300420612494076722707453206704419591401332323901082498509609/
           1237940039285380274899124224000000000000000000000000000000) *
      (pi * -6000000000000000 +
       arctan(4131275018871243593760317/4131263554720769638778880) *
       12000000000000000)) <=
     pi * 672000000000000 +
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * -20000000000000000)),
    inference(resolve, [], [refute_0_254, refute_0_11])).

cnf(refute_0_256, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <= pi * 672000000000000 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_255])).

cnf(refute_0_257, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) <= 0), introduced(tautology, [assume])).

cnf(refute_0_258, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) |
     pi * 6000000000000000 +
     arctan(4131275018871243593760317/4131263554720769638778880) *
     -12000000000000000 <= 0 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) <=
     0 /
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000)), inference(split, [], [refute_0_257])).

cnf(refute_0_259, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) |
     pi * 1/2 <=
     arctan(4131275018871243593760317/4131263554720769638778880)),
    inference(arithmetic, [], [refute_0_258])).

cnf(refute_0_260, plain,
    (4131275018871243593760317/4131263554720769638778880 < 0 |
     0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) |
     pi * 1/2 <= 4131275018871243593760317/4131263554720769638778880),
    inference(resolve, [], [refute_0_259, refute_0_243])).

cnf(refute_0_261, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000) |
     pi <= 4131275018871243593760317/2065631777360384819389440),
    inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (0 <
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 6000000000000000 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      -12000000000000000)),
    inference(resolve, [], [refute_0_261, refute_0_245])).

cnf(refute_0_263, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) *
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <= pi * 672000000000000),
    inference(resolve, [], [refute_0_256, refute_0_262])).

cnf(refute_0_264, plain,
    (pi * 5000000000000000/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) *
     270000000000000000/7 <= 0 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)), inference(split, [], [refute_0_263])).

cnf(refute_0_265, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi * -1/54 |
     sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)), inference(arithmetic, [], [refute_0_264])).

cnf(refute_0_266, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi * -1/54), inference(subst, [], [refute_0_3])).

cnf(refute_0_267, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     pi <=
     -6366875216864375931832498519161/206563750943562179688015850000),
    inference(arithmetic, [], [refute_0_266])).

cnf(refute_0_268, plain,
    (-6366875216864375931832498519161/206563750943562179688015850000 < pi),
    inference(decision, [], [refute_0_14, refute_0_211, refute_0_212])).

cnf(refute_0_269, plain,
    (pi * -1/54 <
     arctan(4131275018871243593760317/4131263554720769638778880)),
    inference(resolve, [], [refute_0_267, refute_0_268])).

cnf(refute_0_270, plain,
    (sqrt(4300420612494076722707453206704419591401332323901082498509609/
          1237940039285380274899124224000000000000000000000000000000) <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)),
    inference(resolve, [], [refute_0_265, refute_0_269])).

cnf(refute_0_271, plain,
    (4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224000000000000000000000000000000 < 0 |
     4300420612494076722707453206704419591401332323901082498509609/
     1237940039285380274899124224000000000000000000000000000000 * 2 /
     (1 +
      4300420612494076722707453206704419591401332323901082498509609/
      1237940039285380274899124224000000000000000000000000000000) <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)),
    inference(resolve, [], [refute_0_270, refute_0_10])).

cnf(refute_0_272, plain,
    (8600841224988153445414906413408839182802664647802164997019218/
     4301658552533362102982352330928419591401332323901082498509609 <=
     pi * 672000000000000 /
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7)), inference(arithmetic, [], [refute_0_271])).

cnf(refute_0_273, plain,
    (8600841224988153445414906413408839182802664647802164997019218/
     4301658552533362102982352330928419591401332323901082498509609 *
     (pi * 5000000000000000/7 +
      arctan(4131275018871243593760317/4131263554720769638778880) *
      270000000000000000/7) <= pi * 672000000000000 |
     pi * 5000000000000000/7 +
     arctan(4131275018871243593760317/4131263554720769638778880) *
     270000000000000000/7 <= 0),
    inference(resolve, [], [refute_0_272, refute_0_5])).

cnf(refute_0_274, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi * -1/54 |
     arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     -9680784138530540771532970536716373365672387749736451068072657/
     987341467154252308784874460722953477617652829467085267514961250),
    inference(arithmetic, [], [refute_0_273])).

cnf(refute_0_275, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     -9680784138530540771532970536716373365672387749736451068072657/
     987341467154252308784874460722953477617652829467085267514961250),
    inference(resolve, [], [refute_0_274, refute_0_269])).

cnf(refute_0_276, plain,
    (4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     -9680784138530540771532970536716373365672387749736451068072657/
     987341467154252308784874460722953477617652829467085267514961250),
    inference(resolve, [], [refute_0_275, refute_0_4])).

cnf(refute_0_277, plain,
    (pi <=
     -186260145712797667393230202932602294740679812420114169825307374591821002574660697156660667/
     3199518533967695638618866366203592605865956751639834655499639083305240362916831948181520),
    inference(arithmetic, [], [refute_0_276])).

cnf(refute_0_278, plain,
    (-186260145712797667393230202932602294740679812420114169825307374591821002574660697156660667/
     3199518533967695638618866366203592605865956751639834655499639083305240362916831948181520 <
     pi),
    inference(decision, [], [refute_0_14, refute_0_211, refute_0_212])).

cnf(refute_0_279, plain, ($false),
    inference(resolve, [], [refute_0_277, refute_0_278])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-6-scaled.tptp
