% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-2-8.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(ln_upper_bound_cf3, axiom,
    (X <= 0 | ~ lgen(R, 1/2 * (X + 5) * (X - 1) / (2 * X + 1), Y) |
     lgen(R, ln(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_2_8', conjecture,
    (! [X] :
       (X = 63357/512 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6))).

cnf(sqrt_lower_bound_1_scaled, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            4000000000000000 * X * (X + 10 ^ 30) /
            (X ^ 2 + 6000000000000000000000000000000 * X + 10 ^ 60)) |
     lgen(R, Y, sqrt(X)))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 63357/512 =>
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) <
        12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_2_8'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 63357/512 =>
          20 / ln(10) *
          ln(36400000000000 /
             sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                  10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                  4000000000000000000000000 * X ^ 2)) <
          12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) + 6)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) + 6 <=
        20 / ln(10) *
        ln(36400000000000 /
           sqrt(9 * X ^ 10 + 309000000 * X ^ 8 +
                10300000000000000 * X ^ 6 + 10 ^ 22 * X ^ 4 +
                4000000000000000000000000 * X ^ 2)) & X = 63357/512)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) &
     skoXC1 = 63357/512), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 63357/512),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (X_000084 < arctan(X_000083) | ~ lgen(1, X_000084, arctan(X_000083))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (X_000083 < 0 |
     ~ lgen(1, X_000084, 3 * X_000083 / (X_000083 ^ 2 + 3)) |
     lgen(1, X_000084, arctan(X_000083))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_2, plain,
    (X_000083 < 0 | X_000084 < arctan(X_000083) |
     ~ lgen(1, X_000084, 3 * X_000083 / (X_000083 ^ 2 + 3))),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000083 < 0 | X_000084 < arctan(X_000083) |
     X_000083 * 3 / (3 + X_000083 * X_000083) <= X_000084),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (2920572702739709197/868167867892207981271040 < 0 |
     pi * -1/2 < arctan(2920572702739709197/868167867892207981271040) |
     2920572702739709197/868167867892207981271040 * 3 /
     (3 +
      2920572702739709197/868167867892207981271040 *
      2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (arctan(X_000041) < X_000042 | ~ lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_6, plain,
    (X_000041 < 0 | ~ lgen(1, X_000041, X_000042) |
     lgen(1, arctan(X_000041), X_000042)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_7, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 |
     ~ lgen(1, X_000041, X_000042)),
    inference(resolve, [], [refute_0_6, refute_0_5])).

cnf(refute_0_8, plain,
    (arctan(X_000041) < X_000042 | X_000041 < 0 | X_000042 <= X_000041),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (2920572702739709197/868167867892207981271040 < 0 |
     arctan(2920572702739709197/868167867892207981271040) < pi * 1/54 |
     pi * 1/54 <= 2920572702739709197/868167867892207981271040),
    inference(subst, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <
     51561165713647475739063467591501125371488356164628726240482021526727511561246188233151554429011342060550564313832324000000000000000/
     33142296319886949630441422032958678236884083403094396718605631862533647554390335748207888607252835515137641078458081 |
     51561165713647475739063467591501125371488356164628726240482021526727511561246188233151554429011342060550564313832324000000000000000/
     33142296319886949630441422032958678236884083403094396718605631862533647554390335748207888607252835515137641078458081 *
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <=
     pi * 728000000000000 |
     pi * 5/7 +
     arctan(2920572702739709197/868167867892207981271040) * -270/7 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_11, plain,
    (45/14 <
     (pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) |
     0 <=
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * -6 +
      arctan(2920572702739709197/868167867892207981271040) * -12) |
     45/14 *
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     pi * 728000000000000 +
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) * (pi * -20)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_12, plain,
    ((36400000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * -1) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) | 0 <= pi * -20 |
     (36400000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * -1) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12)),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_13, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain, (skoXC1 = 63357/512),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_15, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 63357/512 |
     skoXC1 ^ 2 = 63357/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_17, plain, (skoXC1 != 63357/512 | skoXC1 ^ 2 = 63357/512 ^ 2),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain, (skoXC1 ^ 2 = 63357/512 ^ 2),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 63357/512 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 63357/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (skoXC1 ^ 2 != 63357/512 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain, (29 * skoXC1 ^ 2 = 29 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_21])).

cnf(refute_0_23, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_24, plain,
    (29 * skoXC1 ^ 2 != 29 * 63357/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 63357/512 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain,
    (29 * skoXC1 ^ 2 != 29 * 63357/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 63357/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 63357/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_22, refute_0_25])).

cnf(refute_0_27, plain,
    (63357/512 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_28, plain,
    (63357/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 63357/512 ^ 2 - 90000000 |
     63357/512 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_29, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 63357/512 ^ 2 - 90000000 |
     63357/512 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (63357/512 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_26, refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 63357/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 != 63357/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_14, refute_0_33])).

cnf(refute_0_35, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_36, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_37, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_39, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (63357/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * 63357/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (63357/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * 63357/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_34, refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     63357/512 * (29 * 63357/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_30, refute_0_41])).

cnf(refute_0_43, plain,
    (1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * 63357/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     63357/512 * (29 * 63357/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 63357/512 |
     skoXC1 ^ 4 = 63357/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_49, plain, (skoXC1 != 63357/512 | skoXC1 ^ 4 = 63357/512 ^ 4),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (skoXC1 ^ 4 = 63357/512 ^ 4),
    inference(resolve, [], [refute_0_14, refute_0_49])).

cnf(refute_0_51, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 63357/512 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 63357/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (skoXC1 ^ 4 != 63357/512 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (3 * skoXC1 ^ 4 = 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain,
    (-130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_56, plain,
    (3 * skoXC1 ^ 4 != 3 * 63357/512 ^ 4 |
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_57, plain,
    (3 * skoXC1 ^ 4 != 3 * 63357/512 ^ 4 |
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (-130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_54, refute_0_57])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 63357/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 63357/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_61, plain,
    (skoXC1 ^ 2 != 63357/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_61])).

cnf(refute_0_63, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_64, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 63357/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 63357/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_62, refute_0_65])).

cnf(refute_0_67, plain,
    (-130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_68, plain,
    (-130660000 * 63357/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_58, refute_0_68])).

cnf(refute_0_70, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_71, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_72, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_75, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_76, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_73, refute_0_76])).

cnf(refute_0_78, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_79, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_80, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_77, refute_0_80])).

cnf(refute_0_82, plain,
    (1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_83, plain,
    (1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_46, refute_0_83])).

cnf(refute_0_85, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 + 2000000000000) /
     (63357/512 * (29 * 63357/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_90, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_91, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_89, refute_0_90])).

cnf(refute_0_92, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_88, refute_0_91])).

cnf(refute_0_93, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_94, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_95, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
             2000000000000) /
            (63357/512 * (29 * 63357/512 ^ 2 - 90000000))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_92, refute_0_95])).

cnf(refute_0_97, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 =
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 !=
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_102, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 | skoXC1 ^ 2 != 63357/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 63357/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (skoXC1 ^ 2 != 63357/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_18, refute_0_105])).

cnf(refute_0_107, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 63357/512 ^ 2 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 63357/512 ^ 2 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_106, refute_0_109])).

cnf(refute_0_111, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 63357/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 63357/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoXC1 ^ 4 != 63357/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_113])).

cnf(refute_0_115, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 63357/512 ^ 4 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 63357/512 ^ 4 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 63357/512 |
     skoXC1 ^ 6 = 63357/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (skoXC1 != 63357/512 | skoXC1 ^ 6 = 63357/512 ^ 6),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain, (skoXC1 ^ 6 = 63357/512 ^ 6),
    inference(resolve, [], [refute_0_14, refute_0_121])).

cnf(refute_0_123, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 63357/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 63357/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (skoXC1 ^ 6 != 63357/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_128, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 63357/512 ^ 6 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_129, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 63357/512 ^ 6 |
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_126, refute_0_129])).

cnf(refute_0_131, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 63357/512 |
     skoXC1 ^ 8 = 63357/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (skoXC1 != 63357/512 | skoXC1 ^ 8 = 63357/512 ^ 8),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 ^ 8 = 63357/512 ^ 8),
    inference(resolve, [], [refute_0_14, refute_0_133])).

cnf(refute_0_135, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 63357/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 63357/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (skoXC1 ^ 8 != 63357/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain,
    (9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 63357/512 ^ 8 |
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 63357/512 ^ 8 |
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_138, refute_0_141])).

cnf(refute_0_143, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_144, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 63357/512 |
     skoXC1 ^ 10 = 63357/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_145, plain,
    (skoXC1 != 63357/512 | skoXC1 ^ 10 = 63357/512 ^ 10),
    inference(resolve, [], [refute_0_143, refute_0_144])).

cnf(refute_0_146, plain, (skoXC1 ^ 10 = 63357/512 ^ 10),
    inference(resolve, [], [refute_0_14, refute_0_145])).

cnf(refute_0_147, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 63357/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 63357/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_149, plain,
    (skoXC1 ^ 10 != 63357/512 ^ 10 | 9 * skoXC1 ^ 10 = 9 * 63357/512 ^ 10),
    inference(resolve, [], [refute_0_147, refute_0_148])).

cnf(refute_0_150, plain, (9 * skoXC1 ^ 10 = 9 * 63357/512 ^ 10),
    inference(resolve, [], [refute_0_146, refute_0_149])).

cnf(refute_0_151, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 != 9 * 63357/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 != 9 * 63357/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_150, refute_0_153])).

cnf(refute_0_155, plain,
    (9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_156, plain,
    (9 * 63357/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8),
    inference(resolve, [], [refute_0_142, refute_0_156])).

cnf(refute_0_158, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_163, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6),
    inference(resolve, [], [refute_0_130, refute_0_163])).

cnf(refute_0_165, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_164, refute_0_167])).

cnf(refute_0_169, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_170, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4),
    inference(resolve, [], [refute_0_118, refute_0_170])).

cnf(refute_0_172, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_173, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_174, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_171, refute_0_174])).

cnf(refute_0_176, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_177, plain,
    (9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2),
    inference(resolve, [], [refute_0_110, refute_0_177])).

cnf(refute_0_179, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_180, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_181, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
     10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
     4000000000000000000000000 * 63357/512 ^ 2 |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    inference(resolve, [], [refute_0_178, refute_0_181])).

cnf(refute_0_183, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_184, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) |
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_185, plain,
    (sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2) |
     36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    inference(resolve, [], [refute_0_183, refute_0_184])).

cnf(refute_0_186, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) =
     36400000000000 /
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2)),
    inference(resolve, [], [refute_0_182, refute_0_185])).

cnf(refute_0_187, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_188, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (36400000000000 /
     sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2) !=
     36400000000000 /
     sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
          10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
          4000000000000000000000000 * 63357/512 ^ 2) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_186, refute_0_189])).

cnf(refute_0_191, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_192, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_193, plain,
    (ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2)) |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) =
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_190, refute_0_193])).

cnf(refute_0_195, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 |
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) !=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2)) |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_196, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
             10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
             4000000000000000000000000 * skoXC1 ^ 2)) <
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) + 6 |
     12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_102, refute_0_196])).

cnf(refute_0_198, plain,
    (12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 63357/512 ^ 2 + 3 * 63357/512 ^ 4 +
              2000000000000) /
             (63357/512 * (29 * 63357/512 ^ 2 - 90000000)))) + 6 <=
     20 / ln(10) *
     ln(36400000000000 /
        sqrt(9 * 63357/512 ^ 10 + 309000000 * 63357/512 ^ 8 +
             10300000000000000 * 63357/512 ^ 6 + 10 ^ 22 * 63357/512 ^ 4 +
             4000000000000000000000000 * 63357/512 ^ 2))),
    inference(resolve, [], [refute_0_13, refute_0_197])).

cnf(refute_0_199, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * -20 / ln(10) <=
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     pi | pi = 0), inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    ((pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     pi <
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * -20 / ln(10) |
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 |
     pi <= 0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_201, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * -20 / ln(10) * pi <=
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 |
     pi <= 0 | pi = 0),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 |
     pi <= 0 | pi = 0), inference(arithmetic, [], [refute_0_201])).

cnf(refute_0_203, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_204, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_205, plain, (0 < pi | pi = 0),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_206, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 |
     pi = 0), inference(resolve, [], [refute_0_202, refute_0_205])).

cnf(refute_0_207, plain, (pi != 0),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_208, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) / ln(10) <=
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12),
    inference(resolve, [], [refute_0_206, refute_0_207])).

cnf(refute_0_209, plain,
    (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 <
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) / ln(10) |
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) *
     ln(10) | ln(10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_210, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) <=
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) *
     ln(10) | ln(10) <= 0),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12)),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (0 <= pi * -20 |
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224))),
    inference(split, [], [refute_0_211])).

cnf(refute_0_213, plain,
    (ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) | pi <= 0),
    inference(arithmetic, [], [refute_0_212])).

cnf(refute_0_214, plain, (0 < pi),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_215, plain,
    (ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224))),
    inference(resolve, [], [refute_0_213, refute_0_214])).

cnf(refute_0_216, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_217, plain,
    (~ lgen(1, X_000048 - 1, X_000049) | X_000048 <= 0 |
     lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_218, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, X_000048 - 1, X_000049) |
     X_000048 <= 0), inference(resolve, [], [refute_0_217, refute_0_216])).

cnf(refute_0_219, plain,
    (ln(X_000048) < X_000049 | X_000048 <= 0 | X_000049 <= -1 + X_000048),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (ln(36400000000000 /
        sqrt(3024299845962703243695009305677455293181862232440251200241/
             1237940039285380274899124224)) <
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) |
     36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224)),
    inference(subst, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     -1 +
     36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_215, refute_0_220])).

cnf(refute_0_222, plain,
    (36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <= 0 |
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * -1) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (0 <
     36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) |
     36400000000000 <=
     0 *
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_224, plain,
    (0 <
     36400000000000 /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224)),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) /
     (pi * -20) <=
     (36400000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * -1) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224)),
    inference(resolve, [], [refute_0_222, refute_0_224])).

cnf(refute_0_226, plain,
    (0 <= pi * -20 |
     (36400000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * -1) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) * (pi * -20) <=
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12)),
    inference(resolve, [], [refute_0_225, refute_0_12])).

cnf(refute_0_227, plain,
    ((pi * -728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * 20)) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) |
     pi <= 0), inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain,
    ((pi * -728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * 20)) /
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     ln(10) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12)),
    inference(resolve, [], [refute_0_227, refute_0_214])).

cnf(refute_0_229, plain,
    (0 <=
     pi * -6 + arctan(2920572702739709197/868167867892207981271040) * -12 |
     -((pi * -728000000000000 +
       sqrt(3024299845962703243695009305677455293181862232440251200241/
            1237940039285380274899124224) * (pi * 20)) /
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224)) /
     (pi * -6 +
      arctan(2920572702739709197/868167867892207981271040) * -12) <=
     ln(10)), inference(split, [], [refute_0_228])).

cnf(refute_0_230, plain,
    ((pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     ln(10) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (ln(X_000111) < X_000112 | ~ lgen(1, ln(X_000111), X_000112)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_232, plain,
    (~ lgen(1, 1/2 * (X_000111 + 5) * (X_000111 - 1) / (2 * X_000111 + 1),
            X_000112) | X_000111 <= 0 | lgen(1, ln(X_000111), X_000112)),
    inference(subst, [], [ln_upper_bound_cf3])).

cnf(refute_0_233, plain,
    (ln(X_000111) < X_000112 |
     ~ lgen(1, 1/2 * (X_000111 + 5) * (X_000111 - 1) / (2 * X_000111 + 1),
            X_000112) | X_000111 <= 0),
    inference(resolve, [], [refute_0_232, refute_0_231])).

cnf(refute_0_234, plain,
    (ln(X_000111) < X_000112 | X_000111 <= 0 |
     X_000112 <=
     (-5/2 + X_000111 * (2 + X_000111 * 1/2)) / (1 + X_000111 * 2)),
    inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (ln(10) <
     (pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) |
     10 <= 0 |
     (pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2)),
    inference(subst, [], [refute_0_234])).

cnf(refute_0_236, plain,
    (10 <= 0 |
     (pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     (-5/2 + 10 * (2 + 10 * 1/2)) / (1 + 10 * 2) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_230, refute_0_235])).

cnf(refute_0_237, plain,
    ((pi * 728000000000000 +
      sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) * (pi * -20)) /
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     45/14 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_236])).

cnf(refute_0_238, plain,
    (0 <=
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * -6 +
      arctan(2920572702739709197/868167867892207981271040) * -12) |
     45/14 *
     (sqrt(3024299845962703243695009305677455293181862232440251200241/
           1237940039285380274899124224) *
      (pi * -6 +
       arctan(2920572702739709197/868167867892207981271040) * -12)) <=
     pi * 728000000000000 +
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) * (pi * -20) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_237, refute_0_11])).

cnf(refute_0_239, plain,
    (sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <=
     pi * 728000000000000 |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(2920572702739709197/868167867892207981271040) * 12) <= 0 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_238])).

cnf(refute_0_240, plain,
    (0 <
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 6 +
      arctan(2920572702739709197/868167867892207981271040) * 12) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_241, plain,
    (0 <
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) |
     pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12 <=
     0 |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     0 /
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12)),
    inference(split, [], [refute_0_240])).

cnf(refute_0_242, plain,
    (0 <
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 6 + arctan(2920572702739709197/868167867892207981271040) * 12) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) *
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <=
     pi * 728000000000000 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_239, refute_0_242])).

cnf(refute_0_244, plain,
    (pi * 5/7 +
     arctan(2920572702739709197/868167867892207981271040) * -270/7 <= 0 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2 |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7)),
    inference(split, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (pi * 1/54 <= arctan(2920572702739709197/868167867892207981271040) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2 |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7)),
    inference(arithmetic, [], [refute_0_244])).

cnf(refute_0_246, plain,
    (arctan(2920572702739709197/868167867892207981271040) < pi * 1/54 |
     pi <= 2920572702739709197/16077182738744592245760),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_247, plain,
    (2920572702739709197/16077182738744592245760 < pi),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_248, plain,
    (arctan(2920572702739709197/868167867892207981271040) < pi * 1/54),
    inference(resolve, [], [refute_0_246, refute_0_247])).

cnf(refute_0_249, plain,
    (arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2 |
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7)),
    inference(resolve, [], [refute_0_245, refute_0_248])).

cnf(refute_0_250, plain,
    (X_000126 < sqrt(X_000125) | ~ lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_251, plain,
    (X_000125 < 0 |
     ~ lgen(1, X_000126,
            4000000000000000 * X_000125 * (X_000125 + 10 ^ 30) /
            (X_000125 ^ 2 + 6000000000000000000000000000000 * X_000125 +
             10 ^ 60)) | lgen(1, X_000126, sqrt(X_000125))),
    inference(subst, [], [sqrt_lower_bound_1_scaled])).

cnf(refute_0_252, plain,
    (X_000125 < 0 | X_000126 < sqrt(X_000125) |
     ~ lgen(1, X_000126,
            4000000000000000 * X_000125 * (X_000125 + 10 ^ 30) /
            (X_000125 ^ 2 + 6000000000000000000000000000000 * X_000125 +
             10 ^ 60))),
    inference(resolve, [], [refute_0_251, refute_0_250])).

cnf(refute_0_253, plain,
    (X_000125 < 0 | X_000126 < sqrt(X_000125) |
     X_000125 *
     (4000000000000000000000000000000000000000000000 +
      X_000125 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      X_000125 * (6000000000000000000000000000000 + X_000125)) <=
     X_000126), inference(arithmetic, [], [refute_0_252])).

cnf(refute_0_254, plain,
    (3024299845962703243695009305677455293181862232440251200241/
     1237940039285380274899124224 < 0 |
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <
     sqrt(3024299845962703243695009305677455293181862232440251200241/
          1237940039285380274899124224) |
     3024299845962703243695009305677455293181862232440251200241/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000000 +
      3024299845962703243695009305677455293181862232440251200241/
      1237940039285380274899124224 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      3024299845962703243695009305677455293181862232440251200241/
      1237940039285380274899124224 *
      (6000000000000000000000000000000 +
       3024299845962703243695009305677455293181862232440251200241/
       1237940039285380274899124224)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7)),
    inference(subst, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (3024299845962703243695009305677455293181862232440251200241/
     1237940039285380274899124224 < 0 |
     3024299845962703243695009305677455293181862232440251200241/
     1237940039285380274899124224 *
     (4000000000000000000000000000000000000000000000 +
      3024299845962703243695009305677455293181862232440251200241/
      1237940039285380274899124224 * 4000000000000000) /
     (1000000000000000000000000000000000000000000000000000000000000 +
      3024299845962703243695009305677455293181862232440251200241/
      1237940039285380274899124224 *
      (6000000000000000000000000000000 +
       3024299845962703243695009305677455293181862232440251200241/
       1237940039285380274899124224)) <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_249, refute_0_254])).

cnf(refute_0_256, plain,
    (51561165713647475739063467591501125371488356164628726240482021526727511561246188233151554429011342060550564313832324000000000000000/
     33142296319886949630441422032958678236884083403094396718605631862533647554390335748207888607252835515137641078458081 <=
     pi * 728000000000000 /
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_255])).

cnf(refute_0_257, plain,
    (51561165713647475739063467591501125371488356164628726240482021526727511561246188233151554429011342060550564313832324000000000000000/
     33142296319886949630441422032958678236884083403094396718605631862533647554390335748207888607252835515137641078458081 *
     (pi * 5/7 +
      arctan(2920572702739709197/868167867892207981271040) * -270/7) <=
     pi * 728000000000000 |
     pi * 5/7 +
     arctan(2920572702739709197/868167867892207981271040) * -270/7 <= 0 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_256, refute_0_10])).

cnf(refute_0_258, plain,
    (pi *
     226818077862483375965785334891704598373164519900445371235703590974913494563922934165536662761470052340845414880967447/
     35514068221134740942722286351289040434443510623596316543189147480143949289633854140180917591410873358032276440649815000 <=
     arctan(2920572702739709197/868167867892207981271040) |
     pi * 1/54 <= arctan(2920572702739709197/868167867892207981271040) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain,
    (2920572702739709197/868167867892207981271040 < 0 |
     arctan(2920572702739709197/868167867892207981271040) <
     pi *
     226818077862483375965785334891704598373164519900445371235703590974913494563922934165536662761470052340845414880967447/
     35514068221134740942722286351289040434443510623596316543189147480143949289633854140180917591410873358032276440649815000 |
     pi *
     226818077862483375965785334891704598373164519900445371235703590974913494563922934165536662761470052340845414880967447/
     35514068221134740942722286351289040434443510623596316543189147480143949289633854140180917591410873358032276440649815000 <=
     2920572702739709197/868167867892207981271040),
    inference(subst, [], [refute_0_8])).

cnf(refute_0_260, plain,
    (arctan(2920572702739709197/868167867892207981271040) <
     pi *
     226818077862483375965785334891704598373164519900445371235703590974913494563922934165536662761470052340845414880967447/
     35514068221134740942722286351289040434443510623596316543189147480143949289633854140180917591410873358032276440649815000 |
     pi <=
     1515828561869013666051330127749719248706565407765535854569997362323096302426005430577489580338612331148095101248982420258529456125/
     2877815936870388948232777042864634357521151186983830684323184756774374788061355447280929809240272828826629374751166403153201982766848),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (1515828561869013666051330127749719248706565407765535854569997362323096302426005430577489580338612331148095101248982420258529456125/
     2877815936870388948232777042864634357521151186983830684323184756774374788061355447280929809240272828826629374751166403153201982766848 <
     pi),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_262, plain,
    (arctan(2920572702739709197/868167867892207981271040) <
     pi *
     226818077862483375965785334891704598373164519900445371235703590974913494563922934165536662761470052340845414880967447/
     35514068221134740942722286351289040434443510623596316543189147480143949289633854140180917591410873358032276440649815000),
    inference(resolve, [], [refute_0_260, refute_0_261])).

cnf(refute_0_263, plain,
    (pi * 1/54 <= arctan(2920572702739709197/868167867892207981271040) |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_258, refute_0_262])).

cnf(refute_0_264, plain,
    (2920572702739709197/868167867892207981271040 < 0 |
     pi * 1/54 <= 2920572702739709197/868167867892207981271040 |
     arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_263, refute_0_9])).

cnf(refute_0_265, plain,
    (arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2 |
     pi <= 2920572702739709197/16077182738744592245760),
    inference(arithmetic, [], [refute_0_264])).

cnf(refute_0_266, plain,
    (arctan(2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_265, refute_0_247])).

cnf(refute_0_267, plain,
    (2920572702739709197/868167867892207981271040 < 0 |
     2920572702739709197/868167867892207981271040 * 3 /
     (3 +
      2920572702739709197/868167867892207981271040 *
      2920572702739709197/868167867892207981271040) <= pi * -1/2),
    inference(resolve, [], [refute_0_266, refute_0_4])).

cnf(refute_0_268, plain,
    (pi <=
     -15213284258170299990158843613198432226529280/
     2261146340530036614422822049342213640525954429609),
    inference(arithmetic, [], [refute_0_267])).

cnf(refute_0_269, plain,
    (-15213284258170299990158843613198432226529280/
     2261146340530036614422822049342213640525954429609 < pi),
    inference(decision, [], [refute_0_14, refute_0_203, refute_0_204])).

cnf(refute_0_270, plain, ($false),
    inference(resolve, [], [refute_0_268, refute_0_269])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-2-8.tptp
