% SZS output start CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-6-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(atan_lower_bound_case_14, axiom,
    (X <= 1 | ~ lgen(R, Y, 157079/100000 - 1 / X) |
     lgen(R, Y, arctan(X)))).

cnf(ln_upper_bound_cf1, axiom,
    (X <= 0 | ~ lgen(R, X - 1, Y) | lgen(R, ln(X), Y))).

cnf(log_def, axiom, (log(X) = 868589/2000000 * ln(X))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof('Nichols_Plot_Disk_Drive_Meets_1_6', conjecture,
    (! [X] :
       (X = 354991/512 =>
        10 *
        log(33600000000000 ^ 2 /
            (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
             10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18))).

fof(subgoal_0, plain,
    (! [X] :
       (X = 354991/512 =>
        10 *
        log(33600000000000 ^ 2 /
            (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
             10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
        -12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(strip, [], ['Nichols_Plot_Disk_Drive_Meets_1_6'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X = 354991/512 =>
          10 *
          log(33600000000000 ^ 2 /
              (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
               10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) <
          -12 / pi *
          (-pi -
           arctan(1/1140 *
                  (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                  (X * (29 * X ^ 2 - 90000000)))) - 18)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (-12 / pi *
        (-pi -
         arctan(1/1140 * (-130660000 * X ^ 2 + 3 * X ^ 4 + 2000000000000) /
                (X * (29 * X ^ 2 - 90000000)))) - 18 <=
        10 *
        log(33600000000000 ^ 2 /
            (9 * X ^ 10 + 309000000 * X ^ 8 + 10300000000000000 * X ^ 6 +
             10 ^ 22 * X ^ 4 + 4000000000000000000000000 * X ^ 2)) &
        X = 354991/512)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) & skoXC1 = 354991/512),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 = 354991/512),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (-29254552482376122252275769092443099757560083836061785704147/
     29254561989755623963996280317717140077560083836061785704147 <
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) | 0 <= pi * -868589/200000 |
     -29254552482376122252275769092443099757560083836061785704147/
     29254561989755623963996280317717140077560083836061785704147 *
     (pi * -868589/200000) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 <=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain, (skoXC1 = 354991/512),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_3, plain, (skoXC1 ^ 2 = skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_4, plain,
    (skoXC1 ^ 2 != skoXC1 ^ 2 | skoXC1 != 354991/512 |
     skoXC1 ^ 2 = 354991/512 ^ 2), introduced(tautology, [equality])).

cnf(refute_0_5, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_3, refute_0_4])).

cnf(refute_0_6, plain, (skoXC1 ^ 2 = 354991/512 ^ 2),
    inference(resolve, [], [refute_0_2, refute_0_5])).

cnf(refute_0_7, plain, (29 * skoXC1 ^ 2 = 29 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_8, plain,
    (29 * skoXC1 ^ 2 != 29 * skoXC1 ^ 2 | skoXC1 ^ 2 != 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_9, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 | 29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_7, refute_0_8])).

cnf(refute_0_10, plain, (29 * skoXC1 ^ 2 = 29 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_6, refute_0_9])).

cnf(refute_0_11, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * skoXC1 ^ 2 - 90000000),
    introduced(tautology, [refl])).

cnf(refute_0_12, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 != 29 * skoXC1 ^ 2 - 90000000 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    introduced(tautology, [equality])).

cnf(refute_0_13, plain,
    (29 * skoXC1 ^ 2 != 29 * 354991/512 ^ 2 |
     29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_11, refute_0_12])).

cnf(refute_0_14, plain,
    (29 * skoXC1 ^ 2 - 90000000 = 29 * 354991/512 ^ 2 - 90000000),
    inference(resolve, [], [refute_0_10, refute_0_13])).

cnf(refute_0_15, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_16, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (29 * skoXC1 ^ 2 - 90000000 != 29 * 354991/512 ^ 2 - 90000000 |
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_14, refute_0_17])).

cnf(refute_0_19, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [refl])).

cnf(refute_0_20, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) | skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    introduced(tautology, [equality])).

cnf(refute_0_21, plain,
    (skoXC1 != 354991/512 |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_19, refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * skoXC1 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_2, refute_0_21])).

cnf(refute_0_23, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_24, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_25, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_27, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_25, refute_0_26])).

cnf(refute_0_28, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * skoXC1 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (354991/512 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_22, refute_0_28])).

cnf(refute_0_30, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) =
     354991/512 * (29 * 354991/512 ^ 2 - 90000000)),
    inference(resolve, [], [refute_0_18, refute_0_29])).

cnf(refute_0_31, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_32, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_33, plain,
    (skoXC1 * (29 * skoXC1 ^ 2 - 90000000) !=
     354991/512 * (29 * 354991/512 ^ 2 - 90000000) |
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_30, refute_0_33])).

cnf(refute_0_35, plain, (skoXC1 ^ 4 = skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_36, plain,
    (skoXC1 ^ 4 != skoXC1 ^ 4 | skoXC1 != 354991/512 |
     skoXC1 ^ 4 = 354991/512 ^ 4), introduced(tautology, [equality])).

cnf(refute_0_37, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_35, refute_0_36])).

cnf(refute_0_38, plain, (skoXC1 ^ 4 = 354991/512 ^ 4),
    inference(resolve, [], [refute_0_2, refute_0_37])).

cnf(refute_0_39, plain, (3 * skoXC1 ^ 4 = 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_40, plain,
    (3 * skoXC1 ^ 4 != 3 * skoXC1 ^ 4 | skoXC1 ^ 4 != 354991/512 ^ 4 |
     3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_41, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 | 3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_39, refute_0_40])).

cnf(refute_0_42, plain, (3 * skoXC1 ^ 4 = 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_38, refute_0_41])).

cnf(refute_0_43, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_44, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_45, plain,
    (3 * skoXC1 ^ 4 != 3 * 354991/512 ^ 4 |
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_43, refute_0_44])).

cnf(refute_0_46, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_42, refute_0_45])).

cnf(refute_0_47, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (-130660000 * skoXC1 ^ 2 = -130660000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_6, refute_0_49])).

cnf(refute_0_51, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (-130660000 * skoXC1 ^ 2 != -130660000 * 354991/512 ^ 2 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_56, plain,
    (-130660000 * 354991/512 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_46, refute_0_56])).

cnf(refute_0_58, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000),
    introduced(tautology, [refl])).

cnf(refute_0_59, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    introduced(tautology, [equality])).

cnf(refute_0_60, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 =
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000),
    inference(resolve, [], [refute_0_57, refute_0_60])).

cnf(refute_0_62, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000)),
    introduced(tautology, [refl])).

cnf(refute_0_63, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) |
     -130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    introduced(tautology, [equality])).

cnf(refute_0_64, plain,
    (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000 !=
     -130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000 |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000)),
    inference(resolve, [], [refute_0_61, refute_0_64])).

cnf(refute_0_66, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [refl])).

cnf(refute_0_67, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    introduced(tautology, [equality])).

cnf(refute_0_68, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_65, refute_0_68])).

cnf(refute_0_70, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_71, plain,
    (1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) =
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000))),
    inference(resolve, [], [refute_0_34, refute_0_71])).

cnf(refute_0_73, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_74, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_75, plain,
    (1/1140 * (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
     (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)) !=
     1/1140 *
     (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 + 2000000000000) /
     (354991/512 * (29 * 354991/512 ^ 2 - 90000000)) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_72, refute_0_75])).

cnf(refute_0_77, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))),
    introduced(tautology, [refl])).

cnf(refute_0_78, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) |
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    introduced(tautology, [equality])).

cnf(refute_0_79, plain,
    (arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_77, refute_0_78])).

cnf(refute_0_80, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) =
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))),
    inference(resolve, [], [refute_0_76, refute_0_79])).

cnf(refute_0_81, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))))),
    introduced(tautology, [refl])).

cnf(refute_0_82, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) |
     -pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    introduced(tautology, [equality])).

cnf(refute_0_83, plain,
    (-pi -
     arctan(1/1140 *
            (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
            (skoXC1 * (29 * skoXC1 ^ 2 - 90000000))) !=
     -pi -
     arctan(1/1140 *
            (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
             2000000000000) /
            (354991/512 * (29 * 354991/512 ^ 2 - 90000000))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000))))),
    inference(resolve, [], [refute_0_80, refute_0_83])).

cnf(refute_0_85, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [refl])).

cnf(refute_0_86, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    introduced(tautology, [equality])).

cnf(refute_0_87, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 =
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18),
    inference(resolve, [], [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 !=
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_90, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2))),
    inference(resolve, [], [refute_0_88, refute_0_89])).

cnf(refute_0_91, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_92, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * skoXC1 ^ 2 |
     skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (skoXC1 ^ 2 != 354991/512 ^ 2 |
     4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 =
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_6, refute_0_93])).

cnf(refute_0_95, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_96, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_97, plain,
    (4000000000000000000000000 * skoXC1 ^ 2 !=
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_95, refute_0_96])).

cnf(refute_0_98, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_94, refute_0_97])).

cnf(refute_0_99, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_100, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * skoXC1 ^ 4 |
     skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_101, plain,
    (skoXC1 ^ 4 != 354991/512 ^ 4 |
     10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain, (10 ^ 22 * skoXC1 ^ 4 = 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_38, refute_0_101])).

cnf(refute_0_103, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_104, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_105, plain,
    (10 ^ 22 * skoXC1 ^ 4 != 10 ^ 22 * 354991/512 ^ 4 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_102, refute_0_105])).

cnf(refute_0_107, plain, (skoXC1 ^ 6 = skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (skoXC1 ^ 6 != skoXC1 ^ 6 | skoXC1 != 354991/512 |
     skoXC1 ^ 6 = 354991/512 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain, (skoXC1 ^ 6 = 354991/512 ^ 6),
    inference(resolve, [], [refute_0_2, refute_0_109])).

cnf(refute_0_111, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * skoXC1 ^ 6),
    introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * skoXC1 ^ 6 |
     skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoXC1 ^ 6 != 354991/512 ^ 6 |
     10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (10300000000000000 * skoXC1 ^ 6 = 10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_110, refute_0_113])).

cnf(refute_0_115, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_116, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    introduced(tautology, [equality])).

cnf(refute_0_117, plain,
    (10300000000000000 * skoXC1 ^ 6 != 10300000000000000 * 354991/512 ^ 6 |
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_114, refute_0_117])).

cnf(refute_0_119, plain, (skoXC1 ^ 8 = skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_120, plain,
    (skoXC1 ^ 8 != skoXC1 ^ 8 | skoXC1 != 354991/512 |
     skoXC1 ^ 8 = 354991/512 ^ 8), introduced(tautology, [equality])).

cnf(refute_0_121, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain, (skoXC1 ^ 8 = 354991/512 ^ 8),
    inference(resolve, [], [refute_0_2, refute_0_121])).

cnf(refute_0_123, plain, (309000000 * skoXC1 ^ 8 = 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_124, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * skoXC1 ^ 8 |
     skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_125, plain,
    (skoXC1 ^ 8 != 354991/512 ^ 8 |
     309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_123, refute_0_124])).

cnf(refute_0_126, plain,
    (309000000 * skoXC1 ^ 8 = 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_122, refute_0_125])).

cnf(refute_0_127, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_128, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_129, plain,
    (309000000 * skoXC1 ^ 8 != 309000000 * 354991/512 ^ 8 |
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_126, refute_0_129])).

cnf(refute_0_131, plain, (skoXC1 ^ 10 = skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_132, plain,
    (skoXC1 ^ 10 != skoXC1 ^ 10 | skoXC1 != 354991/512 |
     skoXC1 ^ 10 = 354991/512 ^ 10), introduced(tautology, [equality])).

cnf(refute_0_133, plain,
    (skoXC1 != 354991/512 | skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain, (skoXC1 ^ 10 = 354991/512 ^ 10),
    inference(resolve, [], [refute_0_2, refute_0_133])).

cnf(refute_0_135, plain, (9 * skoXC1 ^ 10 = 9 * skoXC1 ^ 10),
    introduced(tautology, [refl])).

cnf(refute_0_136, plain,
    (9 * skoXC1 ^ 10 != 9 * skoXC1 ^ 10 | skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    introduced(tautology, [equality])).

cnf(refute_0_137, plain,
    (skoXC1 ^ 10 != 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain, (9 * skoXC1 ^ 10 = 9 * 354991/512 ^ 10),
    inference(resolve, [], [refute_0_134, refute_0_137])).

cnf(refute_0_139, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [refl])).

cnf(refute_0_140, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    introduced(tautology, [equality])).

cnf(refute_0_141, plain,
    (9 * skoXC1 ^ 10 != 9 * 354991/512 ^ 10 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_139, refute_0_140])).

cnf(refute_0_142, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8),
    inference(resolve, [], [refute_0_138, refute_0_141])).

cnf(refute_0_143, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_144, plain,
    (9 * 354991/512 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8),
    inference(resolve, [], [refute_0_130, refute_0_144])).

cnf(refute_0_146, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [refl])).

cnf(refute_0_147, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6), introduced(tautology, [equality])).

cnf(refute_0_148, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_146, refute_0_147])).

cnf(refute_0_149, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6),
    inference(resolve, [], [refute_0_145, refute_0_148])).

cnf(refute_0_150, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_151, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6),
    inference(resolve, [], [refute_0_118, refute_0_151])).

cnf(refute_0_153, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [refl])).

cnf(refute_0_154, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    introduced(tautology, [equality])).

cnf(refute_0_155, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4),
    inference(resolve, [], [refute_0_152, refute_0_155])).

cnf(refute_0_157, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_158, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4),
    inference(resolve, [], [refute_0_106, refute_0_158])).

cnf(refute_0_160, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_161, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_162, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2),
    inference(resolve, [], [refute_0_159, refute_0_162])).

cnf(refute_0_164, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_165, plain,
    (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_163, refute_0_164])).

cnf(refute_0_166, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 =
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2),
    inference(resolve, [], [refute_0_98, refute_0_165])).

cnf(refute_0_167, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_168, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_169, plain,
    (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
     10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
     4000000000000000000000000 * skoXC1 ^ 2 !=
     9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
     10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
     4000000000000000000000000 * 354991/512 ^ 2 |
     33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) =
     33600000000000 ^ 2 /
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2)),
    inference(resolve, [], [refute_0_166, refute_0_169])).

cnf(refute_0_171, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_172, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_173, plain,
    (33600000000000 ^ 2 /
     (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
      10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
      4000000000000000000000000 * skoXC1 ^ 2) !=
     33600000000000 ^ 2 /
     (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
      10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
      4000000000000000000000000 * 354991/512 ^ 2) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_171, refute_0_172])).

cnf(refute_0_174, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_170, refute_0_173])).

cnf(refute_0_175, plain,
    (868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_176, plain,
    (868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) |
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    introduced(tautology, [equality])).

cnf(refute_0_177, plain,
    (ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_174, refute_0_177])).

cnf(refute_0_179, plain,
    (log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2))),
    inference(subst, [], [log_def])).

cnf(refute_0_180, plain,
    (868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) |
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_181, plain,
    (868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
         10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
         4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2))),
    inference(resolve, [], [refute_0_178, refute_0_181])).

cnf(refute_0_183, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2))),
    introduced(tautology, [refl])).

cnf(refute_0_184, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) !=
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) |
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    introduced(tautology, [equality])).

cnf(refute_0_185, plain,
    (log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) !=
     868589/2000000 *
     ln(33600000000000 ^ 2 /
        (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
         10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
         4000000000000000000000000 * 354991/512 ^ 2)) |
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    inference(resolve, [], [refute_0_183, refute_0_184])).

cnf(refute_0_186, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) =
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    inference(resolve, [], [refute_0_182, refute_0_185])).

cnf(refute_0_187, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) !=
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2))) |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    introduced(tautology, [equality])).

cnf(refute_0_188, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    inference(resolve, [], [refute_0_186, refute_0_187])).

cnf(refute_0_189, plain,
    (10 *
     log(33600000000000 ^ 2 /
         (9 * skoXC1 ^ 10 + 309000000 * skoXC1 ^ 8 +
          10300000000000000 * skoXC1 ^ 6 + 10 ^ 22 * skoXC1 ^ 4 +
          4000000000000000000000000 * skoXC1 ^ 2)) <
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * skoXC1 ^ 2 + 3 * skoXC1 ^ 4 + 2000000000000) /
             (skoXC1 * (29 * skoXC1 ^ 2 - 90000000)))) - 18 |
     -12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    inference(resolve, [], [refute_0_90, refute_0_188])).

cnf(refute_0_190, plain,
    (-12 / pi *
     (-pi -
      arctan(1/1140 *
             (-130660000 * 354991/512 ^ 2 + 3 * 354991/512 ^ 4 +
              2000000000000) /
             (354991/512 * (29 * 354991/512 ^ 2 - 90000000)))) - 18 <=
     10 *
     (868589/2000000 *
      ln(33600000000000 ^ 2 /
         (9 * 354991/512 ^ 10 + 309000000 * 354991/512 ^ 8 +
          10300000000000000 * 354991/512 ^ 6 + 10 ^ 22 * 354991/512 ^ 4 +
          4000000000000000000000000 * 354991/512 ^ 2)))),
    inference(resolve, [], [refute_0_1, refute_0_189])).

cnf(refute_0_191, plain,
    ((pi * 12 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) /
     pi <=
     18 +
     ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) *
     868589/200000), inference(arithmetic, [], [refute_0_190])).

cnf(refute_0_192, plain,
    (18 +
     ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) *
     868589/200000 <
     (pi * 12 +
      arctan(4131275018871243593760317/4131263554720769638778880) * 12) /
     pi |
     pi * 12 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 12 <=
     (18 +
      ln(9507379501711720511225274040320000000000000000000000/
         29254561989755623963996280317717140077560083836061785704147) *
      868589/200000) * pi | pi <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_193, plain,
    (pi * 12 +
     arctan(4131275018871243593760317/4131263554720769638778880) * 12 <=
     (18 +
      ln(9507379501711720511225274040320000000000000000000000/
         29254561989755623963996280317717140077560083836061785704147) *
      868589/200000) * pi | pi <= 0),
    inference(resolve, [], [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) *
     (pi * -868589/200000) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12 |
     pi <= 0), inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_196, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_197, plain, (0 < pi),
    inference(decision, [], [refute_0_2, refute_0_195, refute_0_196])).

cnf(refute_0_198, plain,
    (ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) *
     (pi * -868589/200000) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12),
    inference(resolve, [], [refute_0_194, refute_0_197])).

cnf(refute_0_199, plain,
    (0 <= pi * -868589/200000 |
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147)),
    inference(split, [], [refute_0_198])).

cnf(refute_0_200, plain,
    ((pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) |
     pi <= 0), inference(arithmetic, [], [refute_0_199])).

cnf(refute_0_201, plain,
    ((pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147)),
    inference(resolve, [], [refute_0_200, refute_0_197])).

cnf(refute_0_202, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_203, plain,
    (~ lgen(1, X_000048 - 1, X_000049) | X_000048 <= 0 |
     lgen(1, ln(X_000048), X_000049)),
    inference(subst, [], [ln_upper_bound_cf1])).

cnf(refute_0_204, plain,
    (ln(X_000048) < X_000049 | ~ lgen(1, X_000048 - 1, X_000049) |
     X_000048 <= 0), inference(resolve, [], [refute_0_203, refute_0_202])).

cnf(refute_0_205, plain,
    (ln(X_000048) < X_000049 | X_000048 <= 0 | X_000049 <= -1 + X_000048),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (ln(9507379501711720511225274040320000000000000000000000/
        29254561989755623963996280317717140077560083836061785704147) <
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) |
     9507379501711720511225274040320000000000000000000000/
     29254561989755623963996280317717140077560083836061785704147 <= 0 |
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     -1 +
     9507379501711720511225274040320000000000000000000000/
     29254561989755623963996280317717140077560083836061785704147),
    inference(subst, [], [refute_0_205])).

cnf(refute_0_207, plain,
    (9507379501711720511225274040320000000000000000000000/
     29254561989755623963996280317717140077560083836061785704147 <= 0 |
     (pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     -1 +
     9507379501711720511225274040320000000000000000000000/
     29254561989755623963996280317717140077560083836061785704147),
    inference(resolve, [], [refute_0_201, refute_0_206])).

cnf(refute_0_208, plain,
    ((pi * 6 +
      arctan(4131275018871243593760317/4131263554720769638778880) * -12) /
     (pi * -868589/200000) <=
     -29254552482376122252275769092443099757560083836061785704147/
     29254561989755623963996280317717140077560083836061785704147),
    inference(arithmetic, [], [refute_0_207])).

cnf(refute_0_209, plain,
    (0 <= pi * -868589/200000 |
     -29254552482376122252275769092443099757560083836061785704147/
     29254561989755623963996280317717140077560083836061785704147 *
     (pi * -868589/200000) <=
     pi * 6 +
     arctan(4131275018871243593760317/4131263554720769638778880) * -12),
    inference(resolve, [], [refute_0_208, refute_0_0])).

cnf(refute_0_210, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     9695291901592155105813578381024508517752744944193072461997061417/
     70210948775413497513591072762521136186144201206548285689952800000 |
     pi <= 0), inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain,
    (arctan(4131275018871243593760317/4131263554720769638778880) <=
     pi *
     9695291901592155105813578381024508517752744944193072461997061417/
     70210948775413497513591072762521136186144201206548285689952800000),
    inference(resolve, [], [refute_0_210, refute_0_197])).

cnf(refute_0_212, plain,
    (X_000056 < arctan(X_000055) | ~ lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_213, plain,
    (~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1 |
     lgen(1, X_000056, arctan(X_000055))),
    inference(subst, [], [atan_lower_bound_case_14])).

cnf(refute_0_214, plain,
    (X_000056 < arctan(X_000055) |
     ~ lgen(1, X_000056, 157079/100000 - 1 / X_000055) | X_000055 <= 1),
    inference(resolve, [], [refute_0_213, refute_0_212])).

cnf(refute_0_215, plain,
    (X_000056 < arctan(X_000055) | X_000055 <= 1 |
     157079/100000 + 1 / X_000055 * -1 <= X_000056),
    inference(arithmetic, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (pi *
     9695291901592155105813578381024508517752744944193072461997061417/
     70210948775413497513591072762521136186144201206548285689952800000 <
     arctan(4131275018871243593760317/4131263554720769638778880) |
     4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     9695291901592155105813578381024508517752744944193072461997061417/
     70210948775413497513591072762521136186144201206548285689952800000),
    inference(subst, [], [refute_0_215])).

cnf(refute_0_217, plain,
    (4131275018871243593760317/4131263554720769638778880 <= 1 |
     157079/100000 +
     1 / 4131275018871243593760317/4131263554720769638778880 * -1 <=
     pi *
     9695291901592155105813578381024508517752744944193072461997061417/
     70210948775413497513591072762521136186144201206548285689952800000),
    inference(resolve, [], [refute_0_211, refute_0_216])).

cnf(refute_0_218, plain,
    (165564573966931259905093513717868706436159833262323706511384332970507557844142841917031704/
     40053917233712345772327229818456597095577902984862061464395784936697500525476971026389189 <=
     pi), inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (pi <
     165564573966931259905093513717868706436159833262323706511384332970507557844142841917031704/
     40053917233712345772327229818456597095577902984862061464395784936697500525476971026389189),
    inference(decision, [], [refute_0_2, refute_0_195, refute_0_196])).

cnf(refute_0_220, plain, ($false),
    inference(resolve, [], [refute_0_218, refute_0_219])).
% SZS output end CNFRefutation for Nichols-Plot-Disk-Drive-Meets-1-6-weak.tptp
