% SZS output start CNFRefutation for Nichols-Plot-Simple-Example-1.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(atan_lower_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, Y, -3927/2500 - 3 * X / (1 + 3 * X ^ 2)) |
     lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, Y, X) | lgen(R, Y, arctan(X)))).

cnf(atan_lower_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, arctan(X)))).

cnf(atan_upper_bound_case_11, axiom,
    (-1 <= X | ~ lgen(R, -157079/100000 - 1 / X, Y) |
     lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_12, axiom,
    (0 <= X | ~ lgen(R, 3 * X / (X ^ 2 + 3), Y) | lgen(R, arctan(X), Y))).

cnf(atan_upper_bound_case_13, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, arctan(X), Y))).

fof(nichols_plot_1, conjecture,
    (! [X] :
       ((423/1000 <= X & X <= 383/500) =>
        interval(1, -3/2, 1, 1/2,
                 arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 -
                         3194) /
                        (800 * X ^ 8 + 48035 * X ^ 6 + 12054 * X ^ 4 -
                         1597 * X ^ 2 + 55)))))).

fof(subgoal_0, plain,
    (! [X] :
       ((423/1000 <= X & X <= 383/500) =>
        interval(1, -3/2, 1, 1/2,
                 arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 -
                         3194) /
                        (800 * X ^ 8 + 48035 * X ^ 6 + 12054 * X ^ 4 -
                         1597 * X ^ 2 + 55))))),
    inference(strip, [], [nichols_plot_1])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((423/1000 <= X & X <= 383/500) =>
          interval(1, -3/2, 1, 1/2,
                   arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 -
                           3194) /
                          (800 * X ^ 8 + 48035 * X ^ 6 + 12054 * X ^ 4 -
                           1597 * X ^ 2 + 55))))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (~ interval(1, -3/2, 1, 1/2,
                   arctan((797 * X ^ 6 + 14382 * X ^ 4 + 755 * X ^ 2 -
                           3194) /
                          (800 * X ^ 8 + 48035 * X ^ 6 + 12054 * X ^ 4 -
                           1597 * X ^ 2 + 55))) & X <= 383/500 &
        423/1000 <= X)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (~ interval(1, -3/2, 1, 1/2,
                arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 +
                        755 * skoXC1 ^ 2 - 3194) /
                       (800 * skoXC1 ^ 8 + 48035 * skoXC1 ^ 6 +
                        12054 * skoXC1 ^ 4 - 1597 * skoXC1 ^ 2 + 55))) &
     423/1000 <= skoXC1 & skoXC1 <= 383/500),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (~ interval(1, -3/2, 1, 1/2,
                arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 +
                        755 * skoXC1 ^ 2 - 3194) /
                       (800 * skoXC1 ^ 8 + 48035 * skoXC1 ^ 6 +
                        12054 * skoXC1 ^ 4 - 1597 * skoXC1 ^ 2 + 55)))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 383/500),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (423/1000 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (-3/2 <
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) <=
     -3/2 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_2, plain,
    (X_000062 < arctan(X_000061) | ~ lgen(1, X_000062, arctan(X_000061))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (X_000061 < 0 |
     ~ lgen(1, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3)) |
     lgen(1, X_000062, arctan(X_000061))),
    inference(subst, [], [atan_lower_bound_case_13])).

cnf(refute_0_4, plain,
    (X_000061 < 0 | X_000062 < arctan(X_000061) |
     ~ lgen(1, X_000062, 3 * X_000061 / (X_000061 ^ 2 + 3))),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000061 < 0 | X_000062 < arctan(X_000061) |
     X_000061 * 3 / (3 + X_000061 * X_000061) <= X_000062),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     -3/2 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (3 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))))) <=
     -3/2), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (~ interval(1, -3/2, 1, 1/2,
                arctan((797 * skoXC1 ^ 6 + 14382 * skoXC1 ^ 4 +
                        755 * skoXC1 ^ 2 - 3194) /
                       (800 * skoXC1 ^ 8 + 48035 * skoXC1 ^ 6 +
                        12054 * skoXC1 ^ 4 - 1597 * skoXC1 ^ 2 + 55)))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (~ interval(1, -3/2, 1, 1/2,
                arctan((-3194 +
                        skoXC1 *
                        (skoXC1 *
                         (755 +
                          skoXC1 *
                          (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                       (55 +
                        skoXC1 *
                        (skoXC1 *
                         (-1597 +
                          skoXC1 *
                          (skoXC1 *
                           (12054 +
                            skoXC1 *
                            (skoXC1 *
                             (48035 + skoXC1 * (skoXC1 * 800))))))))))),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(1, -3/2,
            arctan((-3194 +
                    skoXC1 *
                    (skoXC1 *
                     (755 +
                      skoXC1 *
                      (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                   (55 +
                    skoXC1 *
                    (skoXC1 *
                     (-1597 +
                      skoXC1 *
                      (skoXC1 *
                       (12054 +
                        skoXC1 *
                        (skoXC1 *
                         (48035 + skoXC1 * (skoXC1 * 800)))))))))) |
     ~ lgen(1,
            arctan((-3194 +
                    skoXC1 *
                    (skoXC1 *
                     (755 +
                      skoXC1 *
                      (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                   (55 +
                    skoXC1 *
                    (skoXC1 *
                     (-1597 +
                      skoXC1 *
                      (skoXC1 *
                       (12054 +
                        skoXC1 *
                        (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))),
            1/2) |
     interval(1, -3/2, 1, 1/2,
              arctan((-3194 +
                      skoXC1 *
                      (skoXC1 *
                       (755 +
                        skoXC1 *
                        (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                     (55 +
                      skoXC1 *
                      (skoXC1 *
                       (-1597 +
                        skoXC1 *
                        (skoXC1 *
                         (12054 +
                          skoXC1 *
                          (skoXC1 *
                           (48035 + skoXC1 * (skoXC1 * 800))))))))))),
    inference(subst, [], [interval_intro])).

cnf(refute_0_10, plain,
    (~ lgen(1, -3/2,
            arctan((-3194 +
                    skoXC1 *
                    (skoXC1 *
                     (755 +
                      skoXC1 *
                      (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                   (55 +
                    skoXC1 *
                    (skoXC1 *
                     (-1597 +
                      skoXC1 *
                      (skoXC1 *
                       (12054 +
                        skoXC1 *
                        (skoXC1 *
                         (48035 + skoXC1 * (skoXC1 * 800)))))))))) |
     ~ lgen(1,
            arctan((-3194 +
                    skoXC1 *
                    (skoXC1 *
                     (755 +
                      skoXC1 *
                      (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
                   (55 +
                    skoXC1 *
                    (skoXC1 *
                     (-1597 +
                      skoXC1 *
                      (skoXC1 *
                       (12054 +
                        skoXC1 *
                        (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))),
            1/2)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (1/2 <=
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) <=
     -3/2), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_13, plain,
    ((-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) < 1/2 |
     1/2 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_14, plain,
    (arctan(X_000068) < X_000069 | ~ lgen(1, arctan(X_000068), X_000069)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_15, plain,
    (~ lgen(1, 3 * X_000068 / (X_000068 ^ 2 + 3), X_000069) |
     0 <= X_000068 | lgen(1, arctan(X_000068), X_000069)),
    inference(subst, [], [atan_upper_bound_case_12])).

cnf(refute_0_16, plain,
    (arctan(X_000068) < X_000069 |
     ~ lgen(1, 3 * X_000068 / (X_000068 ^ 2 + 3), X_000069) |
     0 <= X_000068), inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    (arctan(X_000068) < X_000069 |
     X_000069 <= X_000068 * 3 / (3 + X_000068 * X_000068) | 0 <= X_000068),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (3 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))))),
    inference(subst, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-33216626 +
           skoXC1 *
           (skoXC1 *
            (-8036874 +
             skoXC1 *
             (skoXC1 *
              (1155472580 +
               skoXC1 *
               (skoXC1 *
                (2326647625 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) =
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 <= 383/500),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_21, plain, (423/1000 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-33216626 +
           skoXC1 *
           (skoXC1 *
            (-8036874 +
             skoXC1 *
             (skoXC1 *
              (1155472580 +
               skoXC1 *
               (skoXC1 *
                (2326647625 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) !=
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(resolve, [], [refute_0_19, refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) != -55),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_25, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000))))))))))))))))),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(resolve, [], [refute_0_25, refute_0_13])).

cnf(refute_0_27, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-18102004 +
       skoXC1 *
       (skoXC1 *
        (76909031 +
         skoXC1 *
         (skoXC1 *
          (460229570 +
           skoXC1 *
           (skoXC1 *
            (-525429746 +
             skoXC1 *
             (skoXC1 *
              (-358479900 +
               skoXC1 *
               (skoXC1 *
                (3340920557 +
                 skoXC1 *
                 (skoXC1 *
                  (113371200 + skoXC1 * (skoXC1 * 960000))))))))))))))) <=
     -11264731/2 |
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     -10210711), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < -1 |
     -1 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_29, plain,
    (-207079/100000 <
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) /
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) |
     0 <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     -207079/100000 *
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) <=
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_30, plain,
    (arctan(X_000054) < X_000055 | ~ lgen(1, arctan(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_31, plain,
    (~ lgen(1, -157079/100000 - 1 / X_000054, X_000055) | -1 <= X_000054 |
     lgen(1, arctan(X_000054), X_000055)),
    inference(subst, [], [atan_upper_bound_case_11])).

cnf(refute_0_32, plain,
    (arctan(X_000054) < X_000055 |
     ~ lgen(1, -157079/100000 - 1 / X_000054, X_000055) | -1 <= X_000054),
    inference(resolve, [], [refute_0_31, refute_0_30])).

cnf(refute_0_33, plain,
    (arctan(X_000054) < X_000055 |
     X_000055 <= -157079/100000 + 1 / X_000054 * -1 | -1 <= X_000054),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     1/2 <=
     -157079/100000 +
     1 /
     ((-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) *
     -1), inference(subst, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) /
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) <=
     -207079/100000), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     0 <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     -207079/100000 *
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) <=
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))),
    inference(resolve, [], [refute_0_35, refute_0_29])).

cnf(refute_0_37, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194 |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000), inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <
     1/2 |
     1/2 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_39, plain,
    (arctan(X_000040) < X_000041 | ~ lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_40, plain,
    (X_000040 < 0 | ~ lgen(1, X_000040, X_000041) |
     lgen(1, arctan(X_000040), X_000041)),
    inference(subst, [], [atan_upper_bound_case_13])).

cnf(refute_0_41, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 |
     ~ lgen(1, X_000040, X_000041)),
    inference(resolve, [], [refute_0_40, refute_0_39])).

cnf(refute_0_42, plain,
    (arctan(X_000040) < X_000041 | X_000040 < 0 | X_000041 <= X_000040),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     1/2 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_44, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     1/2 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_43, refute_0_38])).

cnf(refute_0_45, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (-3107/2 +
       skoXC1 *
       (skoXC1 *
        (-8355 +
         skoXC1 * (skoXC1 * (46441/2 + skoXC1 * (skoXC1 * 400))))))) <=
     -6443/2), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-3107/2 +
       skoXC1 *
       (skoXC1 *
        (-8355 +
         skoXC1 * (skoXC1 * (46441/2 + skoXC1 * (skoXC1 * 400))))))) <=
     -6443/2), inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_47, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (-3107/2 +
       skoXC1 *
       (skoXC1 *
        (-8355 +
         skoXC1 * (skoXC1 * (46441/2 + skoXC1 * (skoXC1 * 400))))))) <=
     -6443/2), inference(resolve, [], [refute_0_45, refute_0_46])).

cnf(refute_0_48, plain,
    (-6443/2 <
     skoXC1 *
     (skoXC1 *
      (-3107/2 +
       skoXC1 *
       (skoXC1 *
        (-8355 +
         skoXC1 * (skoXC1 * (46441/2 + skoXC1 * (skoXC1 * 400)))))))),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_49, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_49])).

cnf(refute_0_51, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     -3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797)))))),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_53, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797)))))),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000),
    inference(decision, [], [refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_55, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000),
    inference(resolve, [], [refute_0_37, refute_0_54])).

cnf(refute_0_56, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -1 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000 |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_55, refute_0_28])).

cnf(refute_0_57, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139),
    inference(decision, [], [refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_59, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) <=
     -327955163/50000 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -327955163/50000 <
     skoXC1 *
     (skoXC1 *
      (671071/20000 +
       skoXC1 *
       (skoXC1 *
        (-2091805089/50000 +
         skoXC1 *
         (skoXC1 * (-4968541963/100000 + skoXC1 * (skoXC1 * -800))))))) |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139),
    inference(decision, [], [refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_61, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(resolve, [], [refute_0_59, refute_0_60])).

cnf(refute_0_62, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -11264731/2 <
     skoXC1 *
     (skoXC1 *
      (-18102004 +
       skoXC1 *
       (skoXC1 *
        (76909031 +
         skoXC1 *
         (skoXC1 *
          (460229570 +
           skoXC1 *
           (skoXC1 *
            (-525429746 +
             skoXC1 *
             (skoXC1 *
              (-358479900 +
               skoXC1 *
               (skoXC1 *
                (3340920557 +
                 skoXC1 *
                 (skoXC1 *
                  (113371200 + skoXC1 * (skoXC1 * 960000))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     -10210711),
    inference(decision, [],
              [refute_0_61, refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_63, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     -10210711), inference(resolve, [], [refute_0_27, refute_0_62])).

cnf(refute_0_64, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -10210711 <
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000)))))))))))))))),
    inference(decision, [],
              [refute_0_61, refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_65, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_65, refute_0_12])).

cnf(refute_0_67, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     -55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194),
    inference(decision, [],
              [refute_0_61, refute_0_53, refute_0_20, refute_0_21])).

cnf(refute_0_69, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) < 1/2),
    inference(resolve, [], [refute_0_69, refute_0_53])).

cnf(refute_0_71, plain,
    (arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) <=
     -3/2), inference(resolve, [], [refute_0_11, refute_0_70])).

cnf(refute_0_72, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (3 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))))) <=
     -3/2), inference(resolve, [], [refute_0_71, refute_0_6])).

cnf(refute_0_73, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -3/2 |
     skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-33216626 +
           skoXC1 *
           (skoXC1 *
            (-8036874 +
             skoXC1 *
             (skoXC1 *
              (1155472580 +
               skoXC1 *
               (skoXC1 *
                (2326647625 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) =
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -3/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(resolve, [], [refute_0_73, refute_0_22])).

cnf(refute_0_75, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) <= -3/2),
    inference(resolve, [], [refute_0_74, refute_0_24])).

cnf(refute_0_76, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) <=
     -3/2 *
     (10210711 +
      skoXC1 *
      (skoXC1 *
       (-5349950 +
        skoXC1 *
        (skoXC1 *
         (-79673144 +
          skoXC1 *
          (skoXC1 *
           (-83024294 +
            skoXC1 *
            (skoXC1 *
             (183934772 +
              skoXC1 *
              (skoXC1 *
               (3489342648 +
                skoXC1 *
                (skoXC1 *
                 (6980578084 +
                  skoXC1 *
                  (skoXC1 *
                   (230568000 +
                    skoXC1 * (skoXC1 * 1920000)))))))))))))))) |
     10210711 +
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     0), inference(resolve, [], [refute_0_75, refute_0_1])).

cnf(refute_0_77, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) <=
     -10210711 |
     skoXC1 *
     (skoXC1 *
      (7402104 +
       skoXC1 *
       (skoXC1 *
        (-236255319 +
         skoXC1 *
         (skoXC1 *
          (-626278158 +
           skoXC1 *
           (skoXC1 *
            (893299290 +
             skoXC1 *
             (skoXC1 *
              (7337165196 +
               skoXC1 *
               (skoXC1 *
                (10620235611 +
                 skoXC1 *
                 (skoXC1 *
                  (347764800 + skoXC1 * (skoXC1 * 2880000))))))))))))))) <=
     -29578113/2), inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    ((-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <
     -177/2500 |
     -177/2500 *
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     30607933 +
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (16660126 +
           skoXC1 *
           (skoXC1 *
            (616099308 +
             skoXC1 *
             (skoXC1 *
              (1224247304 +
               skoXC1 *
               (skoXC1 *
                (2328553252 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_79, plain,
    (X_000076 < arctan(X_000075) | ~ lgen(1, X_000076, arctan(X_000075))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_80, plain,
    (~ lgen(1, X_000076,
            -3927/2500 - 3 * X_000075 / (1 + 3 * X_000075 ^ 2)) |
     -1 <= X_000075 | lgen(1, X_000076, arctan(X_000075))),
    inference(subst, [], [atan_lower_bound_case_11])).

cnf(refute_0_81, plain,
    (X_000076 < arctan(X_000075) |
     ~ lgen(1, X_000076,
            -3927/2500 - 3 * X_000075 / (1 + 3 * X_000075 ^ 2)) |
     -1 <= X_000075), inference(resolve, [], [refute_0_80, refute_0_79])).

cnf(refute_0_82, plain,
    (X_000076 < arctan(X_000075) | -1 <= X_000075 |
     -3927/2500 + X_000075 * 3 / (1 + X_000075 * (X_000075 * 3)) * -1 <=
     X_000076), inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (-3/2 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) |
     -1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -3927/2500 +
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (1 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
       3)) * -1 <= -3/2), inference(subst, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -3927/2500 +
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) * 3 /
     (1 +
      (-3194 +
       skoXC1 *
       (skoXC1 *
        (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
      (55 +
       skoXC1 *
       (skoXC1 *
        (-1597 +
         skoXC1 *
         (skoXC1 *
          (12054 +
           skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
      ((-3194 +
        skoXC1 *
        (skoXC1 *
         (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
       (55 +
        skoXC1 *
        (skoXC1 *
         (-1597 +
          skoXC1 *
          (skoXC1 *
           (12054 +
            skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) *
       3)) * -1 <= -3/2),
    inference(resolve, [], [refute_0_71, refute_0_83])).

cnf(refute_0_85, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -177/2500 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-175670 +
       skoXC1 *
       (skoXC1 *
        (3876349 +
         skoXC1 *
         (skoXC1 *
          (-33216626 +
           skoXC1 *
           (skoXC1 *
            (-8036874 +
             skoXC1 *
             (skoXC1 *
              (1155472580 +
               skoXC1 *
               (skoXC1 *
                (2326647625 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) =
     -3025 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -177/2500 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) = -55),
    inference(resolve, [], [refute_0_85, refute_0_22])).

cnf(refute_0_87, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -177/2500 <=
     (-527010 +
      skoXC1 *
      (skoXC1 *
       (15427029 +
        skoXC1 *
        (skoXC1 *
         (-116745603 +
          skoXC1 *
          (skoXC1 *
           (-501741717 +
            skoXC1 *
            (skoXC1 *
             (617397132 +
              skoXC1 *
              (skoXC1 *
               (2103151224 +
                skoXC1 *
                (skoXC1 *
                 (149368485 + skoXC1 * (skoXC1 * 1912800)))))))))))))) /
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))))),
    inference(resolve, [], [refute_0_86, refute_0_24])).

cnf(refute_0_88, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -177/2500 *
     (30607933 +
      skoXC1 *
      (skoXC1 *
       (-14644490 +
        skoXC1 *
        (skoXC1 *
         (-270030224 +
          skoXC1 *
          (skoXC1 *
           (16660126 +
            skoXC1 *
            (skoXC1 *
             (616099308 +
              skoXC1 *
              (skoXC1 *
               (1224247304 +
                skoXC1 *
                (skoXC1 *
                 (2328553252 +
                  skoXC1 *
                  (skoXC1 *
                   (76856000 + skoXC1 * (skoXC1 * 640000)))))))))))))))) <=
     -527010 +
     skoXC1 *
     (skoXC1 *
      (15427029 +
       skoXC1 *
       (skoXC1 *
        (-116745603 +
         skoXC1 *
         (skoXC1 *
          (-501741717 +
           skoXC1 *
           (skoXC1 *
            (617397132 +
             skoXC1 *
             (skoXC1 *
              (2103151224 +
               skoXC1 *
               (skoXC1 *
                (149368485 + skoXC1 * (skoXC1 * 1912800))))))))))))) |
     30607933 +
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (16660126 +
           skoXC1 *
           (skoXC1 *
            (616099308 +
             skoXC1 *
             (skoXC1 *
              (1224247304 +
               skoXC1 *
               (skoXC1 *
                (2328553252 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     0), inference(resolve, [], [refute_0_87, refute_0_78])).

cnf(refute_0_89, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (16660126 +
           skoXC1 *
           (skoXC1 *
            (616099308 +
             skoXC1 *
             (skoXC1 *
              (1224247304 +
               skoXC1 *
               (skoXC1 *
                (2328553252 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) <=
     -30607933 |
     skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 +
                   skoXC1 * (skoXC1 * -45312))))))))))))))) <=
     4100079141/2500), inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (-3/2 <
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <=
     -3/2 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_91, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_92, plain,
    (~ lgen(1, X_000034, X_000033) | 0 <= X_000033 |
     lgen(1, X_000034, arctan(X_000033))),
    inference(subst, [], [atan_lower_bound_case_12])).

cnf(refute_0_93, plain,
    (X_000034 < arctan(X_000033) | ~ lgen(1, X_000034, X_000033) |
     0 <= X_000033), inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain,
    (X_000034 < arctan(X_000033) | X_000033 <= X_000034 | 0 <= X_000033),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (-3/2 <
     arctan((-3194 +
             skoXC1 *
             (skoXC1 *
              (755 +
               skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
            (55 +
             skoXC1 *
             (skoXC1 *
              (-1597 +
               skoXC1 *
               (skoXC1 *
                (12054 +
                 skoXC1 *
                 (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))))) |
     0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3/2), inference(subst, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3/2), inference(resolve, [], [refute_0_71, refute_0_95])).

cnf(refute_0_97, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <=
     -3/2 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_90])).

cnf(refute_0_98, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2), inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_100, plain,
    (0 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2), inference(resolve, [], [refute_0_98, refute_0_99])).

cnf(refute_0_101, plain,
    (0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2 |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_100, refute_0_12])).

cnf(refute_0_102, plain,
    (skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(decision, [], [refute_0_21, refute_0_20])).

cnf(refute_0_104, plain,
    (skoXC1 *
     (skoXC1 *
      (-3281/2 +
       skoXC1 *
       (skoXC1 *
        (32463 +
         skoXC1 * (skoXC1 * (145699/2 + skoXC1 * (skoXC1 * 1200))))))) <=
     6223/2 |
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194), inference(resolve, [], [refute_0_102, refute_0_103])).

cnf(refute_0_105, plain,
    (-30607933 <
     skoXC1 *
     (skoXC1 *
      (-14644490 +
       skoXC1 *
       (skoXC1 *
        (-270030224 +
         skoXC1 *
         (skoXC1 *
          (16660126 +
           skoXC1 *
           (skoXC1 *
            (616099308 +
             skoXC1 *
             (skoXC1 *
              (1224247304 +
               skoXC1 *
               (skoXC1 *
                (2328553252 +
                 skoXC1 *
                 (skoXC1 *
                  (76856000 + skoXC1 * (skoXC1 * 640000))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 +
                   skoXC1 * (skoXC1 * -45312))))))))))))))) <=
     4100079141/2500),
    inference(decision, [], [refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_106, plain,
    (-1 <=
     (-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 +
                   skoXC1 * (skoXC1 * -45312))))))))))))))) <=
     4100079141/2500),
    inference(resolve, [], [refute_0_89, refute_0_105])).

cnf(refute_0_107, plain,
    (-1 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) <=
     -3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) |
     skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 +
                   skoXC1 * (skoXC1 * -45312))))))))))))))) <=
     4100079141/2500 |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_106, refute_0_28])).

cnf(refute_0_108, plain,
    (skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 +
                   skoXC1 * (skoXC1 * -45312))))))))))))))) <=
     4100079141/2500 |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (4100079141/2500 <
     skoXC1 *
     (skoXC1 *
      (-3597549777/250 +
       skoXC1 *
       (skoXC1 *
        (84914839287/625 +
         skoXC1 *
         (skoXC1 *
          (625702725099/1250 +
           skoXC1 *
           (skoXC1 *
            (-413135601879/625 +
             skoXC1 *
             (skoXC1 *
              (-1368642458202/625 +
               skoXC1 *
               (skoXC1 *
                (-196393784526/625 +
                 skoXC1 *
                 (skoXC1 *
                  (-36771024/5 + skoXC1 * (skoXC1 * -45312))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139),
    inference(decision, [], [refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_110, plain,
    (skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55 |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(resolve, [], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain,
    (-55 <
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) |
     skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139),
    inference(decision, [], [refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_112, plain,
    (skoXC1 *
     (skoXC1 *
      (842 +
       skoXC1 *
       (skoXC1 *
        (-26436 +
         skoXC1 * (skoXC1 * (-48832 + skoXC1 * (skoXC1 * -800))))))) <=
     -3139), inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (-10210711 <
     skoXC1 *
     (skoXC1 *
      (-5349950 +
       skoXC1 *
       (skoXC1 *
        (-79673144 +
         skoXC1 *
         (skoXC1 *
          (-83024294 +
           skoXC1 *
           (skoXC1 *
            (183934772 +
             skoXC1 *
             (skoXC1 *
              (3489342648 +
               skoXC1 *
               (skoXC1 *
                (6980578084 +
                 skoXC1 *
                 (skoXC1 *
                  (230568000 + skoXC1 * (skoXC1 * 1920000))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (7402104 +
       skoXC1 *
       (skoXC1 *
        (-236255319 +
         skoXC1 *
         (skoXC1 *
          (-626278158 +
           skoXC1 *
           (skoXC1 *
            (893299290 +
             skoXC1 *
             (skoXC1 *
              (7337165196 +
               skoXC1 *
               (skoXC1 *
                (10620235611 +
                 skoXC1 *
                 (skoXC1 *
                  (347764800 + skoXC1 * (skoXC1 * 2880000))))))))))))))) <=
     -29578113/2),
    inference(decision, [],
              [refute_0_112, refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_114, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0 |
     skoXC1 *
     (skoXC1 *
      (7402104 +
       skoXC1 *
       (skoXC1 *
        (-236255319 +
         skoXC1 *
         (skoXC1 *
          (-626278158 +
           skoXC1 *
           (skoXC1 *
            (893299290 +
             skoXC1 *
             (skoXC1 *
              (7337165196 +
               skoXC1 *
               (skoXC1 *
                (10620235611 +
                 skoXC1 *
                 (skoXC1 *
                  (347764800 + skoXC1 * (skoXC1 * 2880000))))))))))))))) <=
     -29578113/2), inference(resolve, [], [refute_0_77, refute_0_113])).

cnf(refute_0_115, plain,
    (-29578113/2 <
     skoXC1 *
     (skoXC1 *
      (7402104 +
       skoXC1 *
       (skoXC1 *
        (-236255319 +
         skoXC1 *
         (skoXC1 *
          (-626278158 +
           skoXC1 *
           (skoXC1 *
            (893299290 +
             skoXC1 *
             (skoXC1 *
              (7337165196 +
               skoXC1 *
               (skoXC1 *
                (10620235611 +
                 skoXC1 *
                 (skoXC1 *
                  (347764800 + skoXC1 * (skoXC1 * 2880000)))))))))))))))),
    inference(decision, [],
              [refute_0_112, refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_116, plain,
    ((-3194 +
      skoXC1 *
      (skoXC1 *
       (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797)))))) /
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) < 0),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain,
    (-3194 +
     skoXC1 *
     (skoXC1 *
      (755 + skoXC1 * (skoXC1 * (14382 + skoXC1 * (skoXC1 * 797))))) <
     0 *
     (55 +
      skoXC1 *
      (skoXC1 *
       (-1597 +
        skoXC1 *
        (skoXC1 *
         (12054 +
          skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800)))))))) |
     55 +
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_116])).

cnf(refute_0_118, plain,
    (-3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) |
     skoXC1 *
     (skoXC1 *
      (-1597 +
       skoXC1 *
       (skoXC1 *
        (12054 +
         skoXC1 * (skoXC1 * (48035 + skoXC1 * (skoXC1 * 800))))))) <= -55),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (-3194 <
     skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797)))))),
    inference(resolve, [], [refute_0_118, refute_0_52])).

cnf(refute_0_120, plain,
    (skoXC1 *
     (skoXC1 *
      (-755 + skoXC1 * (skoXC1 * (-14382 + skoXC1 * (skoXC1 * -797))))) <=
     -3194),
    inference(decision, [],
              [refute_0_112, refute_0_104, refute_0_21, refute_0_20])).

cnf(refute_0_121, plain, ($false),
    inference(resolve, [], [refute_0_120, refute_0_119])).
% SZS output end CNFRefutation for Nichols-Plot-Simple-Example-1.tptp
