% SZS output start CNFRefutation for nthrt_cbrt-problem-4-weak.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(nthrt_lower_bound_1, axiom,
    (X < 0 | N < 1 | ~ lgen(R, Y, N * X / (1 + N * X - X)) |
     lgen(R, Y, nthrt(N, X)))).

cnf(nthrt_lower_bound_2_small, axiom,
    (X < 0 | N < 1 | 1 <= X |
     ~ lgen(R, Y,
            (N * (N + 1) * 2 ^ (N + N) * X * X +
             N * (N + N - 1) * 2 ^ N * X) /
            (2 ^ (N + N) * (N + N - 1) * (N - 1) * X * X +
             2 ^ (N + 1) * (N + 1) * (N + N - 1) * X + (N + 1))) |
     lgen(R, Y, nthrt(N, X)))).

cnf(nthrt_lower_bound_2, axiom,
    (X < 0 | N < 1 |
     ~ lgen(R, Y,
            (X + X) * N * (N * X + N + N + X - 1) /
            (N + 1 + (4 * N ^ 2 + 2 * N - 2) * X +
             (2 * N ^ 2 - 3 * N + 1) * X ^ 2)) | lgen(R, Y, nthrt(N, X)))).

cnf(nthrt_lower_bound_2_large, axiom,
    (X < 1 | N < 1 |
     ~ lgen(R, Y,
            4 * N * X * ((N + 1) * X + 2 ^ N * (N + N - 1)) /
            ((N + N - 1) * (N - 1) * X * X +
             2 ^ (N + 1) * (N + 1) * (N + N - 1) * X +
             (N + 1) * 2 ^ (N + N))) | lgen(R, Y, nthrt(N, X)))).

cnf(nthrt_lower_bound_3, axiom,
    (X < 0 | N < 1 |
     ~ lgen(R, Y,
            3 * N * X *
            (6 * N ^ 2 - 5 * N + 1 + (12 * N ^ 2 + 2 * N - 2) * X +
             (2 * N ^ 2 + 3 * N + 1) * X ^ 2) /
            ((6 * N ^ 3 - 11 * N ^ 2 + 6 * N - 1) * X ^ 3 +
             (36 * N ^ 3 - 12 * N ^ 2 - 9 * N + 3) * X ^ 2 +
             (18 * N ^ 3 + 21 * N ^ 2 - 3) * X +
             (2 * N ^ 2 + 3 * N + 1))) | lgen(R, Y, nthrt(N, X)))).

cnf(nthrt_upper_bound_1, axiom,
    (X < 0 | N < 1 | ~ lgen(R, (X - 1 + N) / N, Y) |
     lgen(R, nthrt(N, X), Y))).

cnf(nthrt_upper_bound_2_small, axiom,
    (X < 0 | 1 <= X | N < 1 |
     ~ lgen(R,
            (2 ^ (N + N) * (N + 1) * X * X +
             2 ^ (N + 1) * (N + 1) * (N + N - 1) * X +
             (N - 1) * (N + N - 1)) /
            (4 * N * (2 ^ N * (N + N - 1) * X + N + 1)), Y) |
     lgen(R, nthrt(N, X), Y))).

cnf(nthrt_upper_bound_2, axiom,
    (X < 0 | N < 1 |
     ~ lgen(R,
            ((N + 1) * X * X + (4 * N * N + 2 * N - 2) * X +
             (2 * N * N - 3 * N + 1)) /
            (2 * N * ((2 * N - 1) * X + N + 1)), Y) |
     lgen(R, nthrt(N, X), Y))).

cnf(nthrt_upper_bound_2_large, axiom,
    (X < 1 | N < 1 |
     ~ lgen(R,
            ((N + 1) * X * X + 2 ^ (N + 1) * (N + 1) * (N + N - 1) * X +
             2 ^ (N + N) * (N - 1) * (N + N - 1)) /
            (2 ^ N * N * (N + N - 1) * X + 2 ^ (N + N) * N * (N + 1)), Y) |
     lgen(R, nthrt(N, X), Y))).

cnf(nthrt_lower_bound_3, axiom,
    (X < 0 | N < 1 |
     ~ lgen(R,
            ((2 * N * N + 3 * N + 1) * X * X * X +
             (18 * N * N * N + 21 * N * N - 3) * X * X +
             (36 * N * N * N - 12 * N * N - 9 * N + 3) * X +
             (6 * N * N * N - 11 * N * N + 6 * N - 1)) /
            (3 * N *
             ((6 * N * N - 5 * N + 1) * X * X +
              (12 * N * N + 2 * N - 2) * X + (2 * N * N + 3 * N + 1))),
            Y) | lgen(R, nthrt(N, X), Y))).

fof(cbrt_problem_4, conjecture,
    (! [X] :
       ((7/1000 < X & X < 75) => nthrt(3, 2 * X) < 2 * nthrt(3, X)))).

fof(subgoal_0, plain,
    (! [X] : ((7/1000 < X & X < 75) => nthrt(3, 2 * X) < 2 * nthrt(3, X))),
    inference(strip, [], [cbrt_problem_4])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((7/1000 < X & X < 75) => nthrt(3, 2 * X) < 2 * nthrt(3, X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (7/1000 < X & X < 75 & 2 * nthrt(3, X) <= nthrt(3, 2 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 75 & 7/1000 < skoXC1 &
     2 * nthrt(3, skoXC1) <= nthrt(3, 2 * skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (2 * nthrt(3, skoXC1) <= nthrt(3, 2 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (7/1000 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 75),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (5 + skoXC1 * (320 + skoXC1 * 512) <
     skoXC1 * (5760 + skoXC1 * (152064 + skoXC1 * 737280)) /
     (4 + skoXC1 * (320 + skoXC1 * 640)) |
     skoXC1 * (5760 + skoXC1 * (152064 + skoXC1 * 737280)) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) *
     (4 + skoXC1 * (320 + skoXC1 * 640)) |
     4 + skoXC1 * (320 + skoXC1 * 640) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_1, plain,
    ((5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960) <
     skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) |
     skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) *
     (48 + skoXC1 * 960) <= 5 + skoXC1 * (320 + skoXC1 * 512) |
     48 + skoXC1 * 960 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain, (2 * nthrt(3, skoXC1) <= nthrt(3, 2 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain, (nthrt(3, skoXC1) * 2 <= nthrt(3, skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (nthrt(X_000111, X_000113) < X_000114 |
     ~ lgen(1, nthrt(X_000111, X_000113), X_000114)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_5, plain,
    (X_000113 < 0 | X_000111 < 1 |
     ~ lgen(1,
            (2 ^ (X_000111 + X_000111) * (X_000111 + 1) * X_000113 *
             X_000113 +
             2 ^ (X_000111 + 1) * (X_000111 + 1) *
             (X_000111 + X_000111 - 1) * X_000113 +
             (X_000111 - 1) * (X_000111 + X_000111 - 1)) /
            (4 * X_000111 *
             (2 ^ X_000111 * (X_000111 + X_000111 - 1) * X_000113 +
              X_000111 + 1)), X_000114) | 1 <= X_000113 |
     lgen(1, nthrt(X_000111, X_000113), X_000114)),
    inference(subst, [], [nthrt_upper_bound_2_small])).

cnf(refute_0_6, plain,
    (nthrt(X_000111, X_000113) < X_000114 | X_000113 < 0 | X_000111 < 1 |
     ~ lgen(1,
            (2 ^ (X_000111 + X_000111) * (X_000111 + 1) * X_000113 *
             X_000113 +
             2 ^ (X_000111 + 1) * (X_000111 + 1) *
             (X_000111 + X_000111 - 1) * X_000113 +
             (X_000111 - 1) * (X_000111 + X_000111 - 1)) /
            (4 * X_000111 *
             (2 ^ X_000111 * (X_000111 + X_000111 - 1) * X_000113 +
              X_000111 + 1)), X_000114) | 1 <= X_000113),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (nthrt(X_000111, X_000113) < X_000114 | X_000113 < 0 | X_000111 < 1 |
     X_000114 <=
     (1 + X_000111 * (-3 + X_000111 * 2) +
      2 ^ (X_000111 + X_000111) *
      (X_000113 * (X_000113 * (1 + X_000111))) +
      2 ^ (X_000111 + 1) *
      (X_000113 * (-1 + X_000111 * (1 + X_000111 * 2)))) /
     (X_000111 * (4 + X_000111 * 4) +
      2 ^ X_000111 * (X_000113 * (X_000111 * (-4 + X_000111 * 8)))) |
     1 <= X_000113), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (skoXC1 * 2 < 0 | 3 < 1 | nthrt(3, skoXC1 * 2) < nthrt(3, skoXC1) * 2 |
     1 <= skoXC1 * 2 |
     nthrt(3, skoXC1) * 2 <=
     (1 + 3 * (-3 + 3 * 2) +
      2 ^ (3 + 3) * (skoXC1 * 2 * (skoXC1 * 2 * (1 + 3))) +
      2 ^ (3 + 1) * (skoXC1 * 2 * (-1 + 3 * (1 + 3 * 2)))) /
     (3 * (4 + 3 * 4) + 2 ^ 3 * (skoXC1 * 2 * (3 * (-4 + 3 * 8))))),
    inference(subst, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (skoXC1 * 2 < 0 | 3 < 1 | 1 <= skoXC1 * 2 |
     nthrt(3, skoXC1) * 2 <=
     (1 + 3 * (-3 + 3 * 2) +
      2 ^ (3 + 3) * (skoXC1 * 2 * (skoXC1 * 2 * (1 + 3))) +
      2 ^ (3 + 1) * (skoXC1 * 2 * (-1 + 3 * (1 + 3 * 2)))) /
     (3 * (4 + 3 * 4) + 2 ^ 3 * (skoXC1 * 2 * (3 * (-4 + 3 * 8))))),
    inference(resolve, [], [refute_0_3, refute_0_8])).

cnf(refute_0_10, plain,
    (skoXC1 < 0 | 1/2 <= skoXC1 |
     nthrt(3, skoXC1) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) <
     skoXC1 *
     (90720 +
      skoXC1 *
      (979776 +
       skoXC1 * (2614032 + skoXC1 * (1959552 + skoXC1 * 362880)))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 *
     (90720 +
      skoXC1 *
      (979776 +
       skoXC1 * (2614032 + skoXC1 * (1959552 + skoXC1 * 362880)))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_12, plain,
    ((40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) |
     252 + skoXC1 * (2016 + skoXC1 * 1440) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_13, plain,
    (nthrt(X_000115, X_000117) < X_000118 |
     ~ lgen(1, nthrt(X_000115, X_000117), X_000118)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_14, plain,
    (X_000117 < 0 | X_000115 < 1 |
     ~ lgen(1,
            ((2 * X_000115 * X_000115 + 3 * X_000115 + 1) * X_000117 *
             X_000117 * X_000117 +
             (18 * X_000115 * X_000115 * X_000115 +
              21 * X_000115 * X_000115 - 3) * X_000117 * X_000117 +
             (36 * X_000115 * X_000115 * X_000115 -
              12 * X_000115 * X_000115 - 9 * X_000115 + 3) * X_000117 +
             (6 * X_000115 * X_000115 * X_000115 -
              11 * X_000115 * X_000115 + 6 * X_000115 - 1)) /
            (3 * X_000115 *
             ((6 * X_000115 * X_000115 - 5 * X_000115 + 1) * X_000117 *
              X_000117 +
              (12 * X_000115 * X_000115 + 2 * X_000115 - 2) * X_000117 +
              (2 * X_000115 * X_000115 + 3 * X_000115 + 1))), X_000118) |
     lgen(1, nthrt(X_000115, X_000117), X_000118)),
    inference(subst, [], [nthrt_lower_bound_3])).

cnf(refute_0_15, plain,
    (nthrt(X_000115, X_000117) < X_000118 | X_000117 < 0 | X_000115 < 1 |
     ~ lgen(1,
            ((2 * X_000115 * X_000115 + 3 * X_000115 + 1) * X_000117 *
             X_000117 * X_000117 +
             (18 * X_000115 * X_000115 * X_000115 +
              21 * X_000115 * X_000115 - 3) * X_000117 * X_000117 +
             (36 * X_000115 * X_000115 * X_000115 -
              12 * X_000115 * X_000115 - 9 * X_000115 + 3) * X_000117 +
             (6 * X_000115 * X_000115 * X_000115 -
              11 * X_000115 * X_000115 + 6 * X_000115 - 1)) /
            (3 * X_000115 *
             ((6 * X_000115 * X_000115 - 5 * X_000115 + 1) * X_000117 *
              X_000117 +
              (12 * X_000115 * X_000115 + 2 * X_000115 - 2) * X_000117 +
              (2 * X_000115 * X_000115 + 3 * X_000115 + 1))), X_000118)),
    inference(resolve, [], [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (nthrt(X_000115, X_000117) < X_000118 | X_000117 < 0 | X_000115 < 1 |
     X_000118 <=
     (-1 + X_000115 * (6 + X_000115 * (-11 + X_000115 * 6)) +
      X_000117 *
      (3 + X_000115 * (-9 + X_000115 * (-12 + X_000115 * 36)) +
       X_000117 *
       (-3 + X_000115 * (X_000115 * (21 + X_000115 * 18)) +
        X_000117 * (1 + X_000115 * (3 + X_000115 * 2))))) /
     (X_000115 * (3 + X_000115 * (9 + X_000115 * 6)) +
      X_000117 *
      (X_000115 * (-6 + X_000115 * (6 + X_000115 * 36)) +
       X_000117 * (X_000115 * (3 + X_000115 * (-15 + X_000115 * 18)))))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (skoXC1 * 2 < 0 | 3 < 1 | nthrt(3, skoXC1 * 2) < nthrt(3, skoXC1) * 2 |
     nthrt(3, skoXC1) * 2 <=
     (-1 + 3 * (6 + 3 * (-11 + 3 * 6)) +
      skoXC1 * 2 *
      (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
       skoXC1 * 2 *
       (-3 + 3 * (3 * (21 + 3 * 18)) +
        skoXC1 * 2 * (1 + 3 * (3 + 3 * 2))))) /
     (3 * (3 + 3 * (9 + 3 * 6)) +
      skoXC1 * 2 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * 2 * (3 * (3 + 3 * (-15 + 3 * 18)))))),
    inference(subst, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 * 2 < 0 | 3 < 1 |
     nthrt(3, skoXC1) * 2 <=
     (-1 + 3 * (6 + 3 * (-11 + 3 * 6)) +
      skoXC1 * 2 *
      (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
       skoXC1 * 2 *
       (-3 + 3 * (3 * (21 + 3 * 18)) +
        skoXC1 * 2 * (1 + 3 * (3 + 3 * 2))))) /
     (3 * (3 + 3 * (9 + 3 * 6)) +
      skoXC1 * 2 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * 2 * (3 * (3 + 3 * (-15 + 3 * 18)))))),
    inference(resolve, [], [refute_0_3, refute_0_17])).

cnf(refute_0_19, plain,
    (skoXC1 < 0 |
     nthrt(3, skoXC1) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (7/1000 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_21, plain, (skoXC1 < 75),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_22, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (nthrt(3, skoXC1) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(resolve, [], [refute_0_22, refute_0_19])).

cnf(refute_0_24, plain,
    (X_000104 < nthrt(X_000101, X_000103) |
     ~ lgen(1, X_000104, nthrt(X_000101, X_000103))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_25, plain,
    (X_000103 < 0 | X_000101 < 1 |
     ~ lgen(1, X_000104,
            3 * X_000101 * X_000103 *
            (6 * X_000101 ^ 2 - 5 * X_000101 + 1 +
             (12 * X_000101 ^ 2 + 2 * X_000101 - 2) * X_000103 +
             (2 * X_000101 ^ 2 + 3 * X_000101 + 1) * X_000103 ^ 2) /
            ((6 * X_000101 ^ 3 - 11 * X_000101 ^ 2 + 6 * X_000101 - 1) *
             X_000103 ^ 3 +
             (36 * X_000101 ^ 3 - 12 * X_000101 ^ 2 - 9 * X_000101 + 3) *
             X_000103 ^ 2 +
             (18 * X_000101 ^ 3 + 21 * X_000101 ^ 2 - 3) * X_000103 +
             (2 * X_000101 ^ 2 + 3 * X_000101 + 1))) |
     lgen(1, X_000104, nthrt(X_000101, X_000103))),
    inference(subst, [], [nthrt_lower_bound_3])).

cnf(refute_0_26, plain,
    (X_000103 < 0 | X_000101 < 1 | X_000104 < nthrt(X_000101, X_000103) |
     ~ lgen(1, X_000104,
            3 * X_000101 * X_000103 *
            (6 * X_000101 ^ 2 - 5 * X_000101 + 1 +
             (12 * X_000101 ^ 2 + 2 * X_000101 - 2) * X_000103 +
             (2 * X_000101 ^ 2 + 3 * X_000101 + 1) * X_000103 ^ 2) /
            ((6 * X_000101 ^ 3 - 11 * X_000101 ^ 2 + 6 * X_000101 - 1) *
             X_000103 ^ 3 +
             (36 * X_000101 ^ 3 - 12 * X_000101 ^ 2 - 9 * X_000101 + 3) *
             X_000103 ^ 2 +
             (18 * X_000101 ^ 3 + 21 * X_000101 ^ 2 - 3) * X_000103 +
             (2 * X_000101 ^ 2 + 3 * X_000101 + 1)))),
    inference(resolve, [], [refute_0_25, refute_0_24])).

cnf(refute_0_27, plain,
    (X_000103 < 0 | X_000101 < 1 | X_000104 < nthrt(X_000101, X_000103) |
     X_000103 *
     (X_000101 * (3 + X_000101 * (-15 + X_000101 * 18)) +
      X_000103 *
      (X_000101 * (-6 + X_000101 * (6 + X_000101 * 36)) +
       X_000103 * (X_000101 * (3 + X_000101 * (9 + X_000101 * 6))))) /
     (1 + X_000101 * (3 + X_000101 * 2) +
      X_000103 *
      (-3 + X_000101 * (X_000101 * (21 + X_000101 * 18)) +
       X_000103 *
       (3 + X_000101 * (-9 + X_000101 * (-12 + X_000101 * 36)) +
        X_000103 *
        (-1 + X_000101 * (6 + X_000101 * (-11 + X_000101 * 6)))))) <=
     X_000104), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 < 0 | 3 < 1 |
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) < nthrt(3, skoXC1) |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(resolve, [], [refute_0_23, refute_0_28])).

cnf(refute_0_30, plain,
    (skoXC1 < 0 |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(resolve, [], [refute_0_22, refute_0_30])).

cnf(refute_0_32, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) |
     252 + skoXC1 * (2016 + skoXC1 * 1440) <= 0),
    inference(resolve, [], [refute_0_31, refute_0_12])).

cnf(refute_0_33, plain,
    (skoXC1 * (2016 + skoXC1 * 1440) <= -252 |
     skoXC1 *
     (90720 +
      skoXC1 *
      (979776 +
       skoXC1 * (2614032 + skoXC1 * (1959552 + skoXC1 * 362880)))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (320 + skoXC1 * (320 + skoXC1 * 8) <
     skoXC1 * (368640 + skoXC1 * (152064 + skoXC1 * 11520)) /
     (256 + skoXC1 * (320 + skoXC1 * 10)) |
     skoXC1 * (368640 + skoXC1 * (152064 + skoXC1 * 11520)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) *
     (256 + skoXC1 * (320 + skoXC1 * 10)) |
     256 + skoXC1 * (320 + skoXC1 * 10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_35, plain,
    ((320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     skoXC1 * (480 + skoXC1 * 48) / (256 + skoXC1 * (320 + skoXC1 * 10)) |
     skoXC1 * (480 + skoXC1 * 48) / (256 + skoXC1 * (320 + skoXC1 * 10)) *
     (768 + skoXC1 * 240) <= 320 + skoXC1 * (320 + skoXC1 * 8) |
     768 + skoXC1 * 240 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_36, plain,
    (nthrt(X_000074, X_000076) < X_000077 |
     ~ lgen(1, nthrt(X_000074, X_000076), X_000077)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_37, plain,
    (X_000074 < 1 | X_000076 < 1 |
     ~ lgen(1,
            ((X_000074 + 1) * X_000076 * X_000076 +
             2 ^ (X_000074 + 1) * (X_000074 + 1) *
             (X_000074 + X_000074 - 1) * X_000076 +
             2 ^ (X_000074 + X_000074) * (X_000074 - 1) *
             (X_000074 + X_000074 - 1)) /
            (2 ^ X_000074 * X_000074 * (X_000074 + X_000074 - 1) *
             X_000076 +
             2 ^ (X_000074 + X_000074) * X_000074 * (X_000074 + 1)),
            X_000077) | lgen(1, nthrt(X_000074, X_000076), X_000077)),
    inference(subst, [], [nthrt_upper_bound_2_large])).

cnf(refute_0_38, plain,
    (nthrt(X_000074, X_000076) < X_000077 | X_000074 < 1 | X_000076 < 1 |
     ~ lgen(1,
            ((X_000074 + 1) * X_000076 * X_000076 +
             2 ^ (X_000074 + 1) * (X_000074 + 1) *
             (X_000074 + X_000074 - 1) * X_000076 +
             2 ^ (X_000074 + X_000074) * (X_000074 - 1) *
             (X_000074 + X_000074 - 1)) /
            (2 ^ X_000074 * X_000074 * (X_000074 + X_000074 - 1) *
             X_000076 +
             2 ^ (X_000074 + X_000074) * X_000074 * (X_000074 + 1)),
            X_000077)),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (nthrt(X_000074, X_000076) < X_000077 | X_000074 < 1 | X_000076 < 1 |
     X_000077 <=
     (X_000076 * (X_000076 * (1 + X_000074)) +
      2 ^ (X_000074 + X_000074) * (1 + X_000074 * (-3 + X_000074 * 2)) +
      2 ^ (X_000074 + 1) *
      (X_000076 * (-1 + X_000074 * (1 + X_000074 * 2)))) /
     (2 ^ X_000074 * (X_000076 * (X_000074 * (-1 + X_000074 * 2))) +
      2 ^ (X_000074 + X_000074) * (X_000074 * (1 + X_000074)))),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (3 < 1 | skoXC1 * 2 < 1 | nthrt(3, skoXC1 * 2) < nthrt(3, skoXC1) * 2 |
     nthrt(3, skoXC1) * 2 <=
     (skoXC1 * 2 * (skoXC1 * 2 * (1 + 3)) +
      2 ^ (3 + 3) * (1 + 3 * (-3 + 3 * 2)) +
      2 ^ (3 + 1) * (skoXC1 * 2 * (-1 + 3 * (1 + 3 * 2)))) /
     (2 ^ 3 * (skoXC1 * 2 * (3 * (-1 + 3 * 2))) +
      2 ^ (3 + 3) * (3 * (1 + 3)))), inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (3 < 1 | skoXC1 * 2 < 1 |
     nthrt(3, skoXC1) * 2 <=
     (skoXC1 * 2 * (skoXC1 * 2 * (1 + 3)) +
      2 ^ (3 + 3) * (1 + 3 * (-3 + 3 * 2)) +
      2 ^ (3 + 1) * (skoXC1 * 2 * (-1 + 3 * (1 + 3 * 2)))) /
     (2 ^ 3 * (skoXC1 * 2 * (3 * (-1 + 3 * 2))) +
      2 ^ (3 + 3) * (3 * (1 + 3)))),
    inference(resolve, [], [refute_0_3, refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 < 1/2 |
     nthrt(3, skoXC1) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (X_000073 < nthrt(X_000070, X_000072) |
     ~ lgen(1, X_000073, nthrt(X_000070, X_000072))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain,
    (X_000070 < 1 | X_000072 < 1 |
     ~ lgen(1, X_000073,
            4 * X_000070 * X_000072 *
            ((X_000070 + 1) * X_000072 +
             2 ^ X_000070 * (X_000070 + X_000070 - 1)) /
            ((X_000070 + X_000070 - 1) * (X_000070 - 1) * X_000072 *
             X_000072 +
             2 ^ (X_000070 + 1) * (X_000070 + 1) *
             (X_000070 + X_000070 - 1) * X_000072 +
             (X_000070 + 1) * 2 ^ (X_000070 + X_000070))) |
     lgen(1, X_000073, nthrt(X_000070, X_000072))),
    inference(subst, [], [nthrt_lower_bound_2_large])).

cnf(refute_0_45, plain,
    (X_000070 < 1 | X_000072 < 1 | X_000073 < nthrt(X_000070, X_000072) |
     ~ lgen(1, X_000073,
            4 * X_000070 * X_000072 *
            ((X_000070 + 1) * X_000072 +
             2 ^ X_000070 * (X_000070 + X_000070 - 1)) /
            ((X_000070 + X_000070 - 1) * (X_000070 - 1) * X_000072 *
             X_000072 +
             2 ^ (X_000070 + 1) * (X_000070 + 1) *
             (X_000070 + X_000070 - 1) * X_000072 +
             (X_000070 + 1) * 2 ^ (X_000070 + X_000070)))),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (X_000070 < 1 | X_000072 < 1 | X_000073 < nthrt(X_000070, X_000072) |
     (X_000072 * (X_000072 * (X_000070 * (4 + X_000070 * 4))) +
      2 ^ X_000070 * (X_000072 * (X_000070 * (-4 + X_000070 * 8)))) /
     (X_000072 * (X_000072 * (1 + X_000070 * (-3 + X_000070 * 2))) +
      2 ^ (X_000070 + X_000070) * (1 + X_000070) +
      2 ^ (X_000070 + 1) *
      (X_000072 * (-1 + X_000070 * (1 + X_000070 * 2)))) <= X_000073),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (3 < 1 | skoXC1 < 1 |
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     nthrt(3, skoXC1) |
     (skoXC1 * (skoXC1 * (3 * (4 + 3 * 4))) +
      2 ^ 3 * (skoXC1 * (3 * (-4 + 3 * 8)))) /
     (skoXC1 * (skoXC1 * (1 + 3 * (-3 + 3 * 2))) + 2 ^ (3 + 3) * (1 + 3) +
      2 ^ (3 + 1) * (skoXC1 * (-1 + 3 * (1 + 3 * 2)))) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (skoXC1 < 1/2 | 3 < 1 | skoXC1 < 1 |
     (skoXC1 * (skoXC1 * (3 * (4 + 3 * 4))) +
      2 ^ 3 * (skoXC1 * (3 * (-4 + 3 * 8)))) /
     (skoXC1 * (skoXC1 * (1 + 3 * (-3 + 3 * 2))) + 2 ^ (3 + 3) * (1 + 3) +
      2 ^ (3 + 1) * (skoXC1 * (-1 + 3 * (1 + 3 * 2)))) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_42, refute_0_47])).

cnf(refute_0_49, plain,
    (skoXC1 < 1/2 | skoXC1 < 1 |
     skoXC1 * (480 + skoXC1 * 48) / (256 + skoXC1 * (320 + skoXC1 * 10)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (320 + skoXC1 * (320 + skoXC1 * 8) <
     skoXC1 * (23040 + skoXC1 * (25632 + skoXC1 * 5760)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (23040 + skoXC1 * (25632 + skoXC1 * 5760)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) *
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_51, plain,
    ((320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (768 + skoXC1 * 240) <= 320 + skoXC1 * (320 + skoXC1 * 8) |
     768 + skoXC1 * 240 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_52, plain,
    (X_000065 < nthrt(X_000062, X_000064) |
     ~ lgen(1, X_000065, nthrt(X_000062, X_000064))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_53, plain,
    (X_000064 < 0 | X_000062 < 1 |
     ~ lgen(1, X_000065,
            (X_000064 + X_000064) * X_000062 *
            (X_000062 * X_000064 + X_000062 + X_000062 + X_000064 - 1) /
            (X_000062 + 1 +
             (4 * X_000062 ^ 2 + 2 * X_000062 - 2) * X_000064 +
             (2 * X_000062 ^ 2 - 3 * X_000062 + 1) * X_000064 ^ 2)) |
     lgen(1, X_000065, nthrt(X_000062, X_000064))),
    inference(subst, [], [nthrt_lower_bound_2])).

cnf(refute_0_54, plain,
    (X_000064 < 0 | X_000062 < 1 | X_000065 < nthrt(X_000062, X_000064) |
     ~ lgen(1, X_000065,
            (X_000064 + X_000064) * X_000062 *
            (X_000062 * X_000064 + X_000062 + X_000062 + X_000064 - 1) /
            (X_000062 + 1 +
             (4 * X_000062 ^ 2 + 2 * X_000062 - 2) * X_000064 +
             (2 * X_000062 ^ 2 - 3 * X_000062 + 1) * X_000064 ^ 2))),
    inference(resolve, [], [refute_0_53, refute_0_52])).

cnf(refute_0_55, plain,
    (X_000064 < 0 | X_000062 < 1 | X_000065 < nthrt(X_000062, X_000064) |
     X_000064 *
     (X_000062 * (-2 + X_000062 * 4) +
      X_000064 * (X_000062 * (2 + X_000062 * 2))) /
     (1 + X_000062 +
      X_000064 *
      (-2 + X_000062 * (2 + X_000062 * 4) +
       X_000064 * (1 + X_000062 * (-3 + X_000062 * 2)))) <= X_000065),
    inference(arithmetic, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (skoXC1 < 0 | 3 < 1 |
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     nthrt(3, skoXC1) |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 | 3 < 1 |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_42, refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (5 + skoXC1 * (40 + skoXC1 * 8) <
     skoXC1 * (720 + skoXC1 * (2376 + skoXC1 * 1440)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (720 + skoXC1 * (2376 + skoXC1 * 1440)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) * (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_60, plain,
    ((5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (24 + skoXC1 * 60) <= 5 + skoXC1 * (40 + skoXC1 * 8) |
     24 + skoXC1 * 60 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_61, plain,
    (nthrt(X_000066, X_000068) < X_000069 |
     ~ lgen(1, nthrt(X_000066, X_000068), X_000069)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_62, plain,
    (X_000068 < 0 | X_000066 < 1 |
     ~ lgen(1,
            ((X_000066 + 1) * X_000068 * X_000068 +
             (4 * X_000066 * X_000066 + 2 * X_000066 - 2) * X_000068 +
             (2 * X_000066 * X_000066 - 3 * X_000066 + 1)) /
            (2 * X_000066 *
             ((2 * X_000066 - 1) * X_000068 + X_000066 + 1)), X_000069) |
     lgen(1, nthrt(X_000066, X_000068), X_000069)),
    inference(subst, [], [nthrt_upper_bound_2])).

cnf(refute_0_63, plain,
    (nthrt(X_000066, X_000068) < X_000069 | X_000068 < 0 | X_000066 < 1 |
     ~ lgen(1,
            ((X_000066 + 1) * X_000068 * X_000068 +
             (4 * X_000066 * X_000066 + 2 * X_000066 - 2) * X_000068 +
             (2 * X_000066 * X_000066 - 3 * X_000066 + 1)) /
            (2 * X_000066 *
             ((2 * X_000066 - 1) * X_000068 + X_000066 + 1)), X_000069)),
    inference(resolve, [], [refute_0_62, refute_0_61])).

cnf(refute_0_64, plain,
    (nthrt(X_000066, X_000068) < X_000069 | X_000068 < 0 | X_000066 < 1 |
     X_000069 <=
     (1 + X_000066 * (-3 + X_000066 * 2) +
      X_000068 *
      (-2 + X_000066 * (2 + X_000066 * 4) + X_000068 * (1 + X_000066))) /
     (X_000066 * (2 + X_000066 * 2) +
      X_000068 * (X_000066 * (-2 + X_000066 * 4)))),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 * 2 < 0 | 3 < 1 | nthrt(3, skoXC1 * 2) < nthrt(3, skoXC1) * 2 |
     nthrt(3, skoXC1) * 2 <=
     (1 + 3 * (-3 + 3 * 2) +
      skoXC1 * 2 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * 2 * (1 + 3))) /
     (3 * (2 + 3 * 2) + skoXC1 * 2 * (3 * (-2 + 3 * 4)))),
    inference(subst, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 * 2 < 0 | 3 < 1 |
     nthrt(3, skoXC1) * 2 <=
     (1 + 3 * (-3 + 3 * 2) +
      skoXC1 * 2 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * 2 * (1 + 3))) /
     (3 * (2 + 3 * 2) + skoXC1 * 2 * (3 * (-2 + 3 * 4)))),
    inference(resolve, [], [refute_0_3, refute_0_65])).

cnf(refute_0_67, plain,
    (skoXC1 < 0 |
     nthrt(3, skoXC1) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (nthrt(3, skoXC1) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_22, refute_0_67])).

cnf(refute_0_69, plain,
    (skoXC1 < 0 | 3 < 1 |
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     nthrt(3, skoXC1) |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_70, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain,
    (skoXC1 < 0 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_22, refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (24 + skoXC1 * 60) <= 5 + skoXC1 * (40 + skoXC1 * 8) |
     24 + skoXC1 * 60 <= 0),
    inference(resolve, [], [refute_0_72, refute_0_60])).

cnf(refute_0_74, plain,
    (skoXC1 * (720 + skoXC1 * (2376 + skoXC1 * 1440)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <= 5 + skoXC1 * (40 + skoXC1 * 8) |
     skoXC1 <= -2/5), inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (1/3 + skoXC1 * 1/3 < skoXC1 * 3 / (1 + skoXC1 * 2) |
     skoXC1 * 3 <= (1/3 + skoXC1 * 1/3) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_76, plain,
    (nthrt(X_000042, X_000044) < X_000045 |
     ~ lgen(1, nthrt(X_000042, X_000044), X_000045)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_77, plain,
    (X_000044 < 0 | X_000042 < 1 |
     ~ lgen(1, (X_000044 - 1 + X_000042) / X_000042, X_000045) |
     lgen(1, nthrt(X_000042, X_000044), X_000045)),
    inference(subst, [], [nthrt_upper_bound_1])).

cnf(refute_0_78, plain,
    (nthrt(X_000042, X_000044) < X_000045 | X_000044 < 0 | X_000042 < 1 |
     ~ lgen(1, (X_000044 - 1 + X_000042) / X_000042, X_000045)),
    inference(resolve, [], [refute_0_77, refute_0_76])).

cnf(refute_0_79, plain,
    (nthrt(X_000042, X_000044) < X_000045 | X_000044 < 0 | X_000042 < 1 |
     X_000045 <= (-1 + X_000042 + X_000044) / X_000042),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (skoXC1 * 2 < 0 | 3 < 1 | nthrt(3, skoXC1 * 2) < nthrt(3, skoXC1) * 2 |
     nthrt(3, skoXC1) * 2 <= (-1 + 3 + skoXC1 * 2) / 3),
    inference(subst, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (skoXC1 * 2 < 0 | 3 < 1 |
     nthrt(3, skoXC1) * 2 <= (-1 + 3 + skoXC1 * 2) / 3),
    inference(resolve, [], [refute_0_3, refute_0_80])).

cnf(refute_0_82, plain,
    (skoXC1 < 0 | nthrt(3, skoXC1) <= 1/3 + skoXC1 * 1/3),
    inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain, (nthrt(3, skoXC1) <= 1/3 + skoXC1 * 1/3),
    inference(resolve, [], [refute_0_22, refute_0_82])).

cnf(refute_0_84, plain,
    (X_000049 < nthrt(X_000046, X_000048) |
     ~ lgen(1, X_000049, nthrt(X_000046, X_000048))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_85, plain,
    (X_000048 < 0 | X_000046 < 1 |
     ~ lgen(1, X_000049,
            X_000046 * X_000048 / (1 + X_000046 * X_000048 - X_000048)) |
     lgen(1, X_000049, nthrt(X_000046, X_000048))),
    inference(subst, [], [nthrt_lower_bound_1])).

cnf(refute_0_86, plain,
    (X_000048 < 0 | X_000046 < 1 | X_000049 < nthrt(X_000046, X_000048) |
     ~ lgen(1, X_000049,
            X_000046 * X_000048 / (1 + X_000046 * X_000048 - X_000048))),
    inference(resolve, [], [refute_0_85, refute_0_84])).

cnf(refute_0_87, plain,
    (X_000048 < 0 | X_000046 < 1 | X_000049 < nthrt(X_000046, X_000048) |
     X_000048 * X_000046 / (1 + X_000048 * (-1 + X_000046)) <= X_000049),
    inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 < 0 | 3 < 1 | 1/3 + skoXC1 * 1/3 < nthrt(3, skoXC1) |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <= 1/3 + skoXC1 * 1/3),
    inference(subst, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <= 1/3 + skoXC1 * 1/3),
    inference(resolve, [], [refute_0_83, refute_0_88])).

cnf(refute_0_90, plain,
    (skoXC1 < 0 | skoXC1 * 3 / (1 + skoXC1 * 2) <= 1/3 + skoXC1 * 1/3),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (skoXC1 * 3 / (1 + skoXC1 * 2) <= 1/3 + skoXC1 * 1/3),
    inference(resolve, [], [refute_0_22, refute_0_90])).

cnf(refute_0_92, plain,
    (skoXC1 * 3 <= (1/3 + skoXC1 * 1/3) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(resolve, [], [refute_0_91, refute_0_75])).

cnf(refute_0_93, plain,
    (skoXC1 * (2 + skoXC1 * -2/3) <= 1/3 | skoXC1 <= -1/2),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (-1/2 < skoXC1 | skoXC1 * (2 + skoXC1 * -2/3) <= 1/3),
    inference(decision, [], [refute_0_20, refute_0_21])).

cnf(refute_0_95, plain, (skoXC1 * (2 + skoXC1 * -2/3) <= 1/3),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain, (-2/5 < skoXC1),
    inference(decision, [], [refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_97, plain,
    (skoXC1 * (720 + skoXC1 * (2376 + skoXC1 * 1440)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <= 5 + skoXC1 * (40 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_74, refute_0_96])).

cnf(refute_0_98, plain,
    (skoXC1 * (720 + skoXC1 * (2376 + skoXC1 * 1440)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) * (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(resolve, [], [refute_0_97, refute_0_59])).

cnf(refute_0_99, plain,
    (skoXC1 * (40 + skoXC1 * 10) <= -4 |
     skoXC1 * (360 + skoXC1 * (694 + skoXC1 * (720 + skoXC1 * -80))) <=
     20), inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (1/3 + skoXC1 * 1/3 <
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (30 + skoXC1 * 24) <=
     (1/3 + skoXC1 * 1/3) * (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_101, plain,
    (skoXC1 < 0 | 3 < 1 | 1/3 + skoXC1 * 1/3 < nthrt(3, skoXC1) |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     1/3 + skoXC1 * 1/3), inference(subst, [], [refute_0_55])).

cnf(refute_0_102, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     1/3 + skoXC1 * 1/3),
    inference(resolve, [], [refute_0_83, refute_0_101])).

cnf(refute_0_103, plain,
    (skoXC1 < 0 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     1/3 + skoXC1 * 1/3), inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     1/3 + skoXC1 * 1/3),
    inference(resolve, [], [refute_0_22, refute_0_103])).

cnf(refute_0_105, plain,
    (skoXC1 * (30 + skoXC1 * 24) <=
     (1/3 + skoXC1 * 1/3) * (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(resolve, [], [refute_0_104, refute_0_100])).

cnf(refute_0_106, plain,
    (skoXC1 * (46/3 + skoXC1 * (22/3 + skoXC1 * -10/3)) <= 4/3 |
     skoXC1 * (40 + skoXC1 * 10) <= -4),
    inference(arithmetic, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (5 + skoXC1 * (40 + skoXC1 * 8) <
     skoXC1 * (72 + skoXC1 * 180) / (1 + skoXC1 * 2) |
     skoXC1 * (72 + skoXC1 * 180) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_108, plain,
    ((5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     skoXC1 * 3 / (1 + skoXC1 * 2) |
     skoXC1 * 3 / (1 + skoXC1 * 2) * (24 + skoXC1 * 60) <=
     5 + skoXC1 * (40 + skoXC1 * 8) | 24 + skoXC1 * 60 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_109, plain,
    (skoXC1 < 0 | 3 < 1 |
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     nthrt(3, skoXC1) |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(subst, [], [refute_0_87])).

cnf(refute_0_110, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_68, refute_0_109])).

cnf(refute_0_111, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (1 + skoXC1 * 2) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (skoXC1 * 3 / (1 + skoXC1 * 2) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_22, refute_0_111])).

cnf(refute_0_113, plain,
    (skoXC1 * 3 / (1 + skoXC1 * 2) * (24 + skoXC1 * 60) <=
     5 + skoXC1 * (40 + skoXC1 * 8) | 24 + skoXC1 * 60 <= 0),
    inference(resolve, [], [refute_0_112, refute_0_108])).

cnf(refute_0_114, plain,
    (skoXC1 * (72 + skoXC1 * 180) / (1 + skoXC1 * 2) <=
     5 + skoXC1 * (40 + skoXC1 * 8) | skoXC1 <= -2/5),
    inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 * (72 + skoXC1 * 180) / (1 + skoXC1 * 2) <=
     5 + skoXC1 * (40 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_114, refute_0_96])).

cnf(refute_0_116, plain,
    (skoXC1 * (72 + skoXC1 * 180) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(resolve, [], [refute_0_115, refute_0_107])).

cnf(refute_0_117, plain,
    (skoXC1 * (22 + skoXC1 * (92 + skoXC1 * -16)) <= 5 | skoXC1 <= -1/2),
    inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (-1/2 < skoXC1 | skoXC1 * (22 + skoXC1 * (92 + skoXC1 * -16)) <= 5),
    inference(decision, [], [refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_119, plain,
    (skoXC1 * (22 + skoXC1 * (92 + skoXC1 * -16)) <= 5),
    inference(resolve, [], [refute_0_117, refute_0_118])).

cnf(refute_0_120, plain,
    (-4 < skoXC1 * (40 + skoXC1 * 10) |
     skoXC1 * (46/3 + skoXC1 * (22/3 + skoXC1 * -10/3)) <= 4/3),
    inference(decision, [],
              [refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_121, plain,
    (skoXC1 * (46/3 + skoXC1 * (22/3 + skoXC1 * -10/3)) <= 4/3),
    inference(resolve, [], [refute_0_106, refute_0_120])).

cnf(refute_0_122, plain,
    (-4 < skoXC1 * (40 + skoXC1 * 10) |
     skoXC1 * (360 + skoXC1 * (694 + skoXC1 * (720 + skoXC1 * -80))) <=
     20),
    inference(decision, [],
              [refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_123, plain,
    (skoXC1 * (360 + skoXC1 * (694 + skoXC1 * (720 + skoXC1 * -80))) <=
     20), inference(resolve, [], [refute_0_99, refute_0_122])).

cnf(refute_0_124, plain, (skoXC1 < 1/2 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_123, refute_0_121, refute_0_119, refute_0_95,
               refute_0_20, refute_0_21])).

cnf(refute_0_125, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_124, refute_0_58])).

cnf(refute_0_126, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (768 + skoXC1 * 240) <= 320 + skoXC1 * (320 + skoXC1 * 8) |
     768 + skoXC1 * 240 <= 0),
    inference(resolve, [], [refute_0_125, refute_0_51])).

cnf(refute_0_127, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (23040 + skoXC1 * (25632 + skoXC1 * 5760)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     320 + skoXC1 * (320 + skoXC1 * 8) | skoXC1 <= -16/5),
    inference(arithmetic, [], [refute_0_126])).

cnf(refute_0_128, plain, (skoXC1 < 1/2 | -16/5 < skoXC1),
    inference(decision, [],
              [refute_0_123, refute_0_121, refute_0_119, refute_0_95,
               refute_0_20, refute_0_21])).

cnf(refute_0_129, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (23040 + skoXC1 * (25632 + skoXC1 * 5760)) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     320 + skoXC1 * (320 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_127, refute_0_128])).

cnf(refute_0_130, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (23040 + skoXC1 * (25632 + skoXC1 * 5760)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) *
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(resolve, [], [refute_0_129, refute_0_50])).

cnf(refute_0_131, plain,
    (skoXC1 < 1/2 | skoXC1 * (40 + skoXC1 * 10) <= -4 |
     skoXC1 * (8960 + skoXC1 * (9600 + skoXC1 * (2240 + skoXC1 * -80))) <=
     1280), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (320 + skoXC1 * (320 + skoXC1 * 8) <
     skoXC1 * (2304 + skoXC1 * 720) / (1 + skoXC1 * 2) |
     skoXC1 * (2304 + skoXC1 * 720) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_133, plain,
    ((320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     skoXC1 * 3 / (1 + skoXC1 * 2) |
     skoXC1 * 3 / (1 + skoXC1 * 2) * (768 + skoXC1 * 240) <=
     320 + skoXC1 * (320 + skoXC1 * 8) | 768 + skoXC1 * 240 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_134, plain,
    (skoXC1 < 0 | 3 < 1 |
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240) <
     nthrt(3, skoXC1) |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(subst, [], [refute_0_87])).

cnf(refute_0_135, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 | 3 < 1 |
     skoXC1 * 3 / (1 + skoXC1 * (-1 + 3)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_42, refute_0_134])).

cnf(refute_0_136, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 |
     skoXC1 * 3 / (1 + skoXC1 * 2) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(arithmetic, [], [refute_0_135])).

cnf(refute_0_137, plain,
    (skoXC1 < 1/2 |
     skoXC1 * 3 / (1 + skoXC1 * 2) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_124, refute_0_136])).

cnf(refute_0_138, plain,
    (skoXC1 < 1/2 |
     skoXC1 * 3 / (1 + skoXC1 * 2) * (768 + skoXC1 * 240) <=
     320 + skoXC1 * (320 + skoXC1 * 8) | 768 + skoXC1 * 240 <= 0),
    inference(resolve, [], [refute_0_137, refute_0_133])).

cnf(refute_0_139, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (2304 + skoXC1 * 720) / (1 + skoXC1 * 2) <=
     320 + skoXC1 * (320 + skoXC1 * 8) | skoXC1 <= -16/5),
    inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (2304 + skoXC1 * 720) / (1 + skoXC1 * 2) <=
     320 + skoXC1 * (320 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_139, refute_0_128])).

cnf(refute_0_141, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (2304 + skoXC1 * 720) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) * (1 + skoXC1 * 2) |
     1 + skoXC1 * 2 <= 0),
    inference(resolve, [], [refute_0_140, refute_0_132])).

cnf(refute_0_142, plain,
    (skoXC1 < 1/2 | skoXC1 * (1344 + skoXC1 * (72 + skoXC1 * -16)) <= 320 |
     skoXC1 <= -1/2), inference(arithmetic, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (1/2 <= skoXC1 |
     skoXC1 * (1344 + skoXC1 * (72 + skoXC1 * -16)) <= 320 |
     skoXC1 <= -1/2),
    inference(decision, [],
              [refute_0_123, refute_0_121, refute_0_119, refute_0_95,
               refute_0_20, refute_0_21])).

cnf(refute_0_144, plain,
    (skoXC1 * (1344 + skoXC1 * (72 + skoXC1 * -16)) <= 320 |
     skoXC1 <= -1/2),
    inference(resolve, [], [refute_0_143, refute_0_142])).

cnf(refute_0_145, plain,
    (-1/2 < skoXC1 |
     skoXC1 * (1344 + skoXC1 * (72 + skoXC1 * -16)) <= 320),
    inference(decision, [],
              [refute_0_123, refute_0_121, refute_0_119, refute_0_95,
               refute_0_20, refute_0_21])).

cnf(refute_0_146, plain,
    (skoXC1 * (1344 + skoXC1 * (72 + skoXC1 * -16)) <= 320),
    inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (skoXC1 < 1/2 | -4 < skoXC1 * (40 + skoXC1 * 10) |
     skoXC1 * (8960 + skoXC1 * (9600 + skoXC1 * (2240 + skoXC1 * -80))) <=
     1280),
    inference(decision, [],
              [refute_0_146, refute_0_123, refute_0_121, refute_0_119,
               refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_148, plain,
    (skoXC1 < 1/2 |
     skoXC1 * (8960 + skoXC1 * (9600 + skoXC1 * (2240 + skoXC1 * -80))) <=
     1280), inference(resolve, [], [refute_0_131, refute_0_147])).

cnf(refute_0_149, plain,
    (1/2 <= skoXC1 |
     skoXC1 * (8960 + skoXC1 * (9600 + skoXC1 * (2240 + skoXC1 * -80))) <=
     1280),
    inference(decision, [],
              [refute_0_146, refute_0_123, refute_0_121, refute_0_119,
               refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_150, plain,
    (skoXC1 * (8960 + skoXC1 * (9600 + skoXC1 * (2240 + skoXC1 * -80))) <=
     1280), inference(resolve, [], [refute_0_149, refute_0_148])).

cnf(refute_0_151, plain, (skoXC1 < 1 | 1/2 <= skoXC1),
    inference(decision, [],
              [refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_152, plain,
    (skoXC1 < 1 |
     skoXC1 * (480 + skoXC1 * 48) / (256 + skoXC1 * (320 + skoXC1 * 10)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) / (768 + skoXC1 * 240)),
    inference(resolve, [], [refute_0_151, refute_0_49])).

cnf(refute_0_153, plain,
    (skoXC1 < 1 |
     skoXC1 * (480 + skoXC1 * 48) / (256 + skoXC1 * (320 + skoXC1 * 10)) *
     (768 + skoXC1 * 240) <= 320 + skoXC1 * (320 + skoXC1 * 8) |
     768 + skoXC1 * 240 <= 0),
    inference(resolve, [], [refute_0_152, refute_0_35])).

cnf(refute_0_154, plain,
    (skoXC1 < 1 |
     skoXC1 * (368640 + skoXC1 * (152064 + skoXC1 * 11520)) /
     (256 + skoXC1 * (320 + skoXC1 * 10)) <=
     320 + skoXC1 * (320 + skoXC1 * 8) | skoXC1 <= -16/5),
    inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain, (skoXC1 < 1 | -16/5 < skoXC1),
    inference(decision, [],
              [refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_156, plain,
    (skoXC1 < 1 |
     skoXC1 * (368640 + skoXC1 * (152064 + skoXC1 * 11520)) /
     (256 + skoXC1 * (320 + skoXC1 * 10)) <=
     320 + skoXC1 * (320 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (skoXC1 < 1 |
     skoXC1 * (368640 + skoXC1 * (152064 + skoXC1 * 11520)) <=
     (320 + skoXC1 * (320 + skoXC1 * 8)) *
     (256 + skoXC1 * (320 + skoXC1 * 10)) |
     256 + skoXC1 * (320 + skoXC1 * 10) <= 0),
    inference(resolve, [], [refute_0_156, refute_0_34])).

cnf(refute_0_158, plain,
    (skoXC1 < 1 | skoXC1 * (320 + skoXC1 * 10) <= -256 |
     skoXC1 *
     (184320 + skoXC1 * (44416 + skoXC1 * (5760 + skoXC1 * -80))) <=
     81920), inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    (skoXC1 < 1 |
     81920 <
     skoXC1 *
     (184320 + skoXC1 * (44416 + skoXC1 * (5760 + skoXC1 * -80))) |
     skoXC1 * (320 + skoXC1 * 10) <= -256),
    inference(decision, [],
              [refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_160, plain,
    (skoXC1 < 1 | skoXC1 * (320 + skoXC1 * 10) <= -256),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (skoXC1 < 1 | -256 < skoXC1 * (320 + skoXC1 * 10)),
    inference(decision, [],
              [refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_162, plain, (skoXC1 < 1),
    inference(resolve, [], [refute_0_160, refute_0_161])).

cnf(refute_0_163, plain, (-252 < skoXC1 * (2016 + skoXC1 * 1440)),
    inference(decision, [],
              [refute_0_162, refute_0_150, refute_0_146, refute_0_123,
               refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_164, plain,
    (skoXC1 *
     (90720 +
      skoXC1 *
      (979776 +
       skoXC1 * (2614032 + skoXC1 * (1959552 + skoXC1 * 362880)))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))),
    inference(resolve, [], [refute_0_33, refute_0_163])).

cnf(refute_0_165, plain,
    (skoXC1 *
     (90720 +
      skoXC1 *
      (979776 +
       skoXC1 * (2614032 + skoXC1 * (1959552 + skoXC1 * 362880)))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(resolve, [], [refute_0_164, refute_0_11])).

cnf(refute_0_166, plain,
    (skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= -28 |
     skoXC1 *
     (40320 +
      skoXC1 *
      (344064 +
       skoXC1 *
       (998928 +
        skoXC1 * (688128 + skoXC1 * (161280 + skoXC1 * -8960))))) <= 1120),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (5 + skoXC1 * (40 + skoXC1 * 8) <
     skoXC1 *
     (8640 + skoXC1 * (45792 + skoXC1 * (66528 + skoXC1 * 15120))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 *
     (8640 + skoXC1 * (45792 + skoXC1 * (66528 + skoXC1 * 15120))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_168, plain,
    ((5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (24 + skoXC1 * 60) <= 5 + skoXC1 * (40 + skoXC1 * 8) |
     24 + skoXC1 * 60 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_169, plain,
    (skoXC1 < 0 | 3 < 1 |
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60) <
     nthrt(3, skoXC1) |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_170, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_68, refute_0_169])).

cnf(refute_0_171, plain,
    (skoXC1 < 0 |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(arithmetic, [], [refute_0_170])).

cnf(refute_0_172, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) / (24 + skoXC1 * 60)),
    inference(resolve, [], [refute_0_22, refute_0_171])).

cnf(refute_0_173, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (24 + skoXC1 * 60) <= 5 + skoXC1 * (40 + skoXC1 * 8) |
     24 + skoXC1 * 60 <= 0),
    inference(resolve, [], [refute_0_172, refute_0_168])).

cnf(refute_0_174, plain,
    (skoXC1 *
     (8640 + skoXC1 * (45792 + skoXC1 * (66528 + skoXC1 * 15120))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     5 + skoXC1 * (40 + skoXC1 * 8) | skoXC1 <= -2/5),
    inference(arithmetic, [], [refute_0_173])).

cnf(refute_0_175, plain,
    (skoXC1 *
     (8640 + skoXC1 * (45792 + skoXC1 * (66528 + skoXC1 * 15120))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     5 + skoXC1 * (40 + skoXC1 * 8)),
    inference(resolve, [], [refute_0_174, refute_0_96])).

cnf(refute_0_176, plain,
    (skoXC1 *
     (8640 + skoXC1 * (45792 + skoXC1 * (66528 + skoXC1 * 15120))) <=
     (5 + skoXC1 * (40 + skoXC1 * 8)) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(resolve, [], [refute_0_175, refute_0_167])).

cnf(refute_0_177, plain,
    (skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= -28 |
     skoXC1 *
     (4160 +
      skoXC1 *
      (14488 + skoXC1 * (27152 + skoXC1 * (5200 + skoXC1 * -640)))) <=
     140), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) <
     skoXC1 *
     (7560 + skoXC1 * (66528 + skoXC1 * (91584 + skoXC1 * 34560))) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 *
     (7560 + skoXC1 * (66528 + skoXC1 * (91584 + skoXC1 * 34560))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) *
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_179, plain,
    ((40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) |
     252 + skoXC1 * (2016 + skoXC1 * 1440) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_180, plain,
    (skoXC1 < 0 | 3 < 1 |
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) < nthrt(3, skoXC1) |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(subst, [], [refute_0_55])).

cnf(refute_0_181, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 * (3 * (-2 + 3 * 4) + skoXC1 * (3 * (2 + 3 * 2))) /
     (1 + 3 +
      skoXC1 * (-2 + 3 * (2 + 3 * 4) + skoXC1 * (1 + 3 * (-3 + 3 * 2)))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(resolve, [], [refute_0_23, refute_0_180])).

cnf(refute_0_182, plain,
    (skoXC1 < 0 |
     skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(arithmetic, [], [refute_0_181])).

cnf(refute_0_183, plain,
    (skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) /
     (252 + skoXC1 * (2016 + skoXC1 * 1440))),
    inference(resolve, [], [refute_0_22, refute_0_182])).

cnf(refute_0_184, plain,
    (skoXC1 * (30 + skoXC1 * 24) / (4 + skoXC1 * (40 + skoXC1 * 10)) *
     (252 + skoXC1 * (2016 + skoXC1 * 1440)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112)) |
     252 + skoXC1 * (2016 + skoXC1 * 1440) <= 0),
    inference(resolve, [], [refute_0_183, refute_0_179])).

cnf(refute_0_185, plain,
    (skoXC1 * (2016 + skoXC1 * 1440) <= -252 |
     skoXC1 *
     (7560 + skoXC1 * (66528 + skoXC1 * (91584 + skoXC1 * 34560))) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (skoXC1 *
     (7560 + skoXC1 * (66528 + skoXC1 * (91584 + skoXC1 * 34560))) /
     (4 + skoXC1 * (40 + skoXC1 * 10)) <=
     40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))),
    inference(resolve, [], [refute_0_185, refute_0_163])).

cnf(refute_0_187, plain,
    (skoXC1 *
     (7560 + skoXC1 * (66528 + skoXC1 * (91584 + skoXC1 * 34560))) <=
     (40 + skoXC1 * (840 + skoXC1 * (1344 + skoXC1 * 112))) *
     (4 + skoXC1 * (40 + skoXC1 * 10)) |
     4 + skoXC1 * (40 + skoXC1 * 10) <= 0),
    inference(resolve, [], [refute_0_186, refute_0_178])).

cnf(refute_0_188, plain,
    (skoXC1 * (40 + skoXC1 * 10) <= -4 |
     skoXC1 *
     (2600 +
      skoXC1 *
      (27152 + skoXC1 * (28976 + skoXC1 * (16640 + skoXC1 * -1120)))) <=
     160), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (-4 < skoXC1 * (40 + skoXC1 * 10) |
     skoXC1 *
     (2600 +
      skoXC1 *
      (27152 + skoXC1 * (28976 + skoXC1 * (16640 + skoXC1 * -1120)))) <=
     160),
    inference(decision, [],
              [refute_0_162, refute_0_150, refute_0_146, refute_0_123,
               refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_190, plain,
    (skoXC1 *
     (2600 +
      skoXC1 *
      (27152 + skoXC1 * (28976 + skoXC1 * (16640 + skoXC1 * -1120)))) <=
     160), inference(resolve, [], [refute_0_188, refute_0_189])).

cnf(refute_0_191, plain,
    (-28 < skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) |
     skoXC1 *
     (4160 +
      skoXC1 *
      (14488 + skoXC1 * (27152 + skoXC1 * (5200 + skoXC1 * -640)))) <=
     140),
    inference(decision, [],
              [refute_0_190, refute_0_162, refute_0_150, refute_0_146,
               refute_0_123, refute_0_121, refute_0_119, refute_0_95,
               refute_0_20, refute_0_21])).

cnf(refute_0_192, plain,
    (skoXC1 *
     (4160 +
      skoXC1 *
      (14488 + skoXC1 * (27152 + skoXC1 * (5200 + skoXC1 * -640)))) <=
     140), inference(resolve, [], [refute_0_177, refute_0_191])).

cnf(refute_0_193, plain,
    (-28 < skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) |
     skoXC1 *
     (40320 +
      skoXC1 *
      (344064 +
       skoXC1 *
       (998928 +
        skoXC1 * (688128 + skoXC1 * (161280 + skoXC1 * -8960))))) <= 1120),
    inference(decision, [],
              [refute_0_192, refute_0_190, refute_0_162, refute_0_150,
               refute_0_146, refute_0_123, refute_0_121, refute_0_119,
               refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_194, plain,
    (skoXC1 *
     (40320 +
      skoXC1 *
      (344064 +
       skoXC1 *
       (998928 +
        skoXC1 * (688128 + skoXC1 * (161280 + skoXC1 * -8960))))) <= 1120),
    inference(resolve, [], [refute_0_166, refute_0_193])).

cnf(refute_0_195, plain, (0 <= skoXC1 | 1/2 <= skoXC1),
    inference(decision, [],
              [refute_0_194, refute_0_192, refute_0_190, refute_0_162,
               refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_196, plain,
    (1/2 <= skoXC1 |
     nthrt(3, skoXC1) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_195, refute_0_10])).

cnf(refute_0_197, plain, (skoXC1 < 1/2),
    inference(decision, [],
              [refute_0_194, refute_0_192, refute_0_190, refute_0_162,
               refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_198, plain,
    (nthrt(3, skoXC1) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_196, refute_0_197])).

cnf(refute_0_199, plain,
    (X_000134 < nthrt(X_000131, X_000133) |
     ~ lgen(1, X_000134, nthrt(X_000131, X_000133))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_200, plain,
    (X_000133 < 0 | X_000131 < 1 |
     ~ lgen(1, X_000134,
            (X_000131 * (X_000131 + 1) * 2 ^ (X_000131 + X_000131) *
             X_000133 * X_000133 +
             X_000131 * (X_000131 + X_000131 - 1) * 2 ^ X_000131 *
             X_000133) /
            (2 ^ (X_000131 + X_000131) * (X_000131 + X_000131 - 1) *
             (X_000131 - 1) * X_000133 * X_000133 +
             2 ^ (X_000131 + 1) * (X_000131 + 1) *
             (X_000131 + X_000131 - 1) * X_000133 + (X_000131 + 1))) |
     1 <= X_000133 | lgen(1, X_000134, nthrt(X_000131, X_000133))),
    inference(subst, [], [nthrt_lower_bound_2_small])).

cnf(refute_0_201, plain,
    (X_000133 < 0 | X_000131 < 1 | X_000134 < nthrt(X_000131, X_000133) |
     ~ lgen(1, X_000134,
            (X_000131 * (X_000131 + 1) * 2 ^ (X_000131 + X_000131) *
             X_000133 * X_000133 +
             X_000131 * (X_000131 + X_000131 - 1) * 2 ^ X_000131 *
             X_000133) /
            (2 ^ (X_000131 + X_000131) * (X_000131 + X_000131 - 1) *
             (X_000131 - 1) * X_000133 * X_000133 +
             2 ^ (X_000131 + 1) * (X_000131 + 1) *
             (X_000131 + X_000131 - 1) * X_000133 + (X_000131 + 1))) |
     1 <= X_000133), inference(resolve, [], [refute_0_200, refute_0_199])).

cnf(refute_0_202, plain,
    (X_000133 < 0 | X_000131 < 1 | X_000134 < nthrt(X_000131, X_000133) |
     1 <= X_000133 |
     (2 ^ X_000131 * (X_000133 * (X_000131 * (-1 + X_000131 * 2))) +
      2 ^ (X_000131 + X_000131) *
      (X_000133 * (X_000133 * (X_000131 * (1 + X_000131))))) /
     (1 + X_000131 +
      2 ^ (X_000131 + X_000131) *
      (X_000133 * (X_000133 * (1 + X_000131 * (-3 + X_000131 * 2)))) +
      2 ^ (X_000131 + 1) *
      (X_000133 * (-1 + X_000131 * (1 + X_000131 * 2)))) <= X_000134),
    inference(arithmetic, [], [refute_0_201])).

cnf(refute_0_203, plain,
    (skoXC1 < 0 | 3 < 1 |
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960) <
     nthrt(3, skoXC1) | 1 <= skoXC1 |
     (2 ^ 3 * (skoXC1 * (3 * (-1 + 3 * 2))) +
      2 ^ (3 + 3) * (skoXC1 * (skoXC1 * (3 * (1 + 3))))) /
     (1 + 3 + 2 ^ (3 + 3) * (skoXC1 * (skoXC1 * (1 + 3 * (-3 + 3 * 2)))) +
      2 ^ (3 + 1) * (skoXC1 * (-1 + 3 * (1 + 3 * 2)))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(subst, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (skoXC1 < 0 | 3 < 1 | 1 <= skoXC1 |
     (2 ^ 3 * (skoXC1 * (3 * (-1 + 3 * 2))) +
      2 ^ (3 + 3) * (skoXC1 * (skoXC1 * (3 * (1 + 3))))) /
     (1 + 3 + 2 ^ (3 + 3) * (skoXC1 * (skoXC1 * (1 + 3 * (-3 + 3 * 2)))) +
      2 ^ (3 + 1) * (skoXC1 * (-1 + 3 * (1 + 3 * 2)))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_198, refute_0_203])).

cnf(refute_0_205, plain,
    (skoXC1 < 0 | 1 <= skoXC1 |
     skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (5 + skoXC1 * (320 + skoXC1 * 512) <
     skoXC1 *
     (17280 + skoXC1 * (393984 + skoXC1 * (979776 + skoXC1 * 241920))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 *
     (17280 + skoXC1 * (393984 + skoXC1 * (979776 + skoXC1 * 241920))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_207, plain,
    ((5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960) <
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (48 + skoXC1 * 960) <= 5 + skoXC1 * (320 + skoXC1 * 512) |
     48 + skoXC1 * 960 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_208, plain,
    (skoXC1 < 0 | 3 < 1 |
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960) <
     nthrt(3, skoXC1) |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_209, plain,
    (skoXC1 < 0 | 3 < 1 |
     skoXC1 *
     (3 * (3 + 3 * (-15 + 3 * 18)) +
      skoXC1 *
      (3 * (-6 + 3 * (6 + 3 * 36)) +
       skoXC1 * (3 * (3 + 3 * (9 + 3 * 6))))) /
     (1 + 3 * (3 + 3 * 2) +
      skoXC1 *
      (-3 + 3 * (3 * (21 + 3 * 18)) +
       skoXC1 *
       (3 + 3 * (-9 + 3 * (-12 + 3 * 36)) +
        skoXC1 * (-1 + 3 * (6 + 3 * (-11 + 3 * 6)))))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_198, refute_0_208])).

cnf(refute_0_210, plain,
    (skoXC1 < 0 |
     skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_22, refute_0_210])).

cnf(refute_0_212, plain,
    (skoXC1 * (360 + skoXC1 * (1008 + skoXC1 * 252)) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) *
     (48 + skoXC1 * 960) <= 5 + skoXC1 * (320 + skoXC1 * 512) |
     48 + skoXC1 * 960 <= 0),
    inference(resolve, [], [refute_0_211, refute_0_207])).

cnf(refute_0_213, plain,
    (skoXC1 *
     (17280 + skoXC1 * (393984 + skoXC1 * (979776 + skoXC1 * 241920))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     5 + skoXC1 * (320 + skoXC1 * 512) | skoXC1 <= -1/20),
    inference(arithmetic, [], [refute_0_212])).

cnf(refute_0_214, plain, (-1/20 < skoXC1),
    inference(decision, [],
              [refute_0_194, refute_0_192, refute_0_190, refute_0_162,
               refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_215, plain,
    (skoXC1 *
     (17280 + skoXC1 * (393984 + skoXC1 * (979776 + skoXC1 * 241920))) /
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) <=
     5 + skoXC1 * (320 + skoXC1 * 512)),
    inference(resolve, [], [refute_0_213, refute_0_214])).

cnf(refute_0_216, plain,
    (skoXC1 *
     (17280 + skoXC1 * (393984 + skoXC1 * (979776 + skoXC1 * 241920))) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) *
     (28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80))) |
     28 + skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= 0),
    inference(resolve, [], [refute_0_215, refute_0_206])).

cnf(refute_0_217, plain,
    (skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) <= -28 |
     skoXC1 *
     (4960 +
      skoXC1 *
      (160408 +
       skoXC1 * (366512 + skoXC1 * (-213760 + skoXC1 * -40960)))) <= 140),
    inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (-28 < skoXC1 * (672 + skoXC1 * (840 + skoXC1 * 80)) |
     skoXC1 *
     (4960 +
      skoXC1 *
      (160408 +
       skoXC1 * (366512 + skoXC1 * (-213760 + skoXC1 * -40960)))) <= 140),
    inference(decision, [],
              [refute_0_194, refute_0_192, refute_0_190, refute_0_162,
               refute_0_150, refute_0_146, refute_0_123, refute_0_121,
               refute_0_119, refute_0_95, refute_0_20, refute_0_21])).

cnf(refute_0_219, plain,
    (skoXC1 *
     (4960 +
      skoXC1 *
      (160408 +
       skoXC1 * (366512 + skoXC1 * (-213760 + skoXC1 * -40960)))) <= 140),
    inference(resolve, [], [refute_0_217, refute_0_218])).

cnf(refute_0_220, plain, (0 <= skoXC1 | 1 <= skoXC1),
    inference(decision, [],
              [refute_0_219, refute_0_194, refute_0_192, refute_0_190,
               refute_0_162, refute_0_150, refute_0_146, refute_0_123,
               refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_221, plain,
    (1 <= skoXC1 |
     skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_220, refute_0_205])).

cnf(refute_0_222, plain,
    (skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) / (48 + skoXC1 * 960)),
    inference(resolve, [], [refute_0_221, refute_0_162])).

cnf(refute_0_223, plain,
    (skoXC1 * (120 + skoXC1 * 768) / (4 + skoXC1 * (320 + skoXC1 * 640)) *
     (48 + skoXC1 * 960) <= 5 + skoXC1 * (320 + skoXC1 * 512) |
     48 + skoXC1 * 960 <= 0),
    inference(resolve, [], [refute_0_222, refute_0_1])).

cnf(refute_0_224, plain,
    (skoXC1 * (5760 + skoXC1 * (152064 + skoXC1 * 737280)) /
     (4 + skoXC1 * (320 + skoXC1 * 640)) <=
     5 + skoXC1 * (320 + skoXC1 * 512) | skoXC1 <= -1/20),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (skoXC1 * (5760 + skoXC1 * (152064 + skoXC1 * 737280)) /
     (4 + skoXC1 * (320 + skoXC1 * 640)) <=
     5 + skoXC1 * (320 + skoXC1 * 512)),
    inference(resolve, [], [refute_0_224, refute_0_214])).

cnf(refute_0_226, plain,
    (skoXC1 * (5760 + skoXC1 * (152064 + skoXC1 * 737280)) <=
     (5 + skoXC1 * (320 + skoXC1 * 512)) *
     (4 + skoXC1 * (320 + skoXC1 * 640)) |
     4 + skoXC1 * (320 + skoXC1 * 640) <= 0),
    inference(resolve, [], [refute_0_225, refute_0_0])).

cnf(refute_0_227, plain,
    (skoXC1 * (320 + skoXC1 * 640) <= -4 |
     skoXC1 *
     (2880 + skoXC1 * (44416 + skoXC1 * (368640 + skoXC1 * -327680))) <=
     20), inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain,
    (20 <
     skoXC1 *
     (2880 + skoXC1 * (44416 + skoXC1 * (368640 + skoXC1 * -327680))) |
     skoXC1 * (320 + skoXC1 * 640) <= -4),
    inference(decision, [],
              [refute_0_219, refute_0_194, refute_0_192, refute_0_190,
               refute_0_162, refute_0_150, refute_0_146, refute_0_123,
               refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_229, plain, (skoXC1 * (320 + skoXC1 * 640) <= -4),
    inference(resolve, [], [refute_0_227, refute_0_228])).

cnf(refute_0_230, plain, (-4 < skoXC1 * (320 + skoXC1 * 640)),
    inference(decision, [],
              [refute_0_219, refute_0_194, refute_0_192, refute_0_190,
               refute_0_162, refute_0_150, refute_0_146, refute_0_123,
               refute_0_121, refute_0_119, refute_0_95, refute_0_20,
               refute_0_21])).

cnf(refute_0_231, plain, ($false),
    inference(resolve, [], [refute_0_229, refute_0_230])).
% SZS output end CNFRefutation for nthrt_cbrt-problem-4-weak.tptp
