% SZS output start CNFRefutation for polypaver-sqrt43_A.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(sqrt_le, axiom, (Y < 0 | ~ lgen(R, X, Y ^ 2) | lgen(R, sqrt(X), Y))).

fof(sqrt43, conjecture,
    (! [X, R] :
       ((1/2 <= X & X <= 2 & 0 <= R & R <= 3 & -(X * X / 4) + X <= R &
         R <= X * X / 4 + 1 &
         R <=
         (1 + 2 ^ -23) *
         ((1 + 2 ^ -23) * (R + ((1 + 2 ^ -23) * (X / R) + 2 ^ -126)) +
          2 ^ -126) / 2 + 2 ^ -126 &
         (1 - 2 ^ -23) *
         ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
          2 ^ -126) / 2 - 2 ^ -126 <= R) =>
        (1 - 4 * 2 ^ -23) * sqrt(X) <=
        (1 - 2 ^ -23) *
        ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
         2 ^ -126) / 2 - 2 ^ -126))).

fof(subgoal_0, plain,
    (! [X, R] :
       ((1/2 <= X & X <= 2 & 0 <= R & R <= 3 & -(X * X / 4) + X <= R &
         R <= X * X / 4 + 1 &
         R <=
         (1 + 2 ^ -23) *
         ((1 + 2 ^ -23) * (R + ((1 + 2 ^ -23) * (X / R) + 2 ^ -126)) +
          2 ^ -126) / 2 + 2 ^ -126 &
         (1 - 2 ^ -23) *
         ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
          2 ^ -126) / 2 - 2 ^ -126 <= R) =>
        (1 - 4 * 2 ^ -23) * sqrt(X) <=
        (1 - 2 ^ -23) *
        ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
         2 ^ -126) / 2 - 2 ^ -126)), inference(strip, [], [sqrt43])).

fof(negate_0_0, plain,
    (~ ! [X, R] :
         ((1/2 <= X & X <= 2 & 0 <= R & R <= 3 & -(X * X / 4) + X <= R &
           R <= X * X / 4 + 1 &
           R <=
           (1 + 2 ^ -23) *
           ((1 + 2 ^ -23) * (R + ((1 + 2 ^ -23) * (X / R) + 2 ^ -126)) +
            2 ^ -126) / 2 + 2 ^ -126 &
           (1 - 2 ^ -23) *
           ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
            2 ^ -126) / 2 - 2 ^ -126 <= R) =>
          (1 - 4 * 2 ^ -23) * sqrt(X) <=
          (1 - 2 ^ -23) *
          ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
           2 ^ -126) / 2 - 2 ^ -126)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [R, X] :
       ((1 - 2 ^ -23) *
        ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
         2 ^ -126) / 2 - 2 ^ -126 < (1 - 4 * 2 ^ -23) * sqrt(X) & R <= 3 &
        R <= X * X / 4 + 1 &
        R <=
        (1 + 2 ^ -23) *
        ((1 + 2 ^ -23) * (R + ((1 + 2 ^ -23) * (X / R) + 2 ^ -126)) +
         2 ^ -126) / 2 + 2 ^ -126 & X <= 2 & 0 <= R & 1/2 <= X &
        -(X * X / 4) + X <= R &
        (1 - 2 ^ -23) *
        ((1 - 2 ^ -23) * (R + ((1 - 2 ^ -23) * (X / R) - 2 ^ -126)) -
         2 ^ -126) / 2 - 2 ^ -126 <= R)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    ((1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 < (1 - 4 * 2 ^ -23) * sqrt(skoXC1) &
     0 <= skoRC1 & 1/2 <= skoXC1 &
     -(skoXC1 * skoXC1 / 4) + skoXC1 <= skoRC1 &
     (1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 <= skoRC1 & skoRC1 <= 3 &
     skoRC1 <=
     (1 + 2 ^ -23) *
     ((1 + 2 ^ -23) *
      (skoRC1 + ((1 + 2 ^ -23) * (skoXC1 / skoRC1) + 2 ^ -126)) +
      2 ^ -126) / 2 + 2 ^ -126 & skoRC1 <= skoXC1 * skoXC1 / 4 + 1 &
     skoXC1 <= 2), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    ((1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 < (1 - 4 * 2 ^ -23) * sqrt(skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain,
    (skoRC1 <=
     (1 + 2 ^ -23) *
     ((1 + 2 ^ -23) *
      (skoRC1 + ((1 + 2 ^ -23) * (skoXC1 / skoRC1) + 2 ^ -126)) +
      2 ^ -126) / 2 + 2 ^ -126), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoRC1 <= skoXC1 * skoXC1 / 4 + 1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (-(skoXC1 * skoXC1 / 4) + skoXC1 <= skoRC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoXC1 <= 2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_7, plain, (skoRC1 <= 3),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_8, plain, (1/2 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_9, plain, (0 <= skoRC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_10, plain,
    ((1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 <= skoRC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoRC1 *
     (-281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 8388608/8388607) + skoXC1 <
     0 * (skoRC1 * 1180591057767457882112/590295599252498284543) |
     0 <=
     (skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) |
     skoRC1 * 1180591057767457882112/590295599252498284543 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (skoXC1 * -1 *
     (skoRC1 *
      (skoRC1 *
       1393795245680485074080166545583836489580544/
       348448894496866053305320636211157792718849)) <
     skoRC1 *
     (skoRC1 *
      (-79228148347166085253208997889/
       35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
       skoRC1 *
       (281474951544833/
        2993153926006251682276173773965386945099154627493888 +
        skoRC1 * -70368744177664/70368727400449))) +
     skoXC1 *
     (skoRC1 *
      (281474951544833/
       2993154282818047252893760493211356385005797661212672 +
       skoRC1 * -16777216/8388607) + skoXC1 * -1) |
     skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849) <= 0 |
     (skoRC1 *
      (skoRC1 *
       (-79228148347166085253208997889/
        35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
        skoRC1 *
        (281474951544833/
         2993153926006251682276173773965386945099154627493888 +
         skoRC1 * -70368744177664/70368727400449))) +
      skoXC1 *
      (skoRC1 *
       (281474951544833/
        2993154282818047252893760493211356385005797661212672 +
        skoRC1 * -16777216/8388607) + skoXC1 * -1)) /
     (skoRC1 *
      (skoRC1 *
       1393795245680485074080166545583836489580544/
       348448894496866053305320636211157792718849)) <= skoXC1 * -1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_2, plain,
    ((1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 < (1 - 4 * 2 ^ -23) * sqrt(skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (skoXC1 / skoRC1 <
     281474951544833/5986308565636094505787520986422712770011595322425344 +
     skoRC1 * -8388608/8388607 +
     sqrt(skoXC1) * 1180591057767457882112/590295599252498284543),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (skoXC1 <
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * -8388608/8388607 +
      sqrt(skoXC1) * 1180591057767457882112/590295599252498284543) *
     skoRC1 |
     281474951544833/5986308565636094505787520986422712770011595322425344 +
     skoRC1 * -8388608/8388607 +
     sqrt(skoXC1) * 1180591057767457882112/590295599252498284543 <=
     skoXC1 / skoRC1 | skoRC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_5, plain,
    (skoXC1 <
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * -8388608/8388607 +
      sqrt(skoXC1) * 1180591057767457882112/590295599252498284543) *
     skoRC1 | skoRC1 <= 0),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (skoRC1 *
     (-281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 8388608/8388607) + skoXC1 <
     sqrt(skoXC1) *
     (skoRC1 * 1180591057767457882112/590295599252498284543) |
     skoRC1 <= 0), inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoRC1 <=
     (1 + 2 ^ -23) *
     ((1 + 2 ^ -23) *
      (skoRC1 + ((1 + 2 ^ -23) * (skoXC1 / skoRC1) + 2 ^ -126)) +
      2 ^ -126) / 2 + 2 ^ -126),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_8, plain,
    (-281475001876481/
     5986312847379172623687446570884272788920951721951232 +
     skoRC1 * 590295669621208907776/590296021464963350529 <=
     skoXC1 / skoRC1), inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (skoXC1 / skoRC1 <
     -281475001876481/
     5986312847379172623687446570884272788920951721951232 +
     skoRC1 * 590295669621208907776/590296021464963350529 |
     (-281475001876481/
      5986312847379172623687446570884272788920951721951232 +
      skoRC1 * 590295669621208907776/590296021464963350529) * skoRC1 <=
     skoXC1 | skoRC1 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_10, plain,
    ((-281475001876481/
      5986312847379172623687446570884272788920951721951232 +
      skoRC1 * 590295669621208907776/590296021464963350529) * skoRC1 <=
     skoXC1 | skoRC1 <= 0),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (skoRC1 *
     (-281475001876481/
      5986312847379172623687446570884272788920951721951232 +
      skoRC1 * 590295669621208907776/590296021464963350529) <= skoXC1 |
     skoRC1 <= 0), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain, (skoRC1 <= skoXC1 * skoXC1 / 4 + 1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_13, plain, (skoXC1 * (skoXC1 * -1/4) <= 1 + skoRC1 * -1),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain, (-(skoXC1 * skoXC1 / 4) + skoXC1 <= skoRC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_15, plain, (skoXC1 * (1 + skoXC1 * -1/4) <= skoRC1),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain, (skoXC1 <= 2),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_17, plain, (skoRC1 <= 3),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_18, plain, (1/2 <= skoXC1),
    inference(canonicalize, [], [normalize_0_8])).

cnf(refute_0_19, plain, (0 <= skoRC1),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_20, plain,
    (0 < skoRC1 |
     skoRC1 *
     (-281475001876481/
      5986312847379172623687446570884272788920951721951232 +
      skoRC1 * 590295669621208907776/590296021464963350529) <= skoXC1),
    inference(decision, [],
              [refute_0_13, refute_0_15, refute_0_16, refute_0_17,
               refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (skoRC1 *
     (-281475001876481/
      5986312847379172623687446570884272788920951721951232 +
      skoRC1 * 590295669621208907776/590296021464963350529) <= skoXC1),
    inference(resolve, [], [refute_0_11, refute_0_20])).

cnf(refute_0_22, plain, (0 < skoRC1),
    inference(decision, [],
              [refute_0_21, refute_0_13, refute_0_15, refute_0_16,
               refute_0_17, refute_0_18, refute_0_19])).

cnf(refute_0_23, plain,
    (skoRC1 *
     (-281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 8388608/8388607) + skoXC1 <
     sqrt(skoXC1) *
     (skoRC1 * 1180591057767457882112/590295599252498284543)),
    inference(resolve, [], [refute_0_6, refute_0_22])).

cnf(refute_0_24, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) <
     sqrt(skoXC1) |
     skoRC1 * 1180591057767457882112/590295599252498284543 <= 0),
    inference(split, [], [refute_0_23])).

cnf(refute_0_25, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) <
     sqrt(skoXC1) | skoRC1 <= 0),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) <
     sqrt(skoXC1)), inference(resolve, [], [refute_0_25, refute_0_22])).

cnf(refute_0_27, plain,
    (~ lgen(0, sqrt(X_000045), X_000046) | sqrt(X_000045) <= X_000046),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (X_000046 < 0 | ~ lgen(0, X_000045, X_000046 ^ 2) |
     lgen(0, sqrt(X_000045), X_000046)), inference(subst, [], [sqrt_le])).

cnf(refute_0_29, plain,
    (X_000046 < 0 | ~ lgen(0, X_000045, X_000046 ^ 2) |
     sqrt(X_000045) <= X_000046),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (X_000046 * X_000046 < X_000045 | X_000046 < 0 |
     sqrt(X_000045) <= X_000046),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     (skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) *
     ((skoRC1 *
       (-281474951544833/
        5986308565636094505787520986422712770011595322425344 +
        skoRC1 * 8388608/8388607) + skoXC1) /
      (skoRC1 * 1180591057767457882112/590295599252498284543)) < skoXC1 |
     sqrt(skoXC1) <=
     (skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543)),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     (skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) *
     ((skoRC1 *
       (-281474951544833/
        5986308565636094505787520986422712770011595322425344 +
        skoRC1 * 8388608/8388607) + skoXC1) /
      (skoRC1 * 1180591057767457882112/590295599252498284543)) < skoXC1),
    inference(resolve, [], [refute_0_31, refute_0_26])).

cnf(refute_0_33, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     skoXC1 * -1 <
     (skoRC1 *
      (skoRC1 *
       (-79228148347166085253208997889/
        35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
        skoRC1 *
        (281474951544833/
         2993153926006251682276173773965386945099154627493888 +
         skoRC1 * -70368744177664/70368727400449))) +
      skoXC1 *
      (skoRC1 *
       (281474951544833/
        2993154282818047252893760493211356385005797661212672 +
        skoRC1 * -16777216/8388607) + skoXC1 * -1)) /
     (skoRC1 *
      (skoRC1 *
       1393795245680485074080166545583836489580544/
       348448894496866053305320636211157792718849))),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     skoXC1 * -1 *
     (skoRC1 *
      (skoRC1 *
       1393795245680485074080166545583836489580544/
       348448894496866053305320636211157792718849)) <
     skoRC1 *
     (skoRC1 *
      (-79228148347166085253208997889/
       35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
       skoRC1 *
       (281474951544833/
        2993153926006251682276173773965386945099154627493888 +
        skoRC1 * -70368744177664/70368727400449))) +
     skoXC1 *
     (skoRC1 *
      (281474951544833/
       2993154282818047252893760493211356385005797661212672 +
       skoRC1 * -16777216/8388607) + skoXC1 * -1) |
     skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849) <= 0),
    inference(resolve, [], [refute_0_1, refute_0_33])).

cnf(refute_0_35, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     skoRC1 *
     (skoRC1 *
      (79228148347166085253208997889/
       35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
       skoRC1 *
       (-281474951544833/
        2993153926006251682276173773965386945099154627493888 +
        skoRC1 * 70368744177664/70368727400449))) <
     skoXC1 *
     (skoRC1 *
      (281474951544833/
       2993154282818047252893760493211356385005797661212672 +
       skoRC1 *
       696897373610052748502048717084776288223232/
       348448894496866053305320636211157792718849) + skoXC1 * -1) |
     skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849) <= 0),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    ((1 - 2 ^ -23) *
     ((1 - 2 ^ -23) *
      (skoRC1 + ((1 - 2 ^ -23) * (skoXC1 / skoRC1) - 2 ^ -126)) -
      2 ^ -126) / 2 - 2 ^ -126 <= skoRC1),
    inference(canonicalize, [], [normalize_0_10])).

cnf(refute_0_37, plain,
    (skoXC1 / skoRC1 <=
     281474951544833/5986308565636094505787520986422712770011595322425344 +
     skoRC1 * 590295951096185618432/590295599252498284543),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (281474951544833/5986308565636094505787520986422712770011595322425344 +
     skoRC1 * 590295951096185618432/590295599252498284543 <
     skoXC1 / skoRC1 | skoRC1 <= 0 |
     skoXC1 <=
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 590295951096185618432/590295599252498284543) * skoRC1),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_39, plain,
    (skoRC1 <= 0 |
     skoXC1 <=
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 590295951096185618432/590295599252498284543) * skoRC1),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (skoRC1 <= 0 |
     skoXC1 <=
     skoRC1 *
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 590295951096185618432/590295599252498284543)),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (0 < skoRC1 |
     skoXC1 <=
     skoRC1 *
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 590295951096185618432/590295599252498284543)),
    inference(decision, [],
              [refute_0_21, refute_0_13, refute_0_15, refute_0_16,
               refute_0_17, refute_0_18, refute_0_19])).

cnf(refute_0_42, plain,
    (skoXC1 <=
     skoRC1 *
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 590295951096185618432/590295599252498284543)),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849) <= 0 |
     skoXC1 *
     (skoRC1 *
      (281474951544833/
       2993154282818047252893760493211356385005797661212672 +
       skoRC1 *
       696897373610052748502048717084776288223232/
       348448894496866053305320636211157792718849) + skoXC1 * -1) <=
     skoRC1 *
     (skoRC1 *
      (79228148347166085253208997889/
       35835890243008075201695174122572138901002100160349158413346980827699105518261423012387907685462453518336 +
       skoRC1 *
       (-281474951544833/
        2993153926006251682276173773965386945099154627493888 +
        skoRC1 * 70368744177664/70368727400449)))),
    inference(decision, [],
              [refute_0_42, refute_0_21, refute_0_13, refute_0_15,
               refute_0_16, refute_0_17, refute_0_18, refute_0_19])).

cnf(refute_0_44, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0 |
     skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849) <= 0),
    inference(resolve, [], [refute_0_43, refute_0_35])).

cnf(refute_0_45, plain,
    (0 <
     skoRC1 *
     (skoRC1 *
      1393795245680485074080166545583836489580544/
      348448894496866053305320636211157792718849)),
    inference(decision, [],
              [refute_0_42, refute_0_21, refute_0_13, refute_0_15,
               refute_0_16, refute_0_17, refute_0_18, refute_0_19])).

cnf(refute_0_46, plain,
    ((skoRC1 *
      (-281474951544833/
       5986308565636094505787520986422712770011595322425344 +
       skoRC1 * 8388608/8388607) + skoXC1) /
     (skoRC1 * 1180591057767457882112/590295599252498284543) < 0),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    (skoRC1 *
     (-281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * 8388608/8388607) + skoXC1 <
     0 * (skoRC1 * 1180591057767457882112/590295599252498284543) |
     skoRC1 * 1180591057767457882112/590295599252498284543 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_46])).

cnf(refute_0_48, plain,
    (skoXC1 <
     skoRC1 *
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * -8388608/8388607) | skoRC1 <= 0),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (skoRC1 *
     (281474951544833/
      5986308565636094505787520986422712770011595322425344 +
      skoRC1 * -8388608/8388607) <= skoXC1 | skoRC1 <= 0),
    inference(decision, [],
              [refute_0_42, refute_0_21, refute_0_13, refute_0_15,
               refute_0_16, refute_0_17, refute_0_18, refute_0_19])).

cnf(refute_0_50, plain, (skoRC1 <= 0),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain, ($false),
    inference(resolve, [], [refute_0_50, refute_0_22])).
% SZS output end CNFRefutation for polypaver-sqrt43_A.tptp
