% SZS output start CNFRefutation for sin-cos-problem-diff2a.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, X - X ^ 3 / 6, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_neg, axiom,
    (0 < X |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, Y, X) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

cnf(sin_upper_bound_const_1, axiom,
    (~ lgen(R, 1, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_const_m1, axiom,
    (~ lgen(R, Y, -1) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_const_1, axiom,
    (~ lgen(R, 1, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_const_m1, axiom,
    (~ lgen(R, Y, -1) | lgen(R, Y, cos(X)))).

fof(sin_cos_problem_diff, conjecture,
    (! [X] :
       (X != 0 =>
        abs(-sin(X) / X - 2 * cos(X) / X ^ 2 + 2 * sin(X) / X ^ 3) <=
        1/3))).

fof(subgoal_0, plain,
    (! [X] :
       (X != 0 =>
        abs(-sin(X) / X - 2 * cos(X) / X ^ 2 + 2 * sin(X) / X ^ 3) <=
        1/3)), inference(strip, [], [sin_cos_problem_diff])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (X != 0 =>
          abs(-sin(X) / X - 2 * cos(X) / X ^ 2 + 2 * sin(X) / X ^ 3) <=
          1/3)), inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (1/3 < abs(-sin(X) / X - 2 * cos(X) / X ^ 2 + 2 * sin(X) / X ^ 3) &
        X != 0)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (1/3 <
     abs(-sin(skoXC1) / skoXC1 - 2 * cos(skoXC1) / skoXC1 ^ 2 +
         2 * sin(skoXC1) / skoXC1 ^ 3) & skoXC1 != 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/3 <
     abs(-sin(skoXC1) / skoXC1 - 2 * cos(skoXC1) / skoXC1 ^ 2 +
         2 * sin(skoXC1) / skoXC1 ^ 3)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 != 0),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (~ lgen(0, X_000045, cos(X_000044)) | X_000045 <= cos(X_000044)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000045, -1) | lgen(0, X_000045, cos(X_000044))),
    inference(subst, [], [cos_lower_bound_const_m1])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000045, -1) | X_000045 <= cos(X_000044)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain, (-1 < X_000045 | X_000045 <= cos(X_000044)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (-1 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    (1/3 <
     abs(-sin(skoXC1) / skoXC1 - 2 * cos(skoXC1) / skoXC1 ^ 2 +
         2 * sin(skoXC1) / skoXC1 ^ 3)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (1/3 <
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)))),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) < 0 |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) =
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_10, plain,
    (1/3 <
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) !=
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) <= 1/3),
    introduced(tautology, [equality])).

cnf(refute_0_11, plain,
    (sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) < 0 |
     1/3 <
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) <= 1/3),
    inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) < 0 |
     1/3 <
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_11, refute_0_8])).

cnf(refute_0_13, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain, (skoXC1 != 0),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_15, plain, (skoXC1 * skoXC1 != 0 | skoXC1 = 0),
    inference(decision, [], [refute_0_14])).

cnf(refute_0_16, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(resolve, [], [refute_0_13, refute_0_15])).

cnf(refute_0_17, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_16, refute_0_14])).

cnf(refute_0_18, plain,
    (~ lgen(0, cos(X_000078), X_000079) | cos(X_000078) <= X_000079),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, 1 - X_000078 ^ 2 / 2 + X_000078 ^ 4 / 24, X_000079) |
     lgen(0, cos(X_000078), X_000079)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_20, plain,
    (~ lgen(0, 1 - X_000078 ^ 2 / 2 + X_000078 ^ 4 / 24, X_000079) |
     cos(X_000078) <= X_000079),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (X_000079 <
     1 + X_000078 * (X_000078 * (-1/2 + X_000078 * (X_000078 * 1/24))) |
     cos(X_000078) <= X_000079), inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12)))),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) * (-2 + skoXC1 * skoXC1) |
     0 <= -2 + skoXC1 * skoXC1 |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_24, plain,
    (0 <=
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) =
     -(sin(skoXC1) * -1 / skoXC1 +
      cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
      sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_25, plain,
    (1/3 <
     -(sin(skoXC1) * -1 / skoXC1 +
      cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
      sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) !=
     -(sin(skoXC1) * -1 / skoXC1 +
      cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
      sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) <= 1/3),
    introduced(tautology, [equality])).

cnf(refute_0_26, plain,
    (1/3 <
     -(sin(skoXC1) * -1 / skoXC1 +
      cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
      sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) |
     0 <=
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) |
     abs(sin(skoXC1) * -1 / skoXC1 +
         cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
         sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) <= 1/3),
    inference(resolve, [], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain,
    (1/3 <
     -(sin(skoXC1) * -1 / skoXC1 +
      cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
      sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))) |
     0 <=
     sin(skoXC1) * -1 / skoXC1 + cos(skoXC1) * 2 / (skoXC1 * skoXC1) * -1 +
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_26, refute_0_8])).

cnf(refute_0_28, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(resolve, [], [refute_0_28, refute_0_15])).

cnf(refute_0_30, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_29, refute_0_14])).

cnf(refute_0_31, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <
     skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_32, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <
     skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_31, refute_0_30])).

cnf(refute_0_33, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     sin(skoXC1) * (-2 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain, (skoXC1 != 0 | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(decision, [], [refute_0_14])).

cnf(refute_0_35, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     sin(skoXC1) * (-2 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_33, refute_0_34])).

cnf(refute_0_36, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | 0 <= -2 + skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (~ lgen(0, X_000166, sin(X_000165)) | X_000166 <= sin(X_000165)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_39, plain,
    (X_000165 < 0 | ~ lgen(0, X_000166, X_000165 - X_000165 ^ 3 / 6) |
     lgen(0, X_000166, sin(X_000165))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_40, plain,
    (X_000165 < 0 | ~ lgen(0, X_000166, X_000165 - X_000165 ^ 3 / 6) |
     X_000166 <= sin(X_000165)),
    inference(resolve, [], [refute_0_39, refute_0_38])).

cnf(refute_0_41, plain,
    (X_000165 * (1 + X_000165 * (X_000165 * -1/6)) < X_000166 |
     X_000165 < 0 | X_000166 <= sin(X_000165)),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= sin(skoXC1)),
    inference(subst, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_42, refute_0_37])).

cnf(refute_0_44, plain,
    (~ lgen(0, cos(X_000099), X_000100) | cos(X_000099) <= X_000100),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_45, plain,
    (~ lgen(0,
            1 - X_000099 ^ 2 / 2 + X_000099 ^ 4 / 24 - X_000099 ^ 6 / 720 +
            X_000099 ^ 8 / 40320 - X_000099 ^ 10 / 3628800 +
            X_000099 ^ 12 / 479001600, X_000100) |
     lgen(0, cos(X_000099), X_000100)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_46, plain,
    (~ lgen(0,
            1 - X_000099 ^ 2 / 2 + X_000099 ^ 4 / 24 - X_000099 ^ 6 / 720 +
            X_000099 ^ 8 / 40320 - X_000099 ^ 10 / 3628800 +
            X_000099 ^ 12 / 479001600, X_000100) |
     cos(X_000099) <= X_000100),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain,
    (X_000100 <
     1 +
     X_000099 *
     (X_000099 *
      (-1/2 +
       X_000099 *
       (X_000099 *
        (1/24 +
         X_000099 *
         (X_000099 *
          (-1/720 +
           X_000099 *
           (X_000099 *
            (1/40320 +
             X_000099 *
             (X_000099 *
              (-1/3628800 + X_000099 * (X_000099 * 1/479001600))))))))))) |
     cos(X_000099) <= X_000100), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12)))),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | 2 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_50, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) | 0 <= skoXC1 * (skoXC1 * skoXC1) |
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_51, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) | 0 <= skoXC1 * (skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_50, refute_0_30])).

cnf(refute_0_52, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     sin(skoXC1) * (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_54, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_55, plain,
    (~ lgen(0, sin(X_000011), X_000012) | sin(X_000011) <= X_000012),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_56, plain,
    (~ lgen(0, 1, X_000012) | lgen(0, sin(X_000011), X_000012)),
    inference(subst, [], [sin_upper_bound_const_1])).

cnf(refute_0_57, plain,
    (~ lgen(0, 1, X_000012) | sin(X_000011) <= X_000012),
    inference(resolve, [], [refute_0_56, refute_0_55])).

cnf(refute_0_58, plain, (X_000012 < 1 | sin(X_000011) <= X_000012),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < 1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1)), inference(subst, [], [refute_0_58])).

cnf(refute_0_60, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < sin(skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | -2 + skoXC1 * skoXC1 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_35])).

cnf(refute_0_61, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < sin(skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < 1 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_59, refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     1 * (-2 + skoXC1 * skoXC1) |
     1 <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | -2 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_64, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     1 * (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2 | -2 + skoXC1 * skoXC1 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_63, refute_0_62])).

cnf(refute_0_65, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     cos(skoXC1) * (skoXC1 * 2) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) | skoXC1 * 2 <= 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_65])).

cnf(refute_0_67, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * (skoXC1 * 1/6) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <= skoXC1 * (skoXC1 * 1/6)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_69, plain,
    (-2 + skoXC1 * skoXC1 < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) < 0 |
     0 < -2 + skoXC1 * skoXC1 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_35])).

cnf(refute_0_70, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) < cos(skoXC1) * (skoXC1 * 2) |
     -2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) / (skoXC1 * 2) < cos(skoXC1) |
     skoXC1 * 2 <= 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_72, refute_0_14])).

cnf(refute_0_74, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 0 < skoXC1 * (skoXC1 * skoXC1) |
     2 < skoXC1 * skoXC1 | skoXC1 <= 0),
    inference(decision, [], [refute_0_14])).

cnf(refute_0_75, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * 1/6) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_68, refute_0_75])).

cnf(refute_0_77, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     -1 <
     skoXC1 *
     (skoXC1 *
      (-2/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     2 < skoXC1 * skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-2/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     -1 | skoXC1 <= 0), inference(decision, [], [refute_0_14])).

cnf(refute_0_79, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_78, refute_0_77])).

cnf(refute_0_80, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(decision, [], [refute_0_79, refute_0_14])).

cnf(refute_0_81, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_67, refute_0_80])).

cnf(refute_0_82, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_54, refute_0_81])).

cnf(refute_0_83, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_53, refute_0_82])).

cnf(refute_0_84, plain,
    (2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-4/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 *
             (1/20160 +
              skoXC1 *
              (skoXC1 *
               (-1/1814400 + skoXC1 * (skoXC1 * 1/239500800)))))))))))) |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_86, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_87, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_86, refute_0_81])).

cnf(refute_0_88, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_85, refute_0_87])).

cnf(refute_0_89, plain,
    (2 <
     skoXC1 *
     (2 + skoXC1 * (1 + skoXC1 * (-4/3 + skoXC1 * (skoXC1 * 1/12)))) |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) < 1 * (skoXC1 * 2) |
     1 <= (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) |
     skoXC1 * 2 <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_91, plain,
    (~ lgen(0, cos(X_000031), X_000032) | cos(X_000031) <= X_000032),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_92, plain,
    (~ lgen(0, 1, X_000032) | lgen(0, cos(X_000031), X_000032)),
    inference(subst, [], [cos_upper_bound_const_1])).

cnf(refute_0_93, plain,
    (~ lgen(0, 1, X_000032) | cos(X_000031) <= X_000032),
    inference(resolve, [], [refute_0_92, refute_0_91])).

cnf(refute_0_94, plain, (X_000032 < 1 | cos(X_000031) <= X_000032),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) < 1 |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [refute_0_94])).

cnf(refute_0_96, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * 2) < 1 |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_95, refute_0_81])).

cnf(refute_0_97, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) < 1 * (skoXC1 * 2) |
     skoXC1 * 2 <= 0 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_90, refute_0_96])).

cnf(refute_0_98, plain,
    (2 < skoXC1 * (2 + skoXC1 * (1 + skoXC1 * -1/3)) |
     skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (skoXC1 *
     (2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-4/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 *
             (1/20160 +
              skoXC1 *
              (skoXC1 *
               (-1/1814400 + skoXC1 * (skoXC1 * 1/239500800)))))))))))) <=
     2 | skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_100, plain,
    (skoXC1 * skoXC1 <= 2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_99, refute_0_84])).

cnf(refute_0_101, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_102, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     sin(skoXC1) * (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_52, refute_0_101])).

cnf(refute_0_103, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < sin(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     2 + skoXC1 * (skoXC1 * -1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_102])).

cnf(refute_0_104, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < sin(skoXC1) |
     skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_103])).

cnf(refute_0_105, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_106, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_107, plain,
    (~ lgen(0, X_000021, sin(X_000020)) | X_000021 <= sin(X_000020)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_108, plain,
    (~ lgen(0, X_000021, -1) | lgen(0, X_000021, sin(X_000020))),
    inference(subst, [], [sin_lower_bound_const_m1])).

cnf(refute_0_109, plain,
    (~ lgen(0, X_000021, -1) | X_000021 <= sin(X_000020)),
    inference(resolve, [], [refute_0_108, refute_0_107])).

cnf(refute_0_110, plain, (-1 < X_000021 | X_000021 <= sin(X_000020)),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) <= sin(skoXC1)),
    inference(subst, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (sin(skoXC1) * -2 * (skoXC1 * skoXC1) <
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) +
     sin(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * skoXC1 <= 0 |
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) +
      sin(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1))))) /
     (skoXC1 * skoXC1) <= sin(skoXC1) * -2),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_113, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) < sin(skoXC1) * 2 |
     0 <= skoXC1 * (skoXC1 * skoXC1) |
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_114, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) < sin(skoXC1) * 2 |
     0 <= skoXC1 * (skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_113, refute_0_30])).

cnf(refute_0_115, plain,
    (sin(skoXC1) * -2 <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) +
      sin(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1))))) /
     (skoXC1 * skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (sin(skoXC1) * -2 <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) +
      sin(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1))))) /
     (skoXC1 * skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_115, refute_0_101])).

cnf(refute_0_117, plain,
    (sin(skoXC1) * -2 * (skoXC1 * skoXC1) <
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * 1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) +
     sin(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 * skoXC1 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_112, refute_0_116])).

cnf(refute_0_118, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     sin(skoXC1) * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 * skoXC1 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (0 < skoXC1 * skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_120, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     sin(skoXC1) * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     0 <= skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_124, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_111, refute_0_124])).

cnf(refute_0_126, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     -1 * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     0 <= skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) <= -1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_127, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     -1 * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     0 <= skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_126, refute_0_125])).

cnf(refute_0_128, plain,
    (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)))) <
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)))) /
     (skoXC1 * (skoXC1 * (skoXC1 * -2))) < cos(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_128])).

cnf(refute_0_130, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_129])).

cnf(refute_0_131, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_130, refute_0_14])).

cnf(refute_0_132, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_133, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_134, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= sin(skoXC1)),
    inference(subst, [], [refute_0_110])).

cnf(refute_0_135, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_134, refute_0_37])).

cnf(refute_0_136, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     -1 * (-2 + skoXC1 * skoXC1) | 0 <= -2 + skoXC1 * skoXC1 |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= -1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_137, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     -1 * (-2 + skoXC1 * skoXC1) | 0 <= -2 + skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_136, refute_0_135])).

cnf(refute_0_138, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     cos(skoXC1) * (skoXC1 * 2) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_137])).

cnf(refute_0_139, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_140, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_141, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     1 * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     1 <=
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_142, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) < 1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))))),
    inference(subst, [], [refute_0_58])).

cnf(refute_0_143, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) < sin(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_120])).

cnf(refute_0_144, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
      cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2)))) /
     (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) < 1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_142, refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) <
     1 * (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1)))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_141, refute_0_144])).

cnf(refute_0_146, plain,
    (skoXC1 * (skoXC1 * (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)))) <
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (skoXC1 * (skoXC1 * (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)))) /
     (skoXC1 * (skoXC1 * (skoXC1 * -2))) < cos(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_146])).

cnf(refute_0_148, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_148, refute_0_14])).

cnf(refute_0_150, plain,
    (skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) +
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * 2))) < 0 |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_120])).

cnf(refute_0_151, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) <
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) |
     0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) |
     0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_153, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) <
     cos(skoXC1) * (skoXC1 * (skoXC1 * (skoXC1 * -2))) |
     0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/3)))) /
     (skoXC1 * (skoXC1 * (skoXC1 * -2))) < cos(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -2)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_155, refute_0_14])).

cnf(refute_0_157, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -2)) |
     0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_158, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * 1/6) < cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 * (skoXC1 * 1/6) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_68, refute_0_158])).

cnf(refute_0_160, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     -1 <
     skoXC1 *
     (skoXC1 *
      (-2/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     skoXC1 *
     (skoXC1 *
      (-2/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     -1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_162, plain,
    (0 < skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) |
     0 < skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_161, refute_0_160])).

cnf(refute_0_163, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -2)) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_164, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_149, refute_0_163])).

cnf(refute_0_165, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_140, refute_0_164])).

cnf(refute_0_166, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_139, refute_0_165])).

cnf(refute_0_167, plain,
    (-2 <
     skoXC1 *
     (-2 + skoXC1 * (-1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_166])).

cnf(refute_0_168, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) |
     -2 <
     skoXC1 *
     (-2 + skoXC1 * (-1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_169, plain,
    (-2 <
     skoXC1 *
     (-2 + skoXC1 * (-1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_167, refute_0_168])).

cnf(refute_0_170, plain,
    (skoXC1 < 0 |
     -2 <
     skoXC1 *
     (-2 + skoXC1 * (-1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_171, plain,
    (-2 <
     skoXC1 *
     (-2 + skoXC1 * (-1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     skoXC1 * (skoXC1 * (2 + skoXC1 * (skoXC1 * -1))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_171, refute_0_162, refute_0_100, refute_0_89,
               refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_173, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     cos(skoXC1) * (skoXC1 * 2) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_138, refute_0_172])).

cnf(refute_0_174, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) | skoXC1 * 2 <= 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_173])).

cnf(refute_0_175, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_174])).

cnf(refute_0_176, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_171, refute_0_162, refute_0_100, refute_0_89,
               refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_177, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_133, refute_0_177])).

cnf(refute_0_179, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_132, refute_0_178])).

cnf(refute_0_180, plain,
    (-2 <
     skoXC1 *
     (2 + skoXC1 * (-1 + skoXC1 * (-4/3 + skoXC1 * (skoXC1 * 1/12)))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -2)) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_180, refute_0_171, refute_0_162, refute_0_100,
               refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_182, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_131, refute_0_181])).

cnf(refute_0_183, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_106, refute_0_182])).

cnf(refute_0_184, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_105, refute_0_183])).

cnf(refute_0_185, plain,
    (2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (4/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 *
             (-1/20160 +
              skoXC1 *
              (skoXC1 *
               (1/1814400 + skoXC1 * (skoXC1 * -1/239500800)))))))))))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_187, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_188, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3))) / (skoXC1 * 2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_187, refute_0_177])).

cnf(refute_0_189, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * 2) | skoXC1 * 2 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_186, refute_0_188])).

cnf(refute_0_190, plain,
    (-2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (-4/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 *
             (1/20160 +
              skoXC1 *
              (skoXC1 *
               (-1/1814400 + skoXC1 * (skoXC1 * 1/239500800)))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_189])).

cnf(refute_0_191, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_192, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_193, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_192, refute_0_182])).

cnf(refute_0_194, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_191, refute_0_193])).

cnf(refute_0_195, plain,
    (2 <
     skoXC1 *
     (-2 + skoXC1 * (1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_194])).

cnf(refute_0_196, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) < 1 * (skoXC1 * -2) |
     1 <= (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) |
     skoXC1 * -2 <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_197, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) < 1 |
     cos(skoXC1) <=
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [refute_0_94])).

cnf(refute_0_198, plain,
    ((2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * -2) < 1 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_197, refute_0_182])).

cnf(refute_0_199, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) < 1 * (skoXC1 * -2) |
     skoXC1 * -2 <= 0 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_196, refute_0_198])).

cnf(refute_0_200, plain,
    (2 < skoXC1 * (-2 + skoXC1 * (1 + skoXC1 * 1/3)) | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_199])).

cnf(refute_0_201, plain,
    (skoXC1 < 0 | 2 < skoXC1 * (-2 + skoXC1 * (1 + skoXC1 * 1/3)) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_180, refute_0_171, refute_0_162, refute_0_100,
               refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_202, plain,
    (2 < skoXC1 * (-2 + skoXC1 * (1 + skoXC1 * 1/3)) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_200, refute_0_201])).

cnf(refute_0_203, plain,
    (skoXC1 < 0 |
     2 <
     skoXC1 *
     (-2 + skoXC1 * (1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_202, refute_0_180, refute_0_171, refute_0_162,
               refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_204, plain,
    (2 <
     skoXC1 *
     (-2 + skoXC1 * (1 + skoXC1 * (4/3 + skoXC1 * (skoXC1 * -1/12)))) |
     skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_195, refute_0_203])).

cnf(refute_0_205, plain,
    (-2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (-4/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 *
             (1/20160 +
              skoXC1 *
              (skoXC1 *
               (-1/1814400 + skoXC1 * (skoXC1 * 1/239500800)))))))))))) |
     0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_206, plain,
    (-2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (-4/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 *
             (1/20160 +
              skoXC1 *
              (skoXC1 *
               (-1/1814400 + skoXC1 * (skoXC1 * 1/239500800)))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_190, refute_0_205])).

cnf(refute_0_207, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     skoXC1 *
     (-2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (4/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 *
             (-1/20160 +
              skoXC1 *
              (skoXC1 *
               (1/1814400 + skoXC1 * (skoXC1 * -1/239500800)))))))))))) <=
     2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_206, refute_0_204, refute_0_202, refute_0_180,
               refute_0_171, refute_0_162, refute_0_100, refute_0_89,
               refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_208, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_207, refute_0_185])).

cnf(refute_0_209, plain,
    (skoXC1 < 0 | skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_206, refute_0_204, refute_0_202, refute_0_180,
               refute_0_171, refute_0_162, refute_0_100, refute_0_89,
               refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_210, plain,
    (skoXC1 * (skoXC1 * (-2 + skoXC1 * skoXC1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_210, refute_0_206, refute_0_204, refute_0_202,
               refute_0_180, refute_0_171, refute_0_162, refute_0_100,
               refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_212, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < sin(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_104, refute_0_211])).

cnf(refute_0_213, plain,
    (~ lgen(0, sin(X_000158), X_000159) | sin(X_000158) <= X_000159),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_214, plain,
    (0 < X_000158 | ~ lgen(0, X_000158 - X_000158 ^ 3 / 6, X_000159) |
     lgen(0, sin(X_000158), X_000159)),
    inference(subst, [], [sin_upper_bound_0_neg])).

cnf(refute_0_215, plain,
    (0 < X_000158 | ~ lgen(0, X_000158 - X_000158 ^ 3 / 6, X_000159) |
     sin(X_000158) <= X_000159),
    inference(resolve, [], [refute_0_214, refute_0_213])).

cnf(refute_0_216, plain,
    (0 < X_000158 |
     X_000159 < X_000158 * (1 + X_000158 * (X_000158 * -1/6)) |
     sin(X_000158) <= X_000159),
    inference(arithmetic, [], [refute_0_215])).

cnf(refute_0_217, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) | 0 < skoXC1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))), inference(subst, [], [refute_0_216])).

cnf(refute_0_218, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) | 0 < skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_217, refute_0_212])).

cnf(refute_0_219, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * -2) | skoXC1 * -2 <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_220, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     cos(skoXC1) <=
     (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2)),
    inference(subst, [], [refute_0_47])).

cnf(refute_0_221, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     1 * (2 + skoXC1 * (skoXC1 * -1)) |
     1 <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | 2 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_222, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < 1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))), inference(subst, [], [refute_0_58])).

cnf(refute_0_223, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < 1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_222, refute_0_212])).

cnf(refute_0_224, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     1 * (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     2 + skoXC1 * (skoXC1 * -1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_221, refute_0_223])).

cnf(refute_0_225, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_225, refute_0_211])).

cnf(refute_0_227, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | skoXC1 * -2 <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_226])).

cnf(refute_0_228, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | 0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_227])).

cnf(refute_0_229, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_210, refute_0_206, refute_0_204, refute_0_202,
               refute_0_180, refute_0_171, refute_0_162, refute_0_100,
               refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_230, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     cos(skoXC1) | 0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_228, refute_0_229])).

cnf(refute_0_231, plain,
    ((-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3))) / (skoXC1 * -2) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_220, refute_0_230])).

cnf(refute_0_232, plain,
    (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600)))))))))))) *
     (skoXC1 * -2) | 0 <= skoXC1 | skoXC1 * -2 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_219, refute_0_231])).

cnf(refute_0_233, plain,
    (-2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (4/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 *
             (-1/20160 +
              skoXC1 *
              (skoXC1 *
               (1/1814400 + skoXC1 * (skoXC1 * -1/239500800)))))))))))) |
     0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_232])).

cnf(refute_0_234, plain,
    (skoXC1 < 0 |
     -2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (4/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 *
             (-1/20160 +
              skoXC1 *
              (skoXC1 *
               (1/1814400 + skoXC1 * (skoXC1 * -1/239500800)))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_210, refute_0_206, refute_0_204, refute_0_202,
               refute_0_180, refute_0_171, refute_0_162, refute_0_100,
               refute_0_89, refute_0_98, refute_0_79, refute_0_14])).

cnf(refute_0_235, plain,
    (-2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (-1 +
       skoXC1 *
       (4/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 *
             (-1/20160 +
              skoXC1 *
              (skoXC1 *
               (1/1814400 + skoXC1 * (skoXC1 * -1/239500800)))))))))))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_233, refute_0_234])).

cnf(refute_0_236, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_235, refute_0_210, refute_0_206, refute_0_204,
               refute_0_202, refute_0_180, refute_0_171, refute_0_162,
               refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_237, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) | 0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_218, refute_0_236])).

cnf(refute_0_238, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (2 + skoXC1 * (skoXC1 * -1)) | 0 < skoXC1 |
     2 + skoXC1 * (skoXC1 * -1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_49, refute_0_237])).

cnf(refute_0_239, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)))) <
     cos(skoXC1) * (skoXC1 * -2) | 0 < skoXC1 |
     skoXC1 * (skoXC1 * -1) <= -2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_238])).

cnf(refute_0_240, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_235, refute_0_210, refute_0_206, refute_0_204,
               refute_0_202, refute_0_180, refute_0_171, refute_0_162,
               refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_241, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)))) <
     cos(skoXC1) * (skoXC1 * -2) | 0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_239, refute_0_240])).

cnf(refute_0_242, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)))) /
     (skoXC1 * -2) < cos(skoXC1) | 0 < skoXC1 | skoXC1 * -2 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) | 0 < skoXC1 | 0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) | 0 < skoXC1 | 0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_243, refute_0_14])).

cnf(refute_0_245, plain,
    (0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_235, refute_0_210, refute_0_206, refute_0_204,
               refute_0_202, refute_0_180, refute_0_171, refute_0_162,
               refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_246, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) | 0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_245, refute_0_244])).

cnf(refute_0_247, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_48, refute_0_246])).

cnf(refute_0_248, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) |
     0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 + skoXC1 * (skoXC1 * 1/479001600))))))))))) <=
     0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_235, refute_0_210, refute_0_206, refute_0_204,
               refute_0_202, refute_0_180, refute_0_171, refute_0_162,
               refute_0_100, refute_0_89, refute_0_98, refute_0_79,
               refute_0_14])).

cnf(refute_0_250, plain,
    (0 <= skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_249, refute_0_248])).

cnf(refute_0_251, plain,
    (skoXC1 < 0 | -2 < skoXC1 * (skoXC1 * -1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_250, refute_0_235, refute_0_210, refute_0_206,
               refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_252, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_43, refute_0_251])).

cnf(refute_0_253, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_250, refute_0_235, refute_0_210, refute_0_206,
               refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_254, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_252, refute_0_253])).

cnf(refute_0_255, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_250, refute_0_254])).

cnf(refute_0_256, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) * (-2 + skoXC1 * skoXC1) |
     0 <= -2 + skoXC1 * skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_23, refute_0_255])).

cnf(refute_0_257, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/6)))) <
     cos(skoXC1) * (skoXC1 * 2) | skoXC1 * (skoXC1 * -1) <= -2 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_250, refute_0_235, refute_0_210, refute_0_206,
               refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_259, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/6)))) <
     cos(skoXC1) * (skoXC1 * 2) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_257, refute_0_258])).

cnf(refute_0_260, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/6)))) /
     (skoXC1 * 2) < cos(skoXC1) | skoXC1 * 2 <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0 |
     skoXC1 = 0), inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 <= 0),
    inference(resolve, [], [refute_0_261, refute_0_14])).

cnf(refute_0_263, plain,
    (0 < skoXC1 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_250, refute_0_235, refute_0_210, refute_0_206,
               refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_264, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     cos(skoXC1) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_262, refute_0_263])).

cnf(refute_0_265, plain,
    (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/12))) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_22, refute_0_264])).

cnf(refute_0_266, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/24))) |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1/24))) <= 0 |
     (cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(decision, [],
              [refute_0_250, refute_0_235, refute_0_210, refute_0_206,
               refute_0_204, refute_0_202, refute_0_180, refute_0_171,
               refute_0_162, refute_0_100, refute_0_89, refute_0_98,
               refute_0_79, refute_0_14])).

cnf(refute_0_268, plain,
    ((cos(skoXC1) * (skoXC1 * 2) + sin(skoXC1) * (skoXC1 * skoXC1)) /
     (skoXC1 * (skoXC1 * skoXC1)) <=
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_267, refute_0_266])).

cnf(refute_0_269, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1)) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_268, refute_0_17])).

cnf(refute_0_270, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) <
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_269])).

cnf(refute_0_271, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     sin(skoXC1) * (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_270])).

cnf(refute_0_272, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     sin(skoXC1) * (2 + skoXC1 * (skoXC1 * -1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_271, refute_0_34])).

cnf(refute_0_273, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(split, [], [refute_0_272])).

cnf(refute_0_274, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_273])).

cnf(refute_0_275, plain,
    (~ lgen(0, X_000197, sin(X_000196)) | X_000197 <= sin(X_000196)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_276, plain,
    (X_000196 < 0 |
     ~ lgen(0, X_000197,
            X_000196 - X_000196 ^ 3 / 6 + X_000196 ^ 5 / 120 -
            X_000196 ^ 7 / 5040 + X_000196 ^ 9 / 362880 -
            X_000196 ^ 11 / 39916800) | lgen(0, X_000197, sin(X_000196))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_277, plain,
    (X_000196 < 0 |
     ~ lgen(0, X_000197,
            X_000196 - X_000196 ^ 3 / 6 + X_000196 ^ 5 / 120 -
            X_000196 ^ 7 / 5040 + X_000196 ^ 9 / 362880 -
            X_000196 ^ 11 / 39916800) | X_000197 <= sin(X_000196)),
    inference(resolve, [], [refute_0_276, refute_0_275])).

cnf(refute_0_278, plain,
    (X_000196 *
     (1 +
      X_000196 *
      (X_000196 *
       (-1/6 +
        X_000196 *
        (X_000196 *
         (1/120 +
          X_000196 *
          (X_000196 *
           (-1/5040 +
            X_000196 *
            (X_000196 *
             (1/362880 + X_000196 * (X_000196 * -1/39916800)))))))))) <
     X_000197 | X_000196 < 0 | X_000197 <= sin(X_000196)),
    inference(arithmetic, [], [refute_0_277])).

cnf(refute_0_279, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <= sin(skoXC1)),
    inference(subst, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_279, refute_0_274])).

cnf(refute_0_281, plain,
    (~ lgen(0, X_000092, cos(X_000091)) | X_000092 <= cos(X_000091)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_282, plain,
    (~ lgen(0, X_000092,
            1 - X_000091 ^ 2 / 2 + X_000091 ^ 4 / 24 - X_000091 ^ 6 / 720 +
            X_000091 ^ 8 / 40320 - X_000091 ^ 10 / 3628800) |
     lgen(0, X_000092, cos(X_000091))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_283, plain,
    (~ lgen(0, X_000092,
            1 - X_000091 ^ 2 / 2 + X_000091 ^ 4 / 24 - X_000091 ^ 6 / 720 +
            X_000091 ^ 8 / 40320 - X_000091 ^ 10 / 3628800) |
     X_000092 <= cos(X_000091)),
    inference(resolve, [], [refute_0_282, refute_0_281])).

cnf(refute_0_284, plain,
    (1 +
     X_000091 *
     (X_000091 *
      (-1/2 +
       X_000091 *
       (X_000091 *
        (1/24 +
         X_000091 *
         (X_000091 *
          (-1/720 +
           X_000091 *
           (X_000091 *
            (1/40320 + X_000091 * (X_000091 * -1/3628800))))))))) <
     X_000092 | X_000092 <= cos(X_000091)),
    inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     skoXC1 * (skoXC1 * -1/6) | skoXC1 * (skoXC1 * -1/6) <= cos(skoXC1)),
    inference(subst, [], [refute_0_284])).

cnf(refute_0_286, plain,
    (2 + skoXC1 * (skoXC1 * -1) < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) < 0 |
     0 < 2 + skoXC1 * (skoXC1 * -1) | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(split, [], [refute_0_272])).

cnf(refute_0_287, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) < cos(skoXC1) * (skoXC1 * -2) |
     -2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_286])).

cnf(refute_0_288, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) | 0 <= skoXC1 * (skoXC1 * skoXC1) |
     sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
      sin(skoXC1) * (skoXC1 * skoXC1)) / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_289, plain,
    (sin(skoXC1) * 2 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) +
     sin(skoXC1) * (skoXC1 * skoXC1) | 0 <= skoXC1 * (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_288, refute_0_269])).

cnf(refute_0_290, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     sin(skoXC1) * (-2 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_289])).

cnf(refute_0_291, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(decision, [], [refute_0_14])).

cnf(refute_0_292, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     sin(skoXC1) * (-2 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_290, refute_0_291])).

cnf(refute_0_293, plain,
    (-2 + skoXC1 * skoXC1 < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     0 | 0 < -2 + skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(split, [], [refute_0_292])).

cnf(refute_0_294, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) < cos(skoXC1) * (skoXC1 * 2) |
     -2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     cos(skoXC1) < skoXC1 * (skoXC1 * (skoXC1 * -1/3)) / (skoXC1 * 2) |
     0 <= skoXC1 * 2 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(split, [], [refute_0_294])).

cnf(refute_0_296, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_295])).

cnf(refute_0_297, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_296, refute_0_14])).

cnf(refute_0_298, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_14])).

cnf(refute_0_299, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_297, refute_0_298])).

cnf(refute_0_300, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     skoXC1 * (skoXC1 * -1/6) | 2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_285, refute_0_299])).

cnf(refute_0_301, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     1 <
     skoXC1 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_300])).

cnf(refute_0_302, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 | 0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) <=
     1), inference(decision, [], [refute_0_14])).

cnf(refute_0_303, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_302, refute_0_301])).

cnf(refute_0_304, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 0 < skoXC1 * (skoXC1 * skoXC1) |
     2 < skoXC1 * skoXC1),
    inference(decision, [], [refute_0_303, refute_0_14])).

cnf(refute_0_305, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) < cos(skoXC1) * (skoXC1 * -2) |
     -2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1),
    inference(resolve, [], [refute_0_287, refute_0_304])).

cnf(refute_0_306, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     cos(skoXC1) < skoXC1 * (skoXC1 * (skoXC1 * 1/3)) / (skoXC1 * -2) |
     0 <= skoXC1 * -2), inference(split, [], [refute_0_305])).

cnf(refute_0_307, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1 | skoXC1 <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_307, refute_0_14])).

cnf(refute_0_309, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 | 0 < skoXC1),
    inference(decision, [], [refute_0_303, refute_0_14])).

cnf(refute_0_310, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | cos(skoXC1) < skoXC1 * (skoXC1 * -1/6) |
     2 < skoXC1 * skoXC1),
    inference(resolve, [], [refute_0_308, refute_0_309])).

cnf(refute_0_311, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     skoXC1 * (skoXC1 * -1/6) | 2 < skoXC1 * skoXC1),
    inference(resolve, [], [refute_0_285, refute_0_310])).

cnf(refute_0_312, plain,
    (-2 < skoXC1 * (skoXC1 * -1) |
     1 <
     skoXC1 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     2 < skoXC1 * skoXC1), inference(arithmetic, [], [refute_0_311])).

cnf(refute_0_313, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/3 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) <=
     1), inference(decision, [], [refute_0_303, refute_0_14])).

cnf(refute_0_314, plain,
    (-2 < skoXC1 * (skoXC1 * -1) | 2 < skoXC1 * skoXC1),
    inference(resolve, [], [refute_0_313, refute_0_312])).

cnf(refute_0_315, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * -2/3) |
     1 + skoXC1 * (skoXC1 * -2/3) <= cos(skoXC1)),
    inference(subst, [], [refute_0_284])).

cnf(refute_0_316, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * -1)) |
     2 + skoXC1 * (skoXC1 * -1) <= 0 |
     skoXC1 <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_317, plain,
    (~ lgen(0, sin(X_000144), X_000145) | sin(X_000144) <= X_000145),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_318, plain,
    (X_000144 < 0 | ~ lgen(0, X_000144, X_000145) |
     lgen(0, sin(X_000144), X_000145)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_319, plain,
    (X_000144 < 0 | ~ lgen(0, X_000144, X_000145) |
     sin(X_000144) <= X_000145),
    inference(resolve, [], [refute_0_318, refute_0_317])).

cnf(refute_0_320, plain,
    (X_000145 < X_000144 | X_000144 < 0 | sin(X_000144) <= X_000145),
    inference(arithmetic, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (skoXC1 < 0 |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < skoXC1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))), inference(subst, [], [refute_0_320])).

cnf(refute_0_322, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < sin(skoXC1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | 2 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(split, [], [refute_0_272])).

cnf(refute_0_323, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < sin(skoXC1) |
     skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_322])).

cnf(refute_0_324, plain,
    (skoXC1 < 0 |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_321, refute_0_323])).

cnf(refute_0_325, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (-2 + skoXC1 * skoXC1) | 0 <= -2 + skoXC1 * skoXC1 |
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= skoXC1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_326, plain,
    (~ lgen(0, X_000152, sin(X_000151)) | X_000152 <= sin(X_000151)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_327, plain,
    (0 < X_000151 | ~ lgen(0, X_000152, X_000151) |
     lgen(0, X_000152, sin(X_000151))),
    inference(subst, [], [sin_lower_bound_0_neg])).

cnf(refute_0_328, plain,
    (0 < X_000151 | ~ lgen(0, X_000152, X_000151) |
     X_000152 <= sin(X_000151)),
    inference(resolve, [], [refute_0_327, refute_0_326])).

cnf(refute_0_329, plain,
    (0 < X_000151 | X_000151 < X_000152 | X_000152 <= sin(X_000151)),
    inference(arithmetic, [], [refute_0_328])).

cnf(refute_0_330, plain,
    (skoXC1 <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | 0 < skoXC1 |
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <= sin(skoXC1)),
    inference(subst, [], [refute_0_329])).

cnf(refute_0_331, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | 0 <= -2 + skoXC1 * skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(split, [], [refute_0_292])).

cnf(refute_0_332, plain,
    (sin(skoXC1) <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_331])).

cnf(refute_0_333, plain,
    (skoXC1 <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | 0 < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_330, refute_0_332])).

cnf(refute_0_334, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_284])).

cnf(refute_0_335, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) * (-2 + skoXC1 * skoXC1) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | -2 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_336, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) |
     0 < skoXC1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1)), inference(subst, [], [refute_0_216])).

cnf(refute_0_337, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < sin(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | -2 + skoXC1 * skoXC1 <= 0),
    inference(split, [], [refute_0_292])).

cnf(refute_0_338, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < sin(skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_337])).

cnf(refute_0_339, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) |
     0 < skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_336, refute_0_338])).

cnf(refute_0_340, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 < skoXC1 |
     skoXC1 * skoXC1 <= 2),
    inference(decision, [], [refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_341, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) |
     0 < skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_339, refute_0_340])).

cnf(refute_0_342, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) * (-2 + skoXC1 * skoXC1) |
     0 < skoXC1 | skoXC1 * skoXC1 <= 2 | -2 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_335, refute_0_341])).

cnf(refute_0_343, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * (-5/3 + skoXC1 * (skoXC1 * 1/6)))) <
     cos(skoXC1) * (skoXC1 * 2) | 0 < skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_342])).

cnf(refute_0_344, plain,
    (cos(skoXC1) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * (-5/3 + skoXC1 * (skoXC1 * 1/6)))) /
     (skoXC1 * 2) | 0 < skoXC1 | 0 <= skoXC1 * 2 | skoXC1 * skoXC1 <= 2),
    inference(split, [], [refute_0_343])).

cnf(refute_0_345, plain,
    (cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 < skoXC1 | 0 <= skoXC1 | skoXC1 * skoXC1 <= 2 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_344])).

cnf(refute_0_346, plain,
    (cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 < skoXC1 | 0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_345, refute_0_14])).

cnf(refute_0_347, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (skoXC1 * -2) | 0 <= skoXC1 * -2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_348, plain,
    (~ lgen(0, X_000058, cos(X_000057)) | X_000058 <= cos(X_000057)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_349, plain,
    (~ lgen(0, X_000058, 1 - X_000057 ^ 2 / 2) |
     lgen(0, X_000058, cos(X_000057))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_350, plain,
    (~ lgen(0, X_000058, 1 - X_000057 ^ 2 / 2) |
     X_000058 <= cos(X_000057)),
    inference(resolve, [], [refute_0_349, refute_0_348])).

cnf(refute_0_351, plain,
    (1 + X_000057 * (X_000057 * -1/2) < X_000058 |
     X_000058 <= cos(X_000057)),
    inference(arithmetic, [], [refute_0_350])).

cnf(refute_0_352, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_351])).

cnf(refute_0_353, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     -1 * (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <= -1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_354, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <= sin(skoXC1)),
    inference(subst, [], [refute_0_110])).

cnf(refute_0_355, plain,
    (-1 <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_354, refute_0_274])).

cnf(refute_0_356, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     -1 * (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_353, refute_0_355])).

cnf(refute_0_357, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_356])).

cnf(refute_0_358, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     0 <= skoXC1 * -2 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(split, [], [refute_0_357])).

cnf(refute_0_359, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (skoXC1 * 2) | 0 <= skoXC1 * 2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_361, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_351])).

cnf(refute_0_362, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     1 * (-2 + skoXC1 * skoXC1) |
     1 <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | -2 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_363, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < 1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1)), inference(subst, [], [refute_0_58])).

cnf(refute_0_364, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) < 1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_363, refute_0_338])).

cnf(refute_0_365, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     1 * (-2 + skoXC1 * skoXC1) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2 | -2 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_362, refute_0_364])).

cnf(refute_0_366, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     cos(skoXC1) * (skoXC1 * 2) | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_365])).

cnf(refute_0_367, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 * 2 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(split, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_367])).

cnf(refute_0_369, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1 |
     skoXC1 * skoXC1 <= 2),
    inference(decision, [], [refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_370, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_368, refute_0_369])).

cnf(refute_0_371, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_361, refute_0_370])).

cnf(refute_0_372, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (skoXC1 * 2) | 0 <= skoXC1 * 2 |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_360, refute_0_371])).

cnf(refute_0_373, plain,
    (2 < skoXC1 * (2 + skoXC1 * (1 + skoXC1 * -2/3)) | 0 <= skoXC1 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_372])).

cnf(refute_0_374, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) < -1 * (skoXC1 * 2) |
     0 <= skoXC1 * 2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <= -1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_375, plain,
    (-1 < (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_3])).

cnf(refute_0_376, plain,
    (-1 < (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_375, refute_0_370])).

cnf(refute_0_377, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) < -1 * (skoXC1 * 2) |
     0 <= skoXC1 * 2 | 0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_374, refute_0_376])).

cnf(refute_0_378, plain,
    (2 < skoXC1 * (-2 + skoXC1 * (1 + skoXC1 * 1/3)) | 0 <= skoXC1 |
     skoXC1 * skoXC1 <= 2), inference(arithmetic, [], [refute_0_377])).

cnf(refute_0_379, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_373, refute_0_378, refute_0_314, refute_0_303,
               refute_0_14])).

cnf(refute_0_380, plain,
    (cos(skoXC1) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_359, refute_0_379])).

cnf(refute_0_381, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_352, refute_0_380])).

cnf(refute_0_382, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (skoXC1 * -2) | 0 <= skoXC1 * -2 |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_347, refute_0_381])).

cnf(refute_0_383, plain,
    (2 < skoXC1 * (-2 + skoXC1 * (1 + skoXC1 * 2/3)) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_382])).

cnf(refute_0_384, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) < -1 * (skoXC1 * -2) |
     0 <= skoXC1 * -2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <= -1),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_385, plain,
    (-1 < (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_3])).

cnf(refute_0_386, plain,
    (-1 < (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_385, refute_0_380])).

cnf(refute_0_387, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) < -1 * (skoXC1 * -2) |
     0 <= skoXC1 * -2 | skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_384, refute_0_386])).

cnf(refute_0_388, plain,
    (2 < skoXC1 * (2 + skoXC1 * (1 + skoXC1 * -1/3)) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_387])).

cnf(refute_0_389, plain,
    (0 <= skoXC1 | skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_390, plain,
    (cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_389, refute_0_346])).

cnf(refute_0_391, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_334, refute_0_390])).

cnf(refute_0_392, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_351])).

cnf(refute_0_394, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_393, refute_0_390])).

cnf(refute_0_395, plain,
    (0 < skoXC1 * (skoXC1 * (-1/3 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (-1 < 1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_3])).

cnf(refute_0_397, plain,
    (-1 < 1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_396, refute_0_390])).

cnf(refute_0_398, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     2 < skoXC1 * skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_400, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_398, refute_0_399])).

cnf(refute_0_401, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 < skoXC1 |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(decision, [],
              [refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_402, plain,
    (skoXC1 <
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | 0 < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_333, refute_0_401])).

cnf(refute_0_403, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 * (-2 + skoXC1 * skoXC1) | 0 < skoXC1 |
     0 <= -2 + skoXC1 * skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_325, refute_0_402])).

cnf(refute_0_404, plain,
    (skoXC1 * (2 + skoXC1 * (skoXC1 * -4/3)) < cos(skoXC1) * (skoXC1 * 2) |
     0 < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(arithmetic, [], [refute_0_403])).

cnf(refute_0_405, plain,
    (cos(skoXC1) < skoXC1 * (2 + skoXC1 * (skoXC1 * -4/3)) / (skoXC1 * 2) |
     0 < skoXC1 | 0 <= skoXC1 * 2 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(split, [], [refute_0_404])).

cnf(refute_0_406, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) | 0 < skoXC1 |
     0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_405])).

cnf(refute_0_407, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) | 0 < skoXC1 |
     0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_406, refute_0_14])).

cnf(refute_0_408, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_409, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_408, refute_0_407])).

cnf(refute_0_410, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * -2/3) | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_315, refute_0_409])).

cnf(refute_0_411, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/6 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(arithmetic, [], [refute_0_410])).

cnf(refute_0_412, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 *
     (skoXC1 *
      (-1/6 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) <=
     0),
    inference(decision, [],
              [refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_413, plain, (0 <= skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_412, refute_0_411])).

cnf(refute_0_414, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(decision, [],
              [refute_0_413, refute_0_392, refute_0_395, refute_0_400,
               refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_415, plain,
    (skoXC1 < 0 |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_324, refute_0_414])).

cnf(refute_0_416, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_413, refute_0_415])).

cnf(refute_0_417, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * -1) <= -2 |
     2 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_316, refute_0_416])).

cnf(refute_0_418, plain,
    (skoXC1 * (-2 + skoXC1 * (skoXC1 * 4/3)) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * (skoXC1 * -1) <= -2),
    inference(arithmetic, [], [refute_0_417])).

cnf(refute_0_419, plain,
    (cos(skoXC1) <
     skoXC1 * (-2 + skoXC1 * (skoXC1 * 4/3)) / (skoXC1 * -2) |
     0 <= skoXC1 * -2 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(split, [], [refute_0_418])).

cnf(refute_0_420, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) |
     skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_419])).

cnf(refute_0_421, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) |
     skoXC1 * (skoXC1 * -1) <= -2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_420, refute_0_14])).

cnf(refute_0_422, plain, (0 < skoXC1 | skoXC1 * (skoXC1 * -1) <= -2),
    inference(decision, [],
              [refute_0_413, refute_0_392, refute_0_395, refute_0_400,
               refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_423, plain,
    (cos(skoXC1) < 1 + skoXC1 * (skoXC1 * -2/3) |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_421, refute_0_422])).

cnf(refute_0_424, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * -2/3) | skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_315, refute_0_423])).

cnf(refute_0_425, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/6 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     skoXC1 * (skoXC1 * -1) <= -2),
    inference(arithmetic, [], [refute_0_424])).

cnf(refute_0_426, plain,
    (skoXC1 * (skoXC1 * -1) <= -2 |
     skoXC1 *
     (skoXC1 *
      (-1/6 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) <=
     0),
    inference(decision, [],
              [refute_0_413, refute_0_392, refute_0_395, refute_0_400,
               refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_427, plain, (skoXC1 * (skoXC1 * -1) <= -2),
    inference(resolve, [], [refute_0_426, refute_0_425])).

cnf(refute_0_428, plain, (2 < skoXC1 * skoXC1),
    inference(resolve, [], [refute_0_427, refute_0_314])).

cnf(refute_0_429, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_280, refute_0_428])).

cnf(refute_0_430, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_284])).

cnf(refute_0_431, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-2 + skoXC1 * skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) | -2 + skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_432, plain,
    (~ lgen(0, sin(X_000189), X_000190) | sin(X_000189) <= X_000190),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_433, plain,
    (0 < X_000189 |
     ~ lgen(0,
            X_000189 - X_000189 ^ 3 / 6 + X_000189 ^ 5 / 120 -
            X_000189 ^ 7 / 5040 + X_000189 ^ 9 / 362880 -
            X_000189 ^ 11 / 39916800, X_000190) |
     lgen(0, sin(X_000189), X_000190)),
    inference(subst, [], [sin_upper_bound_2_neg])).

cnf(refute_0_434, plain,
    (0 < X_000189 |
     ~ lgen(0,
            X_000189 - X_000189 ^ 3 / 6 + X_000189 ^ 5 / 120 -
            X_000189 ^ 7 / 5040 + X_000189 ^ 9 / 362880 -
            X_000189 ^ 11 / 39916800, X_000190) |
     sin(X_000189) <= X_000190),
    inference(resolve, [], [refute_0_433, refute_0_432])).

cnf(refute_0_435, plain,
    (0 < X_000189 |
     X_000190 <
     X_000189 *
     (1 +
      X_000189 *
      (X_000189 *
       (-1/6 +
        X_000189 *
        (X_000189 *
         (1/120 +
          X_000189 *
          (X_000189 *
           (-1/5040 +
            X_000189 *
            (X_000189 *
             (1/362880 + X_000189 * (X_000189 * -1/39916800)))))))))) |
     sin(X_000189) <= X_000190),
    inference(arithmetic, [], [refute_0_434])).

cnf(refute_0_436, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 |
     sin(skoXC1) <=
     (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1)), inference(subst, [], [refute_0_435])).

cnf(refute_0_437, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_436, refute_0_338])).

cnf(refute_0_438, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_437, refute_0_428])).

cnf(refute_0_439, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))) *
     (skoXC1 * 2) | 0 <= skoXC1 * 2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_440, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_284])).

cnf(refute_0_441, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_440, refute_0_370])).

cnf(refute_0_442, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3))) / (skoXC1 * 2) |
     0 <= skoXC1), inference(resolve, [], [refute_0_441, refute_0_428])).

cnf(refute_0_443, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * -1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))) *
     (skoXC1 * 2) | 0 <= skoXC1 * 2 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_439, refute_0_442])).

cnf(refute_0_444, plain,
    (2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 * (1/20160 + skoXC1 * (skoXC1 * -1/1814400)))))))))) |
     0 <= skoXC1), inference(arithmetic, [], [refute_0_443])).

cnf(refute_0_445, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))) *
     (skoXC1 * -2) | 0 <= skoXC1 * -2 |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_446, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) <=
     cos(skoXC1)), inference(subst, [], [refute_0_284])).

cnf(refute_0_447, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 * skoXC1 <= 2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_446, refute_0_380])).

cnf(refute_0_448, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3))) / (skoXC1 * -2) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_447, refute_0_428])).

cnf(refute_0_449, plain,
    (2 + skoXC1 * (skoXC1 * (-1 + skoXC1 * 1/3)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))) *
     (skoXC1 * -2) | 0 <= skoXC1 * -2 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_445, refute_0_448])).

cnf(refute_0_450, plain,
    (2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (2/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 * (-1/20160 + skoXC1 * (skoXC1 * 1/1814400)))))))))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_449])).

cnf(refute_0_451, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <=
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_452, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) |
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) <= sin(skoXC1)),
    inference(subst, [], [refute_0_41])).

cnf(refute_0_453, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_452, refute_0_274])).

cnf(refute_0_454, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * skoXC1 <= 2),
    inference(decision, [],
              [refute_0_373, refute_0_378, refute_0_314, refute_0_303,
               refute_0_14])).

cnf(refute_0_455, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1)) | skoXC1 * skoXC1 <= 2),
    inference(resolve, [], [refute_0_453, refute_0_454])).

cnf(refute_0_456, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_455, refute_0_428])).

cnf(refute_0_457, plain,
    (skoXC1 < 0 |
     skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_451, refute_0_456])).

cnf(refute_0_458, plain,
    (skoXC1 < 0 |
     skoXC1 * (-2 + skoXC1 * (skoXC1 * (5/3 + skoXC1 * (skoXC1 * -1/6)))) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_457])).

cnf(refute_0_459, plain, (skoXC1 < 0 | 2 < skoXC1 * skoXC1),
    inference(decision, [],
              [refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_460, plain,
    (skoXC1 < 0 |
     skoXC1 * (-2 + skoXC1 * (skoXC1 * (5/3 + skoXC1 * (skoXC1 * -1/6)))) <
     cos(skoXC1) * (skoXC1 * -2)),
    inference(resolve, [], [refute_0_458, refute_0_459])).

cnf(refute_0_461, plain,
    (skoXC1 < 0 |
     cos(skoXC1) <
     skoXC1 * (-2 + skoXC1 * (skoXC1 * (5/3 + skoXC1 * (skoXC1 * -1/6)))) /
     (skoXC1 * -2) | 0 <= skoXC1 * -2),
    inference(split, [], [refute_0_460])).

cnf(refute_0_462, plain,
    (skoXC1 < 0 |
     cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0 | skoXC1 = 0), inference(arithmetic, [], [refute_0_461])).

cnf(refute_0_463, plain,
    (skoXC1 < 0 |
     cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_462, refute_0_14])).

cnf(refute_0_464, plain, (0 <= skoXC1 | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_465, plain,
    (cos(skoXC1) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_464, refute_0_463])).

cnf(refute_0_466, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_334, refute_0_465])).

cnf(refute_0_467, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_466])).

cnf(refute_0_468, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_469, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800)))))))))),
    inference(resolve, [], [refute_0_467, refute_0_468])).

cnf(refute_0_470, plain,
    (-1 < 1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_396, refute_0_465])).

cnf(refute_0_471, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_470])).

cnf(refute_0_472, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_473, plain,
    (-2 < skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_471, refute_0_472])).

cnf(refute_0_474, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     1 + skoXC1 * (skoXC1 * (-5/6 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_393, refute_0_465])).

cnf(refute_0_475, plain,
    (0 < skoXC1 * (skoXC1 * (-1/3 + skoXC1 * (skoXC1 * 1/12))) |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_474])).

cnf(refute_0_476, plain,
    (0 < skoXC1 * (skoXC1 * (-1/3 + skoXC1 * (skoXC1 * 1/12))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_477, plain,
    (0 < skoXC1 * (skoXC1 * (-1/3 + skoXC1 * (skoXC1 * 1/12)))),
    inference(resolve, [], [refute_0_475, refute_0_476])).

cnf(refute_0_478, plain,
    (2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (2/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 * (-1/20160 + skoXC1 * (skoXC1 * 1/1814400)))))))))) |
     0 < skoXC1),
    inference(decision, [],
              [refute_0_469, refute_0_473, refute_0_477, refute_0_427,
               refute_0_413, refute_0_392, refute_0_395, refute_0_400,
               refute_0_383, refute_0_388, refute_0_373, refute_0_378,
               refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_479, plain,
    (2 <
     skoXC1 *
     (-2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (2/3 +
        skoXC1 *
        (skoXC1 *
         (-1/12 +
          skoXC1 *
          (skoXC1 *
           (1/360 +
            skoXC1 *
            (skoXC1 * (-1/20160 + skoXC1 * (skoXC1 * 1/1814400))))))))))),
    inference(resolve, [], [refute_0_450, refute_0_478])).

cnf(refute_0_480, plain,
    (skoXC1 < 0 |
     2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 * (1/20160 + skoXC1 * (skoXC1 * -1/1814400))))))))))),
    inference(decision, [],
              [refute_0_479, refute_0_469, refute_0_473, refute_0_477,
               refute_0_427, refute_0_413, refute_0_392, refute_0_395,
               refute_0_400, refute_0_383, refute_0_388, refute_0_373,
               refute_0_378, refute_0_314, refute_0_303, refute_0_14])).

cnf(refute_0_481, plain,
    (2 <
     skoXC1 *
     (2 +
      skoXC1 *
      (1 +
       skoXC1 *
       (-2/3 +
        skoXC1 *
        (skoXC1 *
         (1/12 +
          skoXC1 *
          (skoXC1 *
           (-1/360 +
            skoXC1 *
            (skoXC1 * (1/20160 + skoXC1 * (skoXC1 * -1/1814400))))))))))),
    inference(resolve, [], [refute_0_444, refute_0_480])).

cnf(refute_0_482, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_481, refute_0_479, refute_0_469, refute_0_473,
               refute_0_477, refute_0_427, refute_0_413, refute_0_392,
               refute_0_395, refute_0_400, refute_0_383, refute_0_388,
               refute_0_373, refute_0_378, refute_0_314, refute_0_303,
               refute_0_14])).

cnf(refute_0_483, plain,
    ((skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2)) /
     (-2 + skoXC1 * skoXC1) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_438, refute_0_482])).

cnf(refute_0_484, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1/3)) + cos(skoXC1) * (skoXC1 * -2) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-2 + skoXC1 * skoXC1) | 0 < skoXC1 | -2 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_431, refute_0_483])).

cnf(refute_0_485, plain,
    (skoXC1 *
     (2 +
      skoXC1 *
      (skoXC1 *
       (-5/3 +
        skoXC1 *
        (skoXC1 *
         (11/60 +
          skoXC1 *
          (skoXC1 *
           (-11/1260 +
            skoXC1 *
            (skoXC1 *
             (37/181440 +
              skoXC1 *
              (skoXC1 *
               (-1/356400 + skoXC1 * (skoXC1 * 1/39916800)))))))))))) <
     cos(skoXC1) * (skoXC1 * 2) | 0 < skoXC1 | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_484])).

cnf(refute_0_486, plain, (2 < skoXC1 * skoXC1 | 0 < skoXC1),
    inference(decision, [],
              [refute_0_481, refute_0_479, refute_0_469, refute_0_473,
               refute_0_477, refute_0_427, refute_0_413, refute_0_392,
               refute_0_395, refute_0_400, refute_0_383, refute_0_388,
               refute_0_373, refute_0_378, refute_0_314, refute_0_303,
               refute_0_14])).

cnf(refute_0_487, plain,
    (skoXC1 *
     (2 +
      skoXC1 *
      (skoXC1 *
       (-5/3 +
        skoXC1 *
        (skoXC1 *
         (11/60 +
          skoXC1 *
          (skoXC1 *
           (-11/1260 +
            skoXC1 *
            (skoXC1 *
             (37/181440 +
              skoXC1 *
              (skoXC1 *
               (-1/356400 + skoXC1 * (skoXC1 * 1/39916800)))))))))))) <
     cos(skoXC1) * (skoXC1 * 2) | 0 < skoXC1),
    inference(resolve, [], [refute_0_485, refute_0_486])).

cnf(refute_0_488, plain,
    (cos(skoXC1) <
     skoXC1 *
     (2 +
      skoXC1 *
      (skoXC1 *
       (-5/3 +
        skoXC1 *
        (skoXC1 *
         (11/60 +
          skoXC1 *
          (skoXC1 *
           (-11/1260 +
            skoXC1 *
            (skoXC1 *
             (37/181440 +
              skoXC1 *
              (skoXC1 *
               (-1/356400 + skoXC1 * (skoXC1 * 1/39916800)))))))))))) /
     (skoXC1 * 2) | 0 < skoXC1 | 0 <= skoXC1 * 2),
    inference(split, [], [refute_0_487])).

cnf(refute_0_489, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     0 < skoXC1 | 0 <= skoXC1 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_488])).

cnf(refute_0_490, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     0 < skoXC1 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_489, refute_0_14])).

cnf(refute_0_491, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_464, refute_0_490])).

cnf(refute_0_492, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     0 <= skoXC1), inference(resolve, [], [refute_0_430, refute_0_491])).

cnf(refute_0_493, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/20 +
         skoXC1 *
         (skoXC1 *
          (-1/336 +
           skoXC1 *
           (skoXC1 *
            (1/12960 +
             skoXC1 *
             (skoXC1 *
              (-1/887040 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     0 <= skoXC1), inference(arithmetic, [], [refute_0_492])).

cnf(refute_0_494, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (-1/3 +
       skoXC1 *
       (skoXC1 *
        (1/20 +
         skoXC1 *
         (skoXC1 *
          (-1/336 +
           skoXC1 *
           (skoXC1 *
            (1/12960 +
             skoXC1 *
             (skoXC1 *
              (-1/887040 + skoXC1 * (skoXC1 * 1/79833600))))))))))) <= 0),
    inference(decision, [],
              [refute_0_481, refute_0_479, refute_0_469, refute_0_473,
               refute_0_477, refute_0_427, refute_0_413, refute_0_392,
               refute_0_395, refute_0_400, refute_0_383, refute_0_388,
               refute_0_373, refute_0_378, refute_0_314, refute_0_303,
               refute_0_14])).

cnf(refute_0_495, plain, (0 <= skoXC1),
    inference(resolve, [], [refute_0_494, refute_0_493])).

cnf(refute_0_496, plain, (skoXC1 < 0 | 0 < skoXC1 * (skoXC1 * skoXC1)),
    inference(decision, [],
              [refute_0_495, refute_0_481, refute_0_479, refute_0_469,
               refute_0_473, refute_0_477, refute_0_427, refute_0_413,
               refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_497, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_429, refute_0_496])).

cnf(refute_0_498, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2)) /
     (2 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_495, refute_0_497])).

cnf(refute_0_499, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * 1/3)) + cos(skoXC1) * (skoXC1 * 2) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (2 + skoXC1 * (skoXC1 * -1)) | 0 <= 2 + skoXC1 * (skoXC1 * -1)),
    inference(resolve, [], [refute_0_5, refute_0_498])).

cnf(refute_0_500, plain,
    (skoXC1 *
     (-2 +
      skoXC1 *
      (skoXC1 *
       (5/3 +
        skoXC1 *
        (skoXC1 *
         (-11/60 +
          skoXC1 *
          (skoXC1 *
           (11/1260 +
            skoXC1 *
            (skoXC1 *
             (-37/181440 +
              skoXC1 *
              (skoXC1 *
               (1/356400 + skoXC1 * (skoXC1 * -1/39916800)))))))))))) <
     cos(skoXC1) * (skoXC1 * -2) | skoXC1 * skoXC1 <= 2),
    inference(arithmetic, [], [refute_0_499])).

cnf(refute_0_501, plain,
    (skoXC1 *
     (-2 +
      skoXC1 *
      (skoXC1 *
       (5/3 +
        skoXC1 *
        (skoXC1 *
         (-11/60 +
          skoXC1 *
          (skoXC1 *
           (11/1260 +
            skoXC1 *
            (skoXC1 *
             (-37/181440 +
              skoXC1 *
              (skoXC1 *
               (1/356400 + skoXC1 * (skoXC1 * -1/39916800)))))))))))) <
     cos(skoXC1) * (skoXC1 * -2)),
    inference(resolve, [], [refute_0_500, refute_0_428])).

cnf(refute_0_502, plain,
    (cos(skoXC1) <
     skoXC1 *
     (-2 +
      skoXC1 *
      (skoXC1 *
       (5/3 +
        skoXC1 *
        (skoXC1 *
         (-11/60 +
          skoXC1 *
          (skoXC1 *
           (11/1260 +
            skoXC1 *
            (skoXC1 *
             (-37/181440 +
              skoXC1 *
              (skoXC1 *
               (1/356400 + skoXC1 * (skoXC1 * -1/39916800)))))))))))) /
     (skoXC1 * -2) | 0 <= skoXC1 * -2),
    inference(split, [], [refute_0_501])).

cnf(refute_0_503, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     skoXC1 <= 0 | skoXC1 = 0), inference(arithmetic, [], [refute_0_502])).

cnf(refute_0_504, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_503, refute_0_14])).

cnf(refute_0_505, plain, (0 < skoXC1),
    inference(decision, [],
              [refute_0_495, refute_0_481, refute_0_479, refute_0_469,
               refute_0_473, refute_0_477, refute_0_427, refute_0_413,
               refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_506, plain,
    (cos(skoXC1) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600)))))))))))),
    inference(resolve, [], [refute_0_504, refute_0_505])).

cnf(refute_0_507, plain,
    (-1 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600)))))))))))),
    inference(resolve, [], [refute_0_4, refute_0_506])).

cnf(refute_0_508, plain,
    (-2 <
     skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600)))))))))))),
    inference(arithmetic, [], [refute_0_507])).

cnf(refute_0_509, plain,
    (skoXC1 *
     (skoXC1 *
      (-5/6 +
       skoXC1 *
       (skoXC1 *
        (11/120 +
         skoXC1 *
         (skoXC1 *
          (-11/2520 +
           skoXC1 *
           (skoXC1 *
            (37/362880 +
             skoXC1 *
             (skoXC1 *
              (-1/712800 + skoXC1 * (skoXC1 * 1/79833600))))))))))) <= -2),
    inference(decision, [],
              [refute_0_495, refute_0_481, refute_0_479, refute_0_469,
               refute_0_473, refute_0_477, refute_0_427, refute_0_413,
               refute_0_392, refute_0_395, refute_0_400, refute_0_383,
               refute_0_388, refute_0_373, refute_0_378, refute_0_314,
               refute_0_303, refute_0_14])).

cnf(refute_0_510, plain, ($false),
    inference(resolve, [], [refute_0_509, refute_0_508])).
% SZS output end CNFRefutation for sin-cos-problem-diff2a.tptp
