% SZS output start CNFRefutation for sin-recip-3.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_left_mul_divide_neg, axiom, (Y / Z < X | Y <= X * Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sin_upper_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, X - X ^ 3 / 6, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_neg, axiom,
    (0 < X |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(sin_upper_bound_const_1, axiom,
    (~ lgen(R, 1, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_const_m1, axiom,
    (~ lgen(R, Y, -1) | lgen(R, Y, sin(X)))).

cnf(sin_upper_bound_1_neg, axiom,
    (0 < X | ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040, Y) |
     lgen(R, sin(X), Y))).

cnf(sin_upper_bound_3_neg, axiom,
    (0 < X |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040) |
     lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000) | lgen(R, Y, sin(X)))).

fof(sin_recip_3, conjecture,
    (! [X] : (X != 0 => 0 < 3/250 + X ^ 3 * sin(1 / X)))).

fof(subgoal_0, plain, (! [X] : (X != 0 => 0 < 3/250 + X ^ 3 * sin(1 / X))),
    inference(strip, [], [sin_recip_3])).

fof(negate_0_0, plain,
    (~ ! [X] : (X != 0 => 0 < 3/250 + X ^ 3 * sin(1 / X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (X != 0 & 3/250 + X ^ 3 * sin(1 / X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 != 0 & 3/250 + skoXC1 ^ 3 * sin(1 / skoXC1) <= 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (3/250 + skoXC1 ^ 3 * sin(1 / skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 != 0),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (1 < 0 * skoXC1 | 0 <= 1 / skoXC1 | skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) <
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_2, plain, (3/250 + skoXC1 ^ 3 * sin(1 / skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain,
    (sin(1 / skoXC1) * (skoXC1 * (skoXC1 * skoXC1)) <= -3/250),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (skoXC1 * (skoXC1 * skoXC1) <= 0 |
     sin(1 / skoXC1) <= -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(split, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (X_000166 < sin(X_000165) | ~ lgen(1, X_000166, sin(X_000165))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_6, plain,
    (X_000165 < 0 |
     ~ lgen(1, X_000166,
            X_000165 - X_000165 ^ 3 / 6 + X_000165 ^ 5 / 120 -
            X_000165 ^ 7 / 5040 + X_000165 ^ 9 / 362880 -
            X_000165 ^ 11 / 39916800 + X_000165 ^ 13 / 6227020800 -
            X_000165 ^ 15 / 1307674368000) |
     lgen(1, X_000166, sin(X_000165))),
    inference(subst, [], [sin_lower_bound_3_pos])).

cnf(refute_0_7, plain,
    (X_000165 < 0 | X_000166 < sin(X_000165) |
     ~ lgen(1, X_000166,
            X_000165 - X_000165 ^ 3 / 6 + X_000165 ^ 5 / 120 -
            X_000165 ^ 7 / 5040 + X_000165 ^ 9 / 362880 -
            X_000165 ^ 11 / 39916800 + X_000165 ^ 13 / 6227020800 -
            X_000165 ^ 15 / 1307674368000)),
    inference(resolve, [], [refute_0_6, refute_0_5])).

cnf(refute_0_8, plain,
    (X_000165 < 0 | X_000166 < sin(X_000165) |
     X_000165 *
     (1 +
      X_000165 *
      (X_000165 *
       (-1/6 +
        X_000165 *
        (X_000165 *
         (1/120 +
          X_000165 *
          (X_000165 *
           (-1/5040 +
            X_000165 *
            (X_000165 *
             (1/362880 +
              X_000165 *
              (X_000165 *
               (-1/39916800 +
                X_000165 *
                (X_000165 *
                 (1/6227020800 +
                  X_000165 * (X_000165 * -1/1307674368000)))))))))))))) <=
     X_000166), inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (1 / skoXC1 < 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_9])).

cnf(refute_0_11, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain, (1 < 0 * skoXC1 | 0 <= skoXC1 | 1 / skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_13, plain,
    ((-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_14, plain,
    (0 <= skoXC1 * (skoXC1 * skoXC1) |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <= sin(1 / skoXC1)),
    inference(split, [], [refute_0_3])).

cnf(refute_0_15, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <= sin(1 / skoXC1)),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (sin(X_000158) < X_000159 | ~ lgen(1, sin(X_000158), X_000159)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain,
    (0 < X_000158 |
     ~ lgen(1,
            X_000158 - X_000158 ^ 3 / 6 + X_000158 ^ 5 / 120 -
            X_000158 ^ 7 / 5040 + X_000158 ^ 9 / 362880 -
            X_000158 ^ 11 / 39916800 + X_000158 ^ 13 / 6227020800 -
            X_000158 ^ 15 / 1307674368000, X_000159) |
     lgen(1, sin(X_000158), X_000159)),
    inference(subst, [], [sin_upper_bound_3_neg])).

cnf(refute_0_18, plain,
    (0 < X_000158 | sin(X_000158) < X_000159 |
     ~ lgen(1,
            X_000158 - X_000158 ^ 3 / 6 + X_000158 ^ 5 / 120 -
            X_000158 ^ 7 / 5040 + X_000158 ^ 9 / 362880 -
            X_000158 ^ 11 / 39916800 + X_000158 ^ 13 / 6227020800 -
            X_000158 ^ 15 / 1307674368000, X_000159)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (0 < X_000158 | sin(X_000158) < X_000159 |
     X_000159 <=
     X_000158 *
     (1 +
      X_000158 *
      (X_000158 *
       (-1/6 +
        X_000158 *
        (X_000158 *
         (1/120 +
          X_000158 *
          (X_000158 *
           (-1/5040 +
            X_000158 *
            (X_000158 *
             (1/362880 +
              X_000158 *
              (X_000158 *
               (-1/39916800 +
                X_000158 *
                (X_000158 *
                 (1/6227020800 +
                  X_000158 * (X_000158 * -1/1307674368000))))))))))))))),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (sin(1 / skoXC1) < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 < 1 / skoXC1 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000))))))))))))))),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000))))))))))))))),
    inference(resolve, [], [refute_0_15, refute_0_20])).

cnf(refute_0_22, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) <
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_24, plain,
    (X_000138 < sin(X_000137) | ~ lgen(1, X_000138, sin(X_000137))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_25, plain,
    (X_000137 < 0 |
     ~ lgen(1, X_000138,
            X_000137 - X_000137 ^ 3 / 6 + X_000137 ^ 5 / 120 -
            X_000137 ^ 7 / 5040 + X_000137 ^ 9 / 362880 -
            X_000137 ^ 11 / 39916800) | lgen(1, X_000138, sin(X_000137))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_26, plain,
    (X_000137 < 0 | X_000138 < sin(X_000137) |
     ~ lgen(1, X_000138,
            X_000137 - X_000137 ^ 3 / 6 + X_000137 ^ 5 / 120 -
            X_000137 ^ 7 / 5040 + X_000137 ^ 9 / 362880 -
            X_000137 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_25, refute_0_24])).

cnf(refute_0_27, plain,
    (X_000137 < 0 | X_000138 < sin(X_000137) |
     X_000137 *
     (1 +
      X_000137 *
      (X_000137 *
       (-1/6 +
        X_000137 *
        (X_000137 *
         (1/120 +
          X_000137 *
          (X_000137 *
           (-1/5040 +
            X_000137 *
            (X_000137 *
             (1/362880 + X_000137 * (X_000137 * -1/39916800)))))))))) <=
     X_000138), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (1 / skoXC1 < 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 * (1 / skoXC1 * -1/39916800)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 * (1 / skoXC1 * -1/39916800)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_28])).

cnf(refute_0_30, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    ((-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_32, plain,
    (sin(X_000130) < X_000131 | ~ lgen(1, sin(X_000130), X_000131)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_33, plain,
    (0 < X_000130 |
     ~ lgen(1,
            X_000130 - X_000130 ^ 3 / 6 + X_000130 ^ 5 / 120 -
            X_000130 ^ 7 / 5040 + X_000130 ^ 9 / 362880 -
            X_000130 ^ 11 / 39916800, X_000131) |
     lgen(1, sin(X_000130), X_000131)),
    inference(subst, [], [sin_upper_bound_2_neg])).

cnf(refute_0_34, plain,
    (0 < X_000130 | sin(X_000130) < X_000131 |
     ~ lgen(1,
            X_000130 - X_000130 ^ 3 / 6 + X_000130 ^ 5 / 120 -
            X_000130 ^ 7 / 5040 + X_000130 ^ 9 / 362880 -
            X_000130 ^ 11 / 39916800, X_000131)),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (0 < X_000130 | sin(X_000130) < X_000131 |
     X_000131 <=
     X_000130 *
     (1 +
      X_000130 *
      (X_000130 *
       (-1/6 +
        X_000130 *
        (X_000130 *
         (1/120 +
          X_000130 *
          (X_000130 *
           (-1/5040 +
            X_000130 *
            (X_000130 *
             (1/362880 + X_000130 * (X_000130 * -1/39916800))))))))))),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (sin(1 / skoXC1) < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 < 1 / skoXC1 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 + 1 / skoXC1 * (1 / skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 + 1 / skoXC1 * (1 / skoXC1 * -1/39916800))))))))))),
    inference(resolve, [], [refute_0_15, refute_0_36])).

cnf(refute_0_38, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) <
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_40, plain,
    (X_000109 < sin(X_000108) | ~ lgen(1, X_000109, sin(X_000108))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_41, plain,
    (X_000108 < 0 |
     ~ lgen(1, X_000109,
            X_000108 - X_000108 ^ 3 / 6 + X_000108 ^ 5 / 120 -
            X_000108 ^ 7 / 5040) | lgen(1, X_000109, sin(X_000108))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_42, plain,
    (X_000108 < 0 | X_000109 < sin(X_000108) |
     ~ lgen(1, X_000109,
            X_000108 - X_000108 ^ 3 / 6 + X_000108 ^ 5 / 120 -
            X_000108 ^ 7 / 5040)),
    inference(resolve, [], [refute_0_41, refute_0_40])).

cnf(refute_0_43, plain,
    (X_000108 < 0 | X_000109 < sin(X_000108) |
     X_000108 *
     (1 +
      X_000108 *
      (X_000108 *
       (-1/6 +
        X_000108 *
        (X_000108 * (1/120 + X_000108 * (X_000108 * -1/5040)))))) <=
     X_000109), inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (1 / skoXC1 < 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_44])).

cnf(refute_0_46, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    ((-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 <=
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_48, plain,
    (sin(X_000101) < X_000102 | ~ lgen(1, sin(X_000101), X_000102)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_49, plain,
    (0 < X_000101 |
     ~ lgen(1,
            X_000101 - X_000101 ^ 3 / 6 + X_000101 ^ 5 / 120 -
            X_000101 ^ 7 / 5040, X_000102) |
     lgen(1, sin(X_000101), X_000102)),
    inference(subst, [], [sin_upper_bound_1_neg])).

cnf(refute_0_50, plain,
    (0 < X_000101 | sin(X_000101) < X_000102 |
     ~ lgen(1,
            X_000101 - X_000101 ^ 3 / 6 + X_000101 ^ 5 / 120 -
            X_000101 ^ 7 / 5040, X_000102)),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    (0 < X_000101 | sin(X_000101) < X_000102 |
     X_000102 <=
     X_000101 *
     (1 +
      X_000101 *
      (X_000101 *
       (-1/6 +
        X_000101 *
        (X_000101 * (1/120 + X_000101 * (X_000101 * -1/5040))))))),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (sin(1 / skoXC1) < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 < 1 / skoXC1 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040))))))),
    inference(subst, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040))))))),
    inference(resolve, [], [refute_0_15, refute_0_52])).

cnf(refute_0_54, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) <
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <= -3/250 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_56, plain,
    (X_000078 < sin(X_000077) | ~ lgen(1, X_000078, sin(X_000077))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_57, plain,
    (X_000077 < 0 | ~ lgen(1, X_000078, X_000077 - X_000077 ^ 3 / 6) |
     lgen(1, X_000078, sin(X_000077))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_58, plain,
    (X_000077 < 0 | X_000078 < sin(X_000077) |
     ~ lgen(1, X_000078, X_000077 - X_000077 ^ 3 / 6)),
    inference(resolve, [], [refute_0_57, refute_0_56])).

cnf(refute_0_59, plain,
    (X_000077 < 0 | X_000078 < sin(X_000077) |
     X_000077 * (1 + X_000077 * (X_000077 * -1/6)) <= X_000078),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (1 / skoXC1 < 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) < sin(1 / skoXC1) |
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6)) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6)) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_60])).

cnf(refute_0_62, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    ((-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) <
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 <= skoXC1 * (skoXC1 * skoXC1) |
     -1/6 + skoXC1 * skoXC1 <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) * (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_64, plain,
    (sin(X_000070) < X_000071 | ~ lgen(1, sin(X_000070), X_000071)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_65, plain,
    (0 < X_000070 | ~ lgen(1, X_000070 - X_000070 ^ 3 / 6, X_000071) |
     lgen(1, sin(X_000070), X_000071)),
    inference(subst, [], [sin_upper_bound_0_neg])).

cnf(refute_0_66, plain,
    (0 < X_000070 | sin(X_000070) < X_000071 |
     ~ lgen(1, X_000070 - X_000070 ^ 3 / 6, X_000071)),
    inference(resolve, [], [refute_0_65, refute_0_64])).

cnf(refute_0_67, plain,
    (0 < X_000070 | sin(X_000070) < X_000071 |
     X_000071 <= X_000070 * (1 + X_000070 * (X_000070 * -1/6))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (sin(1 / skoXC1) < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 < 1 / skoXC1 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6))),
    inference(subst, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6))),
    inference(resolve, [], [refute_0_15, refute_0_68])).

cnf(refute_0_70, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * skoXC1 = 0), inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain,
    (sin(X_000009) < X_000010 | ~ lgen(1, sin(X_000009), X_000010)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_72, plain,
    (~ lgen(1, 1, X_000010) | lgen(1, sin(X_000009), X_000010)),
    inference(subst, [], [sin_upper_bound_const_1])).

cnf(refute_0_73, plain,
    (sin(X_000009) < X_000010 | ~ lgen(1, 1, X_000010)),
    inference(resolve, [], [refute_0_72, refute_0_71])).

cnf(refute_0_74, plain, (sin(X_000009) < X_000010 | X_000010 <= 1),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (sin(1 / skoXC1) < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <= 1),
    inference(subst, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <= 1),
    inference(resolve, [], [refute_0_15, refute_0_75])).

cnf(refute_0_77, plain,
    (1 < -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     0 <= skoXC1 * (skoXC1 * skoXC1) |
     1 * (skoXC1 * (skoXC1 * skoXC1)) <= -3/250),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_78, plain,
    (0 <= skoXC1 * (skoXC1 * skoXC1) |
     1 * (skoXC1 * (skoXC1 * skoXC1)) <= -3/250 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * skoXC1) <= -3/250),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (X_000019 < sin(X_000018) | ~ lgen(1, X_000019, sin(X_000018))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_81, plain,
    (~ lgen(1, X_000019, -1) | lgen(1, X_000019, sin(X_000018))),
    inference(subst, [], [sin_lower_bound_const_m1])).

cnf(refute_0_82, plain,
    (X_000019 < sin(X_000018) | ~ lgen(1, X_000019, -1)),
    inference(resolve, [], [refute_0_81, refute_0_80])).

cnf(refute_0_83, plain, (X_000019 < sin(X_000018) | -1 <= X_000019),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) < sin(1 / skoXC1) |
     -1 <= -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (-1 <= -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_84])).

cnf(refute_0_86, plain,
    (-3/250 / (skoXC1 * (skoXC1 * skoXC1)) < -1 |
     -1 * (skoXC1 * (skoXC1 * skoXC1)) <= -3/250 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_87, plain,
    (-1 * (skoXC1 * (skoXC1 * skoXC1)) <= -3/250 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * -1)) <= -3/250 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain, (skoXC1 != 0),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_90, plain,
    (skoXC1 * skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(decision, [], [refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_91, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_70, refute_0_90])).

cnf(refute_0_92, plain,
    (0 < 1 / skoXC1 | 0 <= skoXC1 * (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -1/6 + skoXC1 * skoXC1 <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) * (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_91, refute_0_63])).

cnf(refute_0_93, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 58/375 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 58/375),
    inference(decision, [], [refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_95, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 58/375),
    inference(resolve, [], [refute_0_93, refute_0_94])).

cnf(refute_0_96, plain,
    (1 < 0 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 * skoXC1 <= 58/375),
    inference(resolve, [], [refute_0_12, refute_0_95])).

cnf(refute_0_97, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * skoXC1 <= 58/375), inference(arithmetic, [], [refute_0_96])).

cnf(refute_0_98, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1 |
     skoXC1 * skoXC1 <= 58/375),
    inference(decision, [], [refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_99, plain, (0 <= skoXC1 | skoXC1 * skoXC1 <= 58/375),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain,
    (skoXC1 * skoXC1 != 0 | skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(decision, [],
              [refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_101, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_62, refute_0_100])).

cnf(refute_0_102, plain,
    (1 / skoXC1 < 0 |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 * (skoXC1 * skoXC1)) <= -3/250 |
     skoXC1 * (skoXC1 * skoXC1) <= 0),
    inference(resolve, [], [refute_0_101, refute_0_55])).

cnf(refute_0_103, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * skoXC1 <= 58/375 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * skoXC1 <= 58/375 |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_105, plain,
    (1 / skoXC1 < 0 | skoXC1 * skoXC1 <= 58/375 | skoXC1 = 0),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain, (skoXC1 != 0 | skoXC1 * skoXC1 <= 58/375),
    inference(decision, [],
              [refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_107, plain, (1 / skoXC1 < 0 | skoXC1 * skoXC1 <= 58/375),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain,
    (1 < 0 * skoXC1 | skoXC1 * skoXC1 <= 58/375 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_107])).

cnf(refute_0_109, plain, (skoXC1 * skoXC1 <= 58/375 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_108])).

cnf(refute_0_110, plain, (0 < skoXC1 | skoXC1 * skoXC1 <= 58/375),
    inference(decision, [],
              [refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_111, plain, (skoXC1 * skoXC1 <= 58/375),
    inference(resolve, [], [refute_0_109, refute_0_110])).

cnf(refute_0_112, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_113, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_54, refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_115, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_113, refute_0_114])).

cnf(refute_0_116, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(resolve, [], [refute_0_115, refute_0_90])).

cnf(refute_0_117, plain,
    (0 < 1 / skoXC1 |
     0 <=
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(resolve, [], [refute_0_116, refute_0_47])).

cnf(refute_0_118, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 = 0), inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_120, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 = 0),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040),
    inference(decision, [],
              [refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_122, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (1 < 0 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(resolve, [], [refute_0_12, refute_0_122])).

cnf(refute_0_124, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040),
    inference(decision, [],
              [refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_126, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(resolve, [], [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) != 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_128, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_46, refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) != 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_130, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_128, refute_0_129])).

cnf(refute_0_131, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_130, refute_0_100])).

cnf(refute_0_132, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(resolve, [], [refute_0_131, refute_0_39])).

cnf(refute_0_133, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_135, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_137, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 = 0),
    inference(resolve, [], [refute_0_135, refute_0_136])).

cnf(refute_0_138, plain,
    (skoXC1 != 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_139, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_139])).

cnf(refute_0_141, plain,
    (skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040 | skoXC1 <= 0), inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040),
    inference(decision, [],
              [refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_143, plain,
    (skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))) <=
     1/5040), inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) !=
     0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_143, refute_0_126, refute_0_111, refute_0_99,
               refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_145, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_38, refute_0_144])).

cnf(refute_0_146, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_143, refute_0_126, refute_0_111, refute_0_99,
               refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_147, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_147, refute_0_112])).

cnf(refute_0_149, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_148, refute_0_114])).

cnf(refute_0_150, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(resolve, [], [refute_0_149, refute_0_90])).

cnf(refute_0_151, plain,
    (0 < 1 / skoXC1 |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(resolve, [], [refute_0_150, refute_0_31])).

cnf(refute_0_152, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 = 0), inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_143, refute_0_126, refute_0_111, refute_0_99,
               refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_154, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800),
    inference(decision, [],
              [refute_0_143, refute_0_126, refute_0_111, refute_0_99,
               refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_156, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(resolve, [], [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (1 < 0 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(resolve, [], [refute_0_12, refute_0_156])).

cnf(refute_0_158, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800),
    inference(decision, [],
              [refute_0_143, refute_0_126, refute_0_111, refute_0_99,
               refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_160, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) !=
     0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_162, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_30, refute_0_161])).

cnf(refute_0_163, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) != 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_164, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_164, refute_0_127])).

cnf(refute_0_166, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_165, refute_0_129])).

cnf(refute_0_167, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_166, refute_0_100])).

cnf(refute_0_168, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(resolve, [], [refute_0_167, refute_0_23])).

cnf(refute_0_169, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_171, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 * (skoXC1 * skoXC1) <= 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_173, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(resolve, [], [refute_0_171, refute_0_172])).

cnf(refute_0_174, plain,
    (skoXC1 != 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_175, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(resolve, [], [refute_0_173, refute_0_174])).

cnf(refute_0_176, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800 | skoXC1 <= 0), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (0 < skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800),
    inference(decision, [],
              [refute_0_160, refute_0_143, refute_0_126, refute_0_111,
               refute_0_99, refute_0_79, refute_0_88, refute_0_89])).

cnf(refute_0_179, plain,
    (skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))) <=
     1/39916800), inference(resolve, [], [refute_0_177, refute_0_178])).

cnf(refute_0_180, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) !=
     0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_181, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_22, refute_0_180])).

cnf(refute_0_182, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_183, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_181, refute_0_182])).

cnf(refute_0_184, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_183, refute_0_144])).

cnf(refute_0_185, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_184, refute_0_146])).

cnf(refute_0_186, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_185, refute_0_112])).

cnf(refute_0_187, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_186, refute_0_114])).

cnf(refute_0_188, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) <=
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(resolve, [], [refute_0_187, refute_0_90])).

cnf(refute_0_189, plain,
    (0 < 1 / skoXC1 |
     0 <=
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(resolve, [], [refute_0_188, refute_0_13])).

cnf(refute_0_190, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))))))) <=
     1/1307674368000 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_189])).

cnf(refute_0_191, plain,
    (1/1307674368000 <
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))))))) <=
     0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_192, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))))))) <=
     0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_190, refute_0_191])).

cnf(refute_0_193, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * -1)))))))))))))) |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_194, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_192, refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 != 0 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_196, plain,
    (0 < 1 / skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (1 < 0 * skoXC1 | 0 <= skoXC1 |
     skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_196])).

cnf(refute_0_198, plain,
    (0 <= skoXC1 | skoXC1 * (skoXC1 * (skoXC1 * -1)) <= 0),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (0 < skoXC1 * (skoXC1 * (skoXC1 * -1)) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_179, refute_0_160, refute_0_143, refute_0_126,
               refute_0_111, refute_0_99, refute_0_79, refute_0_88,
               refute_0_89])).

cnf(refute_0_200, plain, (0 <= skoXC1),
    inference(resolve, [], [refute_0_198, refute_0_199])).

cnf(refute_0_201, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) !=
     0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_202, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_11, refute_0_201])).

cnf(refute_0_203, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) != 0 |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_204, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_202, refute_0_203])).

cnf(refute_0_205, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_204, refute_0_161])).

cnf(refute_0_206, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_205, refute_0_163])).

cnf(refute_0_207, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_206, refute_0_127])).

cnf(refute_0_208, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_207, refute_0_129])).

cnf(refute_0_209, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_210, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_208, refute_0_209])).

cnf(refute_0_211, plain, (skoXC1 * skoXC1 != 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_212, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_210, refute_0_211])).

cnf(refute_0_213, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -3/250 / (skoXC1 * (skoXC1 * skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(resolve, [], [refute_0_212, refute_0_1])).

cnf(refute_0_214, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))))))) <=
     1/1307674368000 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_213])).

cnf(refute_0_215, plain,
    (1/1307674368000 <
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-58/375 + skoXC1 * skoXC1)))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_216, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 = 0), inference(resolve, [], [refute_0_214, refute_0_215])).

cnf(refute_0_217, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_218, plain, (1 / skoXC1 < 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_216, refute_0_217])).

cnf(refute_0_219, plain, (1 / skoXC1 < 0),
    inference(resolve, [], [refute_0_218, refute_0_89])).

cnf(refute_0_220, plain, (1 < 0 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_219])).

cnf(refute_0_221, plain, (skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain, (0 < skoXC1),
    inference(decision, [],
              [refute_0_200, refute_0_179, refute_0_160, refute_0_143,
               refute_0_126, refute_0_111, refute_0_99, refute_0_79,
               refute_0_88, refute_0_89])).

cnf(refute_0_223, plain, ($false),
    inference(resolve, [], [refute_0_221, refute_0_222])).
% SZS output end CNFRefutation for sin-recip-3.tptp
