% SZS output start CNFRefutation for sin-recip-2.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_const_m1, axiom,
    (~ lgen(R, Y, -1) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040) |
     lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000) | lgen(R, Y, sin(X)))).

fof(sin_recip_2, conjecture,
    (! [X] : (0 < X => 0 < 31/625 + X ^ 2 * sin(1 / X)))).

fof(subgoal_0, plain, (! [X] : (0 < X => 0 < 31/625 + X ^ 2 * sin(1 / X))),
    inference(strip, [], [sin_recip_2])).

fof(negate_0_0, plain,
    (~ ! [X] : (0 < X => 0 < 31/625 + X ^ 2 * sin(1 / X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (0 < X & 31/625 + X ^ 2 * sin(1 / X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 < skoXC1 & 31/625 + skoXC1 ^ 2 * sin(1 / skoXC1) <= 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (31/625 + skoXC1 ^ 2 * sin(1 / skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (1 < 0 * skoXC1 | 0 <= 1 / skoXC1 | skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (-31/625 / (skoXC1 * skoXC1) <
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_2, plain, (31/625 + skoXC1 ^ 2 * sin(1 / skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_3, plain, (sin(1 / skoXC1) * (skoXC1 * skoXC1) <= -31/625),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (skoXC1 * skoXC1 <= 0 |
     sin(1 / skoXC1) <= -31/625 / (skoXC1 * skoXC1)),
    inference(split, [], [refute_0_3])).

cnf(refute_0_5, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_6, plain, (0 < skoXC1 * skoXC1),
    inference(decision, [], [refute_0_5])).

cnf(refute_0_7, plain, (sin(1 / skoXC1) <= -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_4, refute_0_6])).

cnf(refute_0_8, plain,
    (X_000167 < sin(X_000166) | ~ lgen(1, X_000167, sin(X_000166))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (X_000166 < 0 |
     ~ lgen(1, X_000167,
            X_000166 - X_000166 ^ 3 / 6 + X_000166 ^ 5 / 120 -
            X_000166 ^ 7 / 5040 + X_000166 ^ 9 / 362880 -
            X_000166 ^ 11 / 39916800 + X_000166 ^ 13 / 6227020800 -
            X_000166 ^ 15 / 1307674368000) |
     lgen(1, X_000167, sin(X_000166))),
    inference(subst, [], [sin_lower_bound_3_pos])).

cnf(refute_0_10, plain,
    (X_000166 < 0 | X_000167 < sin(X_000166) |
     ~ lgen(1, X_000167,
            X_000166 - X_000166 ^ 3 / 6 + X_000166 ^ 5 / 120 -
            X_000166 ^ 7 / 5040 + X_000166 ^ 9 / 362880 -
            X_000166 ^ 11 / 39916800 + X_000166 ^ 13 / 6227020800 -
            X_000166 ^ 15 / 1307674368000)),
    inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (X_000166 < 0 | X_000167 < sin(X_000166) |
     X_000166 *
     (1 +
      X_000166 *
      (X_000166 *
       (-1/6 +
        X_000166 *
        (X_000166 *
         (1/120 +
          X_000166 *
          (X_000166 *
           (-1/5040 +
            X_000166 *
            (X_000166 *
             (1/362880 +
              X_000166 *
              (X_000166 *
               (-1/39916800 +
                X_000166 *
                (X_000166 *
                 (1/6227020800 +
                  X_000166 * (X_000166 * -1/1307674368000)))))))))))))) <=
     X_000167), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 / skoXC1 < 0 | -31/625 / (skoXC1 * skoXC1) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1)), inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 *
              (1 / skoXC1 *
               (-1/39916800 +
                1 / skoXC1 *
                (1 / skoXC1 *
                 (1/6227020800 +
                  1 / skoXC1 *
                  (1 / skoXC1 * -1/1307674368000)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (-31/625 / (skoXC1 * skoXC1) <
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_16, plain,
    (X_000139 < sin(X_000138) | ~ lgen(1, X_000139, sin(X_000138))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_17, plain,
    (X_000138 < 0 |
     ~ lgen(1, X_000139,
            X_000138 - X_000138 ^ 3 / 6 + X_000138 ^ 5 / 120 -
            X_000138 ^ 7 / 5040 + X_000138 ^ 9 / 362880 -
            X_000138 ^ 11 / 39916800) | lgen(1, X_000139, sin(X_000138))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_18, plain,
    (X_000138 < 0 | X_000139 < sin(X_000138) |
     ~ lgen(1, X_000139,
            X_000138 - X_000138 ^ 3 / 6 + X_000138 ^ 5 / 120 -
            X_000138 ^ 7 / 5040 + X_000138 ^ 9 / 362880 -
            X_000138 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (X_000138 < 0 | X_000139 < sin(X_000138) |
     X_000138 *
     (1 +
      X_000138 *
      (X_000138 *
       (-1/6 +
        X_000138 *
        (X_000138 *
         (1/120 +
          X_000138 *
          (X_000138 *
           (-1/5040 +
            X_000138 *
            (X_000138 *
             (1/362880 + X_000138 * (X_000138 * -1/39916800)))))))))) <=
     X_000139), inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (1 / skoXC1 < 0 | -31/625 / (skoXC1 * skoXC1) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 * (1 / skoXC1 * -1/39916800)))))))))) <=
     -31/625 / (skoXC1 * skoXC1)), inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 *
         (1/120 +
          1 / skoXC1 *
          (1 / skoXC1 *
           (-1/5040 +
            1 / skoXC1 *
            (1 / skoXC1 *
             (1/362880 +
              1 / skoXC1 * (1 / skoXC1 * -1/39916800)))))))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_7, refute_0_20])).

cnf(refute_0_22, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-31/625 / (skoXC1 * skoXC1) <
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_24, plain,
    (X_000110 < sin(X_000109) | ~ lgen(1, X_000110, sin(X_000109))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_25, plain,
    (X_000109 < 0 |
     ~ lgen(1, X_000110,
            X_000109 - X_000109 ^ 3 / 6 + X_000109 ^ 5 / 120 -
            X_000109 ^ 7 / 5040) | lgen(1, X_000110, sin(X_000109))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_26, plain,
    (X_000109 < 0 | X_000110 < sin(X_000109) |
     ~ lgen(1, X_000110,
            X_000109 - X_000109 ^ 3 / 6 + X_000109 ^ 5 / 120 -
            X_000109 ^ 7 / 5040)),
    inference(resolve, [], [refute_0_25, refute_0_24])).

cnf(refute_0_27, plain,
    (X_000109 < 0 | X_000110 < sin(X_000109) |
     X_000109 *
     (1 +
      X_000109 *
      (X_000109 *
       (-1/6 +
        X_000109 *
        (X_000109 * (1/120 + X_000109 * (X_000109 * -1/5040)))))) <=
     X_000110), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (1 / skoXC1 < 0 | -31/625 / (skoXC1 * skoXC1) < sin(1 / skoXC1) |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040)))))) <=
     -31/625 / (skoXC1 * skoXC1)), inference(subst, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 *
     (1 +
      1 / skoXC1 *
      (1 / skoXC1 *
       (-1/6 +
        1 / skoXC1 *
        (1 / skoXC1 * (1/120 + 1 / skoXC1 * (1 / skoXC1 * -1/5040)))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_7, refute_0_28])).

cnf(refute_0_30, plain,
    (1 / skoXC1 < 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (-31/625 / (skoXC1 * skoXC1) <
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) |
     skoXC1 * (skoXC1 * skoXC1) <= 0 |
     -1/6 + skoXC1 * skoXC1 <=
     -31/625 / (skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * skoXC1))),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_32, plain,
    (X_000079 < sin(X_000078) | ~ lgen(1, X_000079, sin(X_000078))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_33, plain,
    (X_000078 < 0 | ~ lgen(1, X_000079, X_000078 - X_000078 ^ 3 / 6) |
     lgen(1, X_000079, sin(X_000078))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_34, plain,
    (X_000078 < 0 | X_000079 < sin(X_000078) |
     ~ lgen(1, X_000079, X_000078 - X_000078 ^ 3 / 6)),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (X_000078 < 0 | X_000079 < sin(X_000078) |
     X_000078 * (1 + X_000078 * (X_000078 * -1/6)) <= X_000079),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (1 / skoXC1 < 0 | -31/625 / (skoXC1 * skoXC1) < sin(1 / skoXC1) |
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6)) <=
     -31/625 / (skoXC1 * skoXC1)), inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (1 / skoXC1 < 0 |
     1 / skoXC1 * (1 + 1 / skoXC1 * (1 / skoXC1 * -1/6)) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_7, refute_0_36])).

cnf(refute_0_38, plain,
    (1 / skoXC1 < 0 |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) <=
     -31/625 / (skoXC1 * skoXC1) | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (X_000019 < sin(X_000018) | ~ lgen(1, X_000019, sin(X_000018))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_40, plain,
    (~ lgen(1, X_000019, -1) | lgen(1, X_000019, sin(X_000018))),
    inference(subst, [], [sin_lower_bound_const_m1])).

cnf(refute_0_41, plain,
    (X_000019 < sin(X_000018) | ~ lgen(1, X_000019, -1)),
    inference(resolve, [], [refute_0_40, refute_0_39])).

cnf(refute_0_42, plain, (X_000019 < sin(X_000018) | -1 <= X_000019),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (-31/625 / (skoXC1 * skoXC1) < sin(1 / skoXC1) |
     -1 <= -31/625 / (skoXC1 * skoXC1)),
    inference(subst, [], [refute_0_42])).

cnf(refute_0_44, plain, (-1 <= -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_7, refute_0_43])).

cnf(refute_0_45, plain,
    (-31/625 / (skoXC1 * skoXC1) < -1 | -1 * (skoXC1 * skoXC1) <= -31/625 |
     skoXC1 * skoXC1 <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_46, plain,
    (-1 * (skoXC1 * skoXC1) <= -31/625 | skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    (skoXC1 * (skoXC1 * -1) <= -31/625 | skoXC1 * skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (0 < skoXC1 * skoXC1 | skoXC1 * (skoXC1 * -1) <= -31/625),
    inference(decision, [], [refute_0_5])).

cnf(refute_0_49, plain, (skoXC1 * (skoXC1 * -1) <= -31/625),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (skoXC1 * skoXC1 != 0),
    inference(decision, [], [refute_0_49, refute_0_5])).

cnf(refute_0_51, plain,
    (1 / skoXC1 < 0 |
     (-1/6 + skoXC1 * skoXC1) / (skoXC1 * (skoXC1 * skoXC1)) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_38, refute_0_50])).

cnf(refute_0_52, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     -1/6 + skoXC1 * skoXC1 <=
     -31/625 / (skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * skoXC1))),
    inference(resolve, [], [refute_0_51, refute_0_31])).

cnf(refute_0_53, plain,
    (1 / skoXC1 < 0 | skoXC1 * (skoXC1 * skoXC1) <= 0 |
     skoXC1 * (31/625 + skoXC1) <= 1/6 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (0 < skoXC1 * (skoXC1 * skoXC1) | skoXC1 * (31/625 + skoXC1) <= 1/6 |
     skoXC1 = 0), inference(decision, [], [refute_0_49, refute_0_5])).

cnf(refute_0_55, plain,
    (1 / skoXC1 < 0 | skoXC1 * (31/625 + skoXC1) <= 1/6 | skoXC1 = 0),
    inference(resolve, [], [refute_0_53, refute_0_54])).

cnf(refute_0_56, plain, (skoXC1 != 0 | skoXC1 * (31/625 + skoXC1) <= 1/6),
    inference(decision, [], [refute_0_49, refute_0_5])).

cnf(refute_0_57, plain,
    (1 / skoXC1 < 0 | skoXC1 * (31/625 + skoXC1) <= 1/6),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (1 < 0 * skoXC1 | skoXC1 * (31/625 + skoXC1) <= 1/6 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_57])).

cnf(refute_0_59, plain,
    (1 < 0 * skoXC1 | skoXC1 * (31/625 + skoXC1) <= 1/6),
    inference(resolve, [], [refute_0_58, refute_0_5])).

cnf(refute_0_60, plain, (skoXC1 * (31/625 + skoXC1) <= 1/6),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [], [refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_62, plain,
    (1 / skoXC1 < 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_30, refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) != 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [], [refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_64, plain,
    (1 / skoXC1 < 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -31/625 / (skoXC1 * skoXC1) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_62, refute_0_63])).

cnf(refute_0_65, plain,
    (1 / skoXC1 < 0 |
     (-1/5040 +
      skoXC1 *
      (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))) /
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_64, refute_0_50])).

cnf(refute_0_66, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     -1/5040 +
     skoXC1 *
     (skoXC1 * (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))),
    inference(resolve, [], [refute_0_65, refute_0_23])).

cnf(refute_0_67, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040 | skoXC1 = 0), inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (0 <
     skoXC1 *
     (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))) |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040 | skoXC1 = 0),
    inference(decision, [], [refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_69, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040 | skoXC1 = 0),
    inference(resolve, [], [refute_0_67, refute_0_68])).

cnf(refute_0_70, plain,
    (skoXC1 != 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040),
    inference(decision, [], [refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_71, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040), inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_71])).

cnf(refute_0_73, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040), inference(resolve, [], [refute_0_72, refute_0_5])).

cnf(refute_0_74, plain,
    (skoXC1 *
     (skoXC1 *
      (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))) <=
     1/5040), inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) !=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_74, refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_76, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_22, refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) != 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_74, refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_78, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_76, refute_0_77])).

cnf(refute_0_79, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_78, refute_0_61])).

cnf(refute_0_80, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_79, refute_0_63])).

cnf(refute_0_81, plain,
    (1 / skoXC1 < 0 |
     (-1/39916800 +
      skoXC1 *
      (skoXC1 *
       (1/362880 +
        skoXC1 *
        (skoXC1 *
         (-1/5040 +
          skoXC1 *
          (skoXC1 *
           (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_80, refute_0_50])).

cnf(refute_0_82, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     -1/39916800 +
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 + skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))),
    inference(resolve, [], [refute_0_81, refute_0_15])).

cnf(refute_0_83, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800 | skoXC1 = 0), inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))) |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(decision, [],
              [refute_0_74, refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_85, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800 | skoXC1 = 0),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (skoXC1 != 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800),
    inference(decision, [],
              [refute_0_74, refute_0_60, refute_0_49, refute_0_5])).

cnf(refute_0_87, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800), inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_87])).

cnf(refute_0_89, plain,
    (1 < 0 * skoXC1 |
     skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800), inference(resolve, [], [refute_0_88, refute_0_5])).

cnf(refute_0_90, plain,
    (skoXC1 *
     (skoXC1 *
      (1/362880 +
       skoXC1 *
       (skoXC1 *
        (-1/5040 +
         skoXC1 *
         (skoXC1 *
          (1/120 +
           skoXC1 * (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))) <=
     1/39916800), inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))) !=
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_90, refute_0_74, refute_0_60, refute_0_49,
               refute_0_5])).

cnf(refute_0_92, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_14, refute_0_91])).

cnf(refute_0_93, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))) != 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_90, refute_0_74, refute_0_60, refute_0_49,
               refute_0_5])).

cnf(refute_0_94, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))) =
     0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_94, refute_0_75])).

cnf(refute_0_96, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_95, refute_0_77])).

cnf(refute_0_97, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (skoXC1 * skoXC1)) = 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_96, refute_0_61])).

cnf(refute_0_98, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_97, refute_0_63])).

cnf(refute_0_99, plain,
    (1 / skoXC1 < 0 |
     (-1/1307674368000 +
      skoXC1 *
      (skoXC1 *
       (1/6227020800 +
        skoXC1 *
        (skoXC1 *
         (-1/39916800 +
          skoXC1 *
          (skoXC1 *
           (1/362880 +
            skoXC1 *
            (skoXC1 *
             (-1/5040 +
              skoXC1 *
              (skoXC1 *
               (1/120 +
                skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1)))))))))))))) <=
     -31/625 / (skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_98, refute_0_50])).

cnf(refute_0_100, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     -1/1307674368000 +
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 * (skoXC1 * (-1/6 + skoXC1 * skoXC1)))))))))))) <=
     -31/625 / (skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 *
                (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))))),
    inference(resolve, [], [refute_0_99, refute_0_1])).

cnf(refute_0_101, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 *
               (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))))))) <=
     1/1307674368000 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (1/1307674368000 <
     skoXC1 *
     (skoXC1 *
      (1/6227020800 +
       skoXC1 *
       (skoXC1 *
        (-1/39916800 +
         skoXC1 *
         (skoXC1 *
          (1/362880 +
           skoXC1 *
           (skoXC1 *
            (-1/5040 +
             skoXC1 *
             (skoXC1 *
              (1/120 +
               skoXC1 *
               (skoXC1 * (-1/6 + skoXC1 * (31/625 + skoXC1))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_90, refute_0_74, refute_0_60, refute_0_49,
               refute_0_5])).

cnf(refute_0_103, plain,
    (1 / skoXC1 < 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 *
               (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) <= 0 |
     skoXC1 = 0), inference(resolve, [], [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (skoXC1 *
        (skoXC1 *
         (skoXC1 *
          (skoXC1 *
           (skoXC1 *
            (skoXC1 *
             (skoXC1 *
              (skoXC1 * (skoXC1 * (skoXC1 * (skoXC1 * skoXC1))))))))))))) |
     skoXC1 = 0),
    inference(decision, [],
              [refute_0_90, refute_0_74, refute_0_60, refute_0_49,
               refute_0_5])).

cnf(refute_0_105, plain, (1 / skoXC1 < 0 | skoXC1 = 0),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain, (skoXC1 != 0),
    inference(decision, [],
              [refute_0_90, refute_0_74, refute_0_60, refute_0_49,
               refute_0_5])).

cnf(refute_0_107, plain, (1 / skoXC1 < 0),
    inference(resolve, [], [refute_0_105, refute_0_106])).

cnf(refute_0_108, plain, (1 < 0 * skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_107])).

cnf(refute_0_109, plain, (1 < 0 * skoXC1),
    inference(resolve, [], [refute_0_108, refute_0_5])).

cnf(refute_0_110, plain, ($false),
    inference(arithmetic, [], [refute_0_109])).
% SZS output end CNFRefutation for sin-recip-2.tptp
