% SZS output start CNFRefutation for sin-3426-2a.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

cnf(sin_upper_bound_0_neg, axiom,
    (0 < X | ~ lgen(R, X - X ^ 3 / 6, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_neg, axiom,
    (0 < X |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

fof(sin_problem_3426, conjecture,
    (! [X] :
       ((X != 0 & abs(X) < pi / 2) =>
        1 / sin(X) ^ 2 - 1/5 <=
        1 / (X + 2 * pi) ^ 2 + 1 / (X + pi) ^ 2 + 1 / X ^ 2 +
        1 / (X - pi) ^ 2 + 1 / (X - 2 * pi) ^ 2))).

fof(subgoal_0, plain,
    (! [X] :
       ((X != 0 & abs(X) < pi / 2) =>
        1 / sin(X) ^ 2 - 1/5 <=
        1 / (X + 2 * pi) ^ 2 + 1 / (X + pi) ^ 2 + 1 / X ^ 2 +
        1 / (X - pi) ^ 2 + 1 / (X - 2 * pi) ^ 2)),
    inference(strip, [], [sin_problem_3426])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((X != 0 & abs(X) < pi / 2) =>
          1 / sin(X) ^ 2 - 1/5 <=
          1 / (X + 2 * pi) ^ 2 + 1 / (X + pi) ^ 2 + 1 / X ^ 2 +
          1 / (X - pi) ^ 2 + 1 / (X - 2 * pi) ^ 2)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (1 / (X + 2 * pi) ^ 2 + 1 / (X + pi) ^ 2 + 1 / X ^ 2 +
        1 / (X - pi) ^ 2 + 1 / (X - 2 * pi) ^ 2 < 1 / sin(X) ^ 2 - 1/5 &
        abs(X) < pi / 2 & X != 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (1 / (skoXC1 + 2 * pi) ^ 2 + 1 / (skoXC1 + pi) ^ 2 + 1 / skoXC1 ^ 2 +
     1 / (skoXC1 - pi) ^ 2 + 1 / (skoXC1 - 2 * pi) ^ 2 <
     1 / sin(skoXC1) ^ 2 - 1/5 & abs(skoXC1) < pi / 2 & skoXC1 != 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1 / (skoXC1 + 2 * pi) ^ 2 + 1 / (skoXC1 + pi) ^ 2 + 1 / skoXC1 ^ 2 +
     1 / (skoXC1 - pi) ^ 2 + 1 / (skoXC1 - 2 * pi) ^ 2 <
     1 / sin(skoXC1) ^ 2 - 1/5), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (abs(skoXC1) < pi / 2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_6, plain, (skoXC1 != 0),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) |
     0 <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
             skoXC1 *
             (skoXC1 *
              (19/30 + pi * (pi * -1/3) +
               skoXC1 * (skoXC1 * 1/30))))))))))) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     0 <=
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <=
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000057, sin(X_000056)) | X_000057 <= sin(X_000056)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (X_000056 < 0 | ~ lgen(0, X_000057, X_000056 - X_000056 ^ 3 / 6) |
     lgen(0, X_000057, sin(X_000056))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_4, plain,
    (X_000056 < 0 | ~ lgen(0, X_000057, X_000056 - X_000056 ^ 3 / 6) |
     X_000057 <= sin(X_000056)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000056 * (1 + X_000056 * (X_000056 * -1/6)) < X_000057 |
     X_000056 < 0 | X_000057 <= sin(X_000056)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <=
     sin(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (-1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) <= -1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_8, plain,
    ((pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) *
     (sin(skoXC1) * sin(skoXC1)) < 1 | sin(skoXC1) * sin(skoXC1) <= 0 |
     1 / (sin(skoXC1) * sin(skoXC1)) <=
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_9, plain,
    (1 / (skoXC1 + 2 * pi) ^ 2 + 1 / (skoXC1 + pi) ^ 2 + 1 / skoXC1 ^ 2 +
     1 / (skoXC1 - pi) ^ 2 + 1 / (skoXC1 - 2 * pi) ^ 2 <
     1 / sin(skoXC1) ^ 2 - 1/5),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    ((pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) <
     1 / (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    ((pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) *
     (sin(skoXC1) * sin(skoXC1)) < 1 | sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_8, refute_0_10])).

cnf(refute_0_12, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain, (abs(skoXC1) < pi / 2),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain, (abs(skoXC1) < pi * 1/2),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1 | abs(skoXC1) = -skoXC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_16, plain,
    (-skoXC1 < pi * 1/2 | abs(skoXC1) != -skoXC1 |
     pi * 1/2 <= abs(skoXC1)), introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (-skoXC1 < pi * 1/2 | 0 <= skoXC1 | pi * 1/2 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain, (-skoXC1 < pi * 1/2 | 0 <= skoXC1),
    inference(resolve, [], [refute_0_17, refute_0_14])).

cnf(refute_0_19, plain, (pi * -1/2 < skoXC1 | 0 <= skoXC1),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 < 0 | abs(skoXC1) = skoXC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_21, plain,
    (skoXC1 < pi * 1/2 | abs(skoXC1) != skoXC1 | pi * 1/2 <= abs(skoXC1)),
    introduced(tautology, [equality])).

cnf(refute_0_22, plain,
    (skoXC1 < 0 | skoXC1 < pi * 1/2 | pi * 1/2 <= abs(skoXC1)),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain, (skoXC1 < 0 | skoXC1 < pi * 1/2),
    inference(resolve, [], [refute_0_22, refute_0_14])).

cnf(refute_0_24, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_25, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_26, plain, (skoXC1 != 0),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_27, plain, (skoXC1 < pi * 1/2 | 0 <= skoXC1),
    inference(decision, [], [refute_0_24, refute_0_25, refute_0_26])).

cnf(refute_0_28, plain, (skoXC1 < pi * 1/2),
    inference(resolve, [], [refute_0_27, refute_0_23])).

cnf(refute_0_29, plain, (skoXC1 < 0 | pi * -1/2 < skoXC1),
    inference(decision, [],
              [refute_0_28, refute_0_24, refute_0_25, refute_0_26])).

cnf(refute_0_30, plain, (pi * -1/2 < skoXC1),
    inference(resolve, [], [refute_0_19, refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) !=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_32, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_12, refute_0_31])).

cnf(refute_0_33, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) != 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_34, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) != 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_36, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain,
    (skoXC1 * (pi * -4 + skoXC1) != pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_38, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 * (pi * -2 + skoXC1) != pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_40, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_38, refute_0_39])).

cnf(refute_0_41, plain,
    (skoXC1 * (pi * 2 + skoXC1) != pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_42, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (skoXC1 * (pi * 4 + skoXC1) != pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_44, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0 | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain, (skoXC1 * skoXC1 != 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_46, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     sin(skoXC1) * sin(skoXC1) <= 0),
    inference(resolve, [], [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     1 / (sin(skoXC1) * sin(skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     1 / (sin(skoXC1) * sin(skoXC1)) <=
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_48, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     1 / (sin(skoXC1) * sin(skoXC1)) *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_47, refute_0_10])).

cnf(refute_0_49, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) = 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * 4))))) +
       skoXC1 *
       (pi * (pi * (pi * (pi * (pi * -4)))) +
        skoXC1 *
        (pi * (pi * (pi * (pi * -7))) +
         skoXC1 *
         (pi * (pi * (pi * 8)) +
          skoXC1 * (pi * (pi * 2) + skoXC1 * (pi * -4 + skoXC1))))))) !=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_51, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_53, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * pi)) +
       skoXC1 * (skoXC1 * (pi * (pi * -2) + skoXC1 * skoXC1)))) = 0 |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 * (skoXC1 * (pi * pi + skoXC1 * (pi * -2 + skoXC1))) = 0 |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_53, refute_0_33])).

cnf(refute_0_55, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 * (pi * -4 + skoXC1) = pi * (pi * -4) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_54, refute_0_35])).

cnf(refute_0_56, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 * (pi * -2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_55, refute_0_37])).

cnf(refute_0_57, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 * (pi * 2 + skoXC1) = pi * (pi * -1) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_56, refute_0_39])).

cnf(refute_0_58, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) |
     skoXC1 * (pi * 4 + skoXC1) = pi * (pi * -4) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_57, refute_0_41])).

cnf(refute_0_59, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) | skoXC1 * skoXC1 = 0),
    inference(resolve, [], [refute_0_58, refute_0_43])).

cnf(refute_0_60, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1))),
    inference(resolve, [], [refute_0_59, refute_0_45])).

cnf(refute_0_61, plain,
    (0 < sin(skoXC1) * sin(skoXC1) | sin(skoXC1) * sin(skoXC1) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_62, plain,
    (0 < sin(skoXC1) * sin(skoXC1) | 1 <= 0 | sin(skoXC1) = 0),
    inference(split, [], [refute_0_61])).

cnf(refute_0_63, plain, (0 < sin(skoXC1) * sin(skoXC1) | sin(skoXC1) = 0),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * 0) | sin(skoXC1) != 0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))),
    introduced(tautology, [equality])).

cnf(refute_0_65, plain,
    (0 < sin(skoXC1) * sin(skoXC1) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * 0) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * sin(skoXC1)) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (0 < sin(skoXC1) * sin(skoXC1) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) /
     (sin(skoXC1) * 0)),
    inference(resolve, [], [refute_0_65, refute_0_60])).

cnf(refute_0_67, plain,
    (0 < sin(skoXC1) * sin(skoXC1) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * -33/5))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * -33/5))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16)))))))),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_69, plain, (0 < sin(skoXC1) * sin(skoXC1)),
    inference(resolve, [], [refute_0_68, refute_0_67])).

cnf(refute_0_70, plain,
    (-1 <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))))),
    inference(resolve, [], [refute_0_46, refute_0_69])).

cnf(refute_0_71, plain,
    (-1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (pi * (pi * 33))) +
            skoXC1 *
            (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))) <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0), inference(resolve, [], [refute_0_7, refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1)))))))) <=
     0), inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * 33))) +
           skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))))),
    inference(decision, [],
              [refute_0_26, refute_0_25, refute_0_24, refute_0_28,
               refute_0_30])).

cnf(refute_0_74, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain,
    (sin(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     0 <=
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))),
    inference(split, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (sin(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (X_000057 < sin(X_000056) | ~ lgen(1, X_000057, sin(X_000056))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_78, plain,
    (X_000056 < 0 | ~ lgen(1, X_000057, X_000056 - X_000056 ^ 3 / 6) |
     lgen(1, X_000057, sin(X_000056))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_79, plain,
    (X_000056 < 0 | X_000057 < sin(X_000056) |
     ~ lgen(1, X_000057, X_000056 - X_000056 ^ 3 / 6)),
    inference(resolve, [], [refute_0_78, refute_0_77])).

cnf(refute_0_80, plain,
    (X_000056 < 0 | X_000057 < sin(X_000056) |
     X_000056 * (1 + X_000056 * (X_000056 * -1/6)) <= X_000057),
    inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <= 0),
    inference(subst, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_83, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <= 0 |
     sin(skoXC1) <=
     0 /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))))),
    inference(split, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) |
     sin(skoXC1) <= 0), inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))),
    inference(decision, [],
              [refute_0_26, refute_0_25, refute_0_24, refute_0_28,
               refute_0_30])).

cnf(refute_0_86, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) |
     sin(skoXC1) <= 0),
    inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (skoXC1 < 0 |
     0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <= 0),
    inference(resolve, [], [refute_0_86, refute_0_81])).

cnf(refute_0_88, plain,
    (~ lgen(0, sin(X_000077), X_000078) | sin(X_000077) <= X_000078),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_89, plain,
    (0 < X_000077 |
     ~ lgen(0,
            X_000077 - X_000077 ^ 3 / 6 + X_000077 ^ 5 / 120 -
            X_000077 ^ 7 / 5040 + X_000077 ^ 9 / 362880 -
            X_000077 ^ 11 / 39916800, X_000078) |
     lgen(0, sin(X_000077), X_000078)),
    inference(subst, [], [sin_upper_bound_2_neg])).

cnf(refute_0_90, plain,
    (0 < X_000077 |
     ~ lgen(0,
            X_000077 - X_000077 ^ 3 / 6 + X_000077 ^ 5 / 120 -
            X_000077 ^ 7 / 5040 + X_000077 ^ 9 / 362880 -
            X_000077 ^ 11 / 39916800, X_000078) |
     sin(X_000077) <= X_000078),
    inference(resolve, [], [refute_0_89, refute_0_88])).

cnf(refute_0_91, plain,
    (0 < X_000077 |
     X_000078 <
     X_000077 *
     (1 +
      X_000077 *
      (X_000077 *
       (-1/6 +
        X_000077 *
        (X_000077 *
         (1/120 +
          X_000077 *
          (X_000077 *
           (-1/5040 +
            X_000077 *
            (X_000077 *
             (1/362880 + X_000077 * (X_000077 * -1/39916800)))))))))) |
     sin(X_000077) <= X_000078), inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 | sin(skoXC1) <= 0), inference(subst, [], [refute_0_91])).

cnf(refute_0_93, plain, (0 < 0 | sin(skoXC1) < 0 | 0 < sin(skoXC1)),
    inference(split, [], [refute_0_69])).

cnf(refute_0_94, plain, (sin(skoXC1) < 0 | 0 < sin(skoXC1)),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <=
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_96, plain,
    (sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) |
     0 <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
             skoXC1 *
             (skoXC1 *
              (19/30 + pi * (pi * -1/3) +
               skoXC1 * (skoXC1 * 1/30))))))))))) |
     sin(skoXC1) <=
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30))))))))))))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_97, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <= 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <=
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_98, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <
     sin(skoXC1) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <= 0),
    inference(split, [], [refute_0_74])).

cnf(refute_0_99, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_100, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) |
     0 <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * -33/5))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))) |
     0 /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <=
     sin(skoXC1)), inference(split, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) |
     0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * 33/5))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) |
     0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_101, refute_0_85])).

cnf(refute_0_103, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <
     sin(skoXC1) | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_98, refute_0_102])).

cnf(refute_0_104, plain,
    (~ lgen(0, sin(X_000049), X_000050) | sin(X_000049) <= X_000050),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_105, plain,
    (0 < X_000049 | ~ lgen(0, X_000049 - X_000049 ^ 3 / 6, X_000050) |
     lgen(0, sin(X_000049), X_000050)),
    inference(subst, [], [sin_upper_bound_0_neg])).

cnf(refute_0_106, plain,
    (0 < X_000049 | ~ lgen(0, X_000049 - X_000049 ^ 3 / 6, X_000050) |
     sin(X_000049) <= X_000050),
    inference(resolve, [], [refute_0_105, refute_0_104])).

cnf(refute_0_107, plain,
    (0 < X_000049 |
     X_000050 < X_000049 * (1 + X_000049 * (X_000049 * -1/6)) |
     sin(X_000049) <= X_000050),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) | 0 < skoXC1 |
     sin(skoXC1) <=
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(subst, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) | 0 < skoXC1 |
     0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_108, refute_0_103])).

cnf(refute_0_110, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <= 0),
    inference(resolve, [], [refute_0_97, refute_0_109])).

cnf(refute_0_111, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30))))))))))))) | 0 < skoXC1 |
     0 <= sin(skoXC1) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30))))))))))))) | 0 < skoXC1 |
     0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_111, refute_0_102])).

cnf(refute_0_113, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30))))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <= 0),
    inference(split, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <= 0),
    inference(resolve, [], [refute_0_114, refute_0_26])).

cnf(refute_0_116, plain,
    (0 <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_117, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) |
     0 < skoXC1 |
     0 <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
             skoXC1 *
             (skoXC1 *
              (19/30 + pi * (pi * -1/3) +
               skoXC1 * (skoXC1 * 1/30))))))))))) | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_96, refute_0_117])).

cnf(refute_0_119, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) | 0 < skoXC1 |
     0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(arithmetic, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) | 0 < skoXC1),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_121, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) | 0 < skoXC1 |
     0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_119, refute_0_120])).

cnf(refute_0_122, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <= 0),
    inference(split, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) < sin(skoXC1) |
     0 < skoXC1 | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_123, refute_0_120])).

cnf(refute_0_125, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 |
     sin(skoXC1) <=
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30))))))))))))),
    inference(subst, [], [refute_0_91])).

cnf(refute_0_126, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1 | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_125, refute_0_124])).

cnf(refute_0_127, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) < 0 | skoXC1 < 0 |
     0 <= sin(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_128, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/6)) |
     0 <= sin(skoXC1)), inference(arithmetic, [], [refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 < 0 | skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/6)) <= 0),
    inference(decision, [],
              [refute_0_30, refute_0_28, refute_0_24, refute_0_25,
               refute_0_26])).

cnf(refute_0_130, plain, (skoXC1 < 0 | 0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_129, refute_0_128])).

cnf(refute_0_131, plain, (0 <= sin(skoXC1) | skoXC1 <= 0),
    inference(decision, [],
              [refute_0_130, refute_0_30, refute_0_28, refute_0_24,
               refute_0_25, refute_0_26])).

cnf(refute_0_132, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_131, refute_0_126])).

cnf(refute_0_133, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * -40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * 33))) +
          skoXC1 * (skoXC1 * (pi * (pi * -10) + skoXC1 * skoXC1))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) | 0 <= sin(skoXC1) |
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <= 0),
    inference(resolve, [], [refute_0_95, refute_0_132])).

cnf(refute_0_134, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (40 + pi * (pi * -32/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (2 + pi * (pi * (-8 + pi * (pi * -22/45))))))) +
          skoXC1 *
          (skoXC1 *
           (pi *
            (pi *
             (-10 +
              pi *
              (pi *
               (-76/15 + pi * (pi * (8/3 + pi * (pi * 142/1575))))))) +
            skoXC1 *
            (skoXC1 *
             (4 +
              pi *
              (pi *
               (14/3 +
                pi *
                (pi *
                 (-337/360 +
                  pi * (pi * (-13/45 + pi * (pi * -73/16200))))))) +
              skoXC1 *
              (skoXC1 *
               (-22/15 +
                pi *
                (pi *
                 (-1/18 +
                  pi *
                  (pi *
                   (329/1800 +
                    pi * (pi * (4/315 + pi * (pi * 2/18711))))))) +
                skoXC1 *
                (skoXC1 *
                 (41/360 +
                  pi *
                  (pi *
                   (-17/420 +
                    pi *
                    (pi *
                     (-16733/1814400 +
                      pi *
                      (pi * (-13/45360 + pi * (pi * -37/24948000))))))) +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1400 +
                    pi *
                    (pi *
                     (223/90720 +
                      pi *
                      (pi *
                       (6571/29937600 +
                        pi *
                        (pi * (29/7484400 + pi * (pi * 1/74844000))))))) +
                    skoXC1 *
                    (skoXC1 *
                     (-251/1814400 +
                      pi *
                      (pi *
                       (-53/855360 +
                        pi *
                        (pi *
                         (-3653/1197504000 + pi * (pi * -1/29937600))))) +
                      skoXC1 *
                      (skoXC1 *
                       (71/14968800 +
                        pi *
                        (pi * (53/59875200 + pi * (pi * 1/36288000))) +
                        skoXC1 *
                        (skoXC1 *
                         (-13/171072000 + pi * (pi * -1/119750400) +
                          skoXC1 *
                          (skoXC1 * 1/1197504000)))))))))))))))))))))) |
     0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (40 + pi * (pi * -32/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (2 + pi * (pi * (-8 + pi * (pi * -22/45))))))) +
          skoXC1 *
          (skoXC1 *
           (pi *
            (pi *
             (-10 +
              pi *
              (pi *
               (-76/15 + pi * (pi * (8/3 + pi * (pi * 142/1575))))))) +
            skoXC1 *
            (skoXC1 *
             (4 +
              pi *
              (pi *
               (14/3 +
                pi *
                (pi *
                 (-337/360 +
                  pi * (pi * (-13/45 + pi * (pi * -73/16200))))))) +
              skoXC1 *
              (skoXC1 *
               (-22/15 +
                pi *
                (pi *
                 (-1/18 +
                  pi *
                  (pi *
                   (329/1800 +
                    pi * (pi * (4/315 + pi * (pi * 2/18711))))))) +
                skoXC1 *
                (skoXC1 *
                 (41/360 +
                  pi *
                  (pi *
                   (-17/420 +
                    pi *
                    (pi *
                     (-16733/1814400 +
                      pi *
                      (pi * (-13/45360 + pi * (pi * -37/24948000))))))) +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1400 +
                    pi *
                    (pi *
                     (223/90720 +
                      pi *
                      (pi *
                       (6571/29937600 +
                        pi *
                        (pi * (29/7484400 + pi * (pi * 1/74844000))))))) +
                    skoXC1 *
                    (skoXC1 *
                     (-251/1814400 +
                      pi *
                      (pi *
                       (-53/855360 +
                        pi *
                        (pi *
                         (-3653/1197504000 + pi * (pi * -1/29937600))))) +
                      skoXC1 *
                      (skoXC1 *
                       (71/14968800 +
                        pi *
                        (pi * (53/59875200 + pi * (pi * 1/36288000))) +
                        skoXC1 *
                        (skoXC1 *
                         (-13/171072000 + pi * (pi * -1/119750400) +
                          skoXC1 *
                          (skoXC1 * 1/1197504000)))))))))))))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(decision, [],
              [refute_0_130, refute_0_30, refute_0_28, refute_0_24,
               refute_0_25, refute_0_26])).

cnf(refute_0_136, plain,
    (0 <= sin(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(resolve, [], [refute_0_135, refute_0_134])).

cnf(refute_0_137, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) | 0 <= sin(skoXC1)),
    inference(decision, [],
              [refute_0_130, refute_0_30, refute_0_28, refute_0_24,
               refute_0_25, refute_0_26])).

cnf(refute_0_138, plain, (0 <= sin(skoXC1)),
    inference(resolve, [], [refute_0_136, refute_0_137])).

cnf(refute_0_139, plain, (0 < sin(skoXC1)),
    inference(resolve, [], [refute_0_138, refute_0_94])).

cnf(refute_0_140, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_92, refute_0_139])).

cnf(refute_0_141, plain,
    (0 < skoXC1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <= 0),
    inference(decision, [],
              [refute_0_26, refute_0_25, refute_0_24, refute_0_28,
               refute_0_30])).

cnf(refute_0_142, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_141, refute_0_140])).

cnf(refute_0_143, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(decision, [],
              [refute_0_142, refute_0_26, refute_0_25, refute_0_24,
               refute_0_28, refute_0_30])).

cnf(refute_0_144, plain,
    (skoXC1 < 0 |
     0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))))),
    inference(resolve, [], [refute_0_87, refute_0_143])).

cnf(refute_0_145, plain, (0 <= skoXC1),
    inference(decision, [],
              [refute_0_142, refute_0_26, refute_0_25, refute_0_24,
               refute_0_28, refute_0_30])).

cnf(refute_0_146, plain,
    (0 <
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5))))))))))),
    inference(resolve, [], [refute_0_145, refute_0_144])).

cnf(refute_0_147, plain,
    (sin(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(resolve, [], [refute_0_76, refute_0_146])).

cnf(refute_0_148, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(resolve, [], [refute_0_6, refute_0_147])).

cnf(refute_0_149, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5)))))))))))),
    inference(resolve, [], [refute_0_145, refute_0_148])).

cnf(refute_0_150, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (sin(skoXC1) *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * -33/5))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))) |
     0 <=
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (-35 + pi * (pi * 8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * -33/5))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * 2) + skoXC1 * (skoXC1 * -1/5))))))))))),
    inference(resolve, [], [refute_0_1, refute_0_149])).

cnf(refute_0_151, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30))))))))))))) |
     sin(skoXC1) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16/5))))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * (35 + pi * (pi * -8))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * 33/5))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * -2) + skoXC1 * (skoXC1 * 1/5)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_150])).

cnf(refute_0_152, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) <
     sin(skoXC1) *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30)))))))))))))),
    inference(resolve, [], [refute_0_151, refute_0_146])).

cnf(refute_0_153, plain,
    (sin(skoXC1) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (pi * (pi * -33))) +
           skoXC1 *
           (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1))))))))) /
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
       skoXC1 *
       (skoXC1 *
        (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi *
           (pi *
            (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
           skoXC1 *
           (skoXC1 *
            (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
             skoXC1 *
             (skoXC1 *
              (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
               skoXC1 *
               (skoXC1 *
                (19/30 + pi * (pi * -1/3) +
                 skoXC1 * (skoXC1 * 1/30))))))))))))) |
     0 <=
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30))))))))))))),
    inference(split, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (sin(skoXC1) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (sin(skoXC1) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30)))))))))))) <= 0),
    inference(resolve, [], [refute_0_154, refute_0_26])).

cnf(refute_0_156, plain,
    (0 <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi *
           (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
            skoXC1 *
            (skoXC1 *
             (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
              skoXC1 *
              (skoXC1 *
               (-19/30 + pi * (pi * 1/3) +
                skoXC1 * (skoXC1 * -1/30))))))))))))),
    inference(decision, [],
              [refute_0_142, refute_0_26, refute_0_25, refute_0_24,
               refute_0_28, refute_0_30])).

cnf(refute_0_157, plain,
    (sin(skoXC1) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30))))))))))))),
    inference(resolve, [], [refute_0_155, refute_0_156])).

cnf(refute_0_158, plain,
    (~ lgen(0, X_000085, sin(X_000084)) | X_000085 <= sin(X_000084)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_159, plain,
    (X_000084 < 0 |
     ~ lgen(0, X_000085,
            X_000084 - X_000084 ^ 3 / 6 + X_000084 ^ 5 / 120 -
            X_000084 ^ 7 / 5040 + X_000084 ^ 9 / 362880 -
            X_000084 ^ 11 / 39916800) | lgen(0, X_000085, sin(X_000084))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_160, plain,
    (X_000084 < 0 |
     ~ lgen(0, X_000085,
            X_000084 - X_000084 ^ 3 / 6 + X_000084 ^ 5 / 120 -
            X_000084 ^ 7 / 5040 + X_000084 ^ 9 / 362880 -
            X_000084 ^ 11 / 39916800) | X_000085 <= sin(X_000084)),
    inference(resolve, [], [refute_0_159, refute_0_158])).

cnf(refute_0_161, plain,
    (X_000084 *
     (1 +
      X_000084 *
      (X_000084 *
       (-1/6 +
        X_000084 *
        (X_000084 *
         (1/120 +
          X_000084 *
          (X_000084 *
           (-1/5040 +
            X_000084 *
            (X_000084 *
             (1/362880 + X_000084 * (X_000084 * -1/39916800)))))))))) <
     X_000085 | X_000084 < 0 | X_000085 <= sin(X_000084)),
    inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) |
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) <= sin(skoXC1)),
    inference(subst, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30))))))))))))),
    inference(resolve, [], [refute_0_162, refute_0_157])).

cnf(refute_0_164, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) /
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30))))))))))))),
    inference(resolve, [], [refute_0_145, refute_0_163])).

cnf(refute_0_165, plain,
    (skoXC1 *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * 40))))) +
        skoXC1 *
        (skoXC1 *
         (pi * (pi * (pi * (pi * -33))) +
          skoXC1 *
          (skoXC1 * (pi * (pi * 10) + skoXC1 * (skoXC1 * -1)))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
      skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
          skoXC1 *
          (skoXC1 *
           (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
            skoXC1 *
            (skoXC1 *
             (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
              skoXC1 *
              (skoXC1 *
               (19/30 + pi * (pi * -1/3) +
                skoXC1 * (skoXC1 * 1/30)))))))))))) |
     0 <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) +
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (-35 + pi * (pi * (8 + pi * (pi * 8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (20 + pi * (pi * (-23/30 + pi * (pi * -4/3))))) +
           skoXC1 *
           (skoXC1 *
            (-5 + pi * (pi * (-4/3 + pi * (pi * 11/10))) +
             skoXC1 *
             (skoXC1 *
              (19/30 + pi * (pi * -1/3) +
               skoXC1 * (skoXC1 * 1/30)))))))))))),
    inference(resolve, [], [refute_0_0, refute_0_164])).

cnf(refute_0_166, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (-40 + pi * (pi * 32/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-2 + pi * (pi * (8 + pi * (pi * 22/45))))))) +
          skoXC1 *
          (skoXC1 *
           (pi *
            (pi *
             (10 +
              pi *
              (pi *
               (76/15 + pi * (pi * (-8/3 + pi * (pi * -142/1575))))))) +
            skoXC1 *
            (skoXC1 *
             (-4 +
              pi *
              (pi *
               (-14/3 +
                pi *
                (pi *
                 (337/360 +
                  pi * (pi * (13/45 + pi * (pi * 73/16200))))))) +
              skoXC1 *
              (skoXC1 *
               (22/15 +
                pi *
                (pi *
                 (1/18 +
                  pi *
                  (pi *
                   (-329/1800 +
                    pi * (pi * (-4/315 + pi * (pi * -2/18711))))))) +
                skoXC1 *
                (skoXC1 *
                 (-41/360 +
                  pi *
                  (pi *
                   (17/420 +
                    pi *
                    (pi *
                     (16733/1814400 +
                      pi *
                      (pi * (13/45360 + pi * (pi * 37/24948000))))))) +
                  skoXC1 *
                  (skoXC1 *
                   (1/1400 +
                    pi *
                    (pi *
                     (-223/90720 +
                      pi *
                      (pi *
                       (-6571/29937600 +
                        pi *
                        (pi *
                         (-29/7484400 + pi * (pi * -1/74844000))))))) +
                    skoXC1 *
                    (skoXC1 *
                     (251/1814400 +
                      pi *
                      (pi *
                       (53/855360 +
                        pi *
                        (pi *
                         (3653/1197504000 + pi * (pi * 1/29937600))))) +
                      skoXC1 *
                      (skoXC1 *
                       (-71/14968800 +
                        pi *
                        (pi * (-53/59875200 + pi * (pi * -1/36288000))) +
                        skoXC1 *
                        (skoXC1 *
                         (13/171072000 + pi * (pi * 1/119750400) +
                          skoXC1 *
                          (skoXC1 * -1/1197504000)))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (pi * (pi * (pi * (pi * (pi * (pi * (-40 + pi * (pi * 32/15))))))) +
        skoXC1 *
        (skoXC1 *
         (pi *
          (pi * (pi * (pi * (-2 + pi * (pi * (8 + pi * (pi * 22/45))))))) +
          skoXC1 *
          (skoXC1 *
           (pi *
            (pi *
             (10 +
              pi *
              (pi *
               (76/15 + pi * (pi * (-8/3 + pi * (pi * -142/1575))))))) +
            skoXC1 *
            (skoXC1 *
             (-4 +
              pi *
              (pi *
               (-14/3 +
                pi *
                (pi *
                 (337/360 +
                  pi * (pi * (13/45 + pi * (pi * 73/16200))))))) +
              skoXC1 *
              (skoXC1 *
               (22/15 +
                pi *
                (pi *
                 (1/18 +
                  pi *
                  (pi *
                   (-329/1800 +
                    pi * (pi * (-4/315 + pi * (pi * -2/18711))))))) +
                skoXC1 *
                (skoXC1 *
                 (-41/360 +
                  pi *
                  (pi *
                   (17/420 +
                    pi *
                    (pi *
                     (16733/1814400 +
                      pi *
                      (pi * (13/45360 + pi * (pi * 37/24948000))))))) +
                  skoXC1 *
                  (skoXC1 *
                   (1/1400 +
                    pi *
                    (pi *
                     (-223/90720 +
                      pi *
                      (pi *
                       (-6571/29937600 +
                        pi *
                        (pi *
                         (-29/7484400 + pi * (pi * -1/74844000))))))) +
                    skoXC1 *
                    (skoXC1 *
                     (251/1814400 +
                      pi *
                      (pi *
                       (53/855360 +
                        pi *
                        (pi *
                         (3653/1197504000 + pi * (pi * 1/29937600))))) +
                      skoXC1 *
                      (skoXC1 *
                       (-71/14968800 +
                        pi *
                        (pi * (-53/59875200 + pi * (pi * -1/36288000))) +
                        skoXC1 *
                        (skoXC1 *
                         (13/171072000 + pi * (pi * 1/119750400) +
                          skoXC1 *
                          (skoXC1 * -1/1197504000)))))))))))))))))))))) <=
     0 |
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(decision, [],
              [refute_0_142, refute_0_26, refute_0_25, refute_0_24,
               refute_0_28, refute_0_30])).

cnf(refute_0_168, plain,
    (skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30))))))))))) <=
     pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16)))))))),
    inference(resolve, [], [refute_0_167, refute_0_166])).

cnf(refute_0_169, plain,
    (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * -16))))))) <
     skoXC1 *
     (skoXC1 *
      (pi * (pi * (pi * (pi * (pi * (pi * (pi * (pi * 8/15))))))) +
       skoXC1 *
       (skoXC1 *
        (pi *
         (pi * (pi * (pi * (35 + pi * (pi * (-8 + pi * (pi * -8/15))))))) +
         skoXC1 *
         (skoXC1 *
          (pi * (pi * (-20 + pi * (pi * (23/30 + pi * (pi * 4/3))))) +
           skoXC1 *
           (skoXC1 *
            (5 + pi * (pi * (4/3 + pi * (pi * -11/10))) +
             skoXC1 *
             (skoXC1 *
              (-19/30 + pi * (pi * 1/3) +
               skoXC1 * (skoXC1 * -1/30)))))))))))),
    inference(decision, [],
              [refute_0_142, refute_0_26, refute_0_25, refute_0_24,
               refute_0_28, refute_0_30])).

cnf(refute_0_170, plain, ($false),
    inference(resolve, [], [refute_0_168, refute_0_169])).
% SZS output end CNFRefutation for sin-3426-2a.tptp
