% SZS output start CNFRefutation for sqrt-1mcosq-8-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sqrt_le, axiom, (Y < 0 | ~ lgen(R, X, Y ^ 2) | lgen(R, sqrt(X), Y))).

cnf(sqrt_lower_bound_id, axiom,
    (X < 0 | 1 < X | ~ lgen(R, Y, X) | lgen(R, Y, sqrt(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof(sqrt_1mcosq_8, conjecture,
    (! [X, Y] :
       ((1/5 < X & X < Y & Y < 3/2) =>
        sqrt(1 - cos(X) ^ 2) / sqrt(1 - cos(Y) ^ 2) <= 32/5 * X / Y))).

fof(subgoal_0, plain,
    (! [X, Y] :
       ((1/5 < X & X < Y & Y < 3/2) =>
        sqrt(1 - cos(X) ^ 2) / sqrt(1 - cos(Y) ^ 2) <= 32/5 * X / Y)),
    inference(strip, [], [sqrt_1mcosq_8])).

fof(negate_0_0, plain,
    (~ ! [X, Y] :
         ((1/5 < X & X < Y & Y < 3/2) =>
          sqrt(1 - cos(X) ^ 2) / sqrt(1 - cos(Y) ^ 2) <= 32/5 * X / Y)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X, Y] :
       (1/5 < X & X < Y & Y < 3/2 &
        32/5 * X / Y < sqrt(1 - cos(X) ^ 2) / sqrt(1 - cos(Y) ^ 2))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoYC1 < 3/2 &
     32/5 * skoXC1 / skoYC1 <
     sqrt(1 - cos(skoXC1) ^ 2) / sqrt(1 - cos(skoYC1) ^ 2) & 1/5 < skoXC1 &
     skoXC1 < skoYC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (32/5 * skoXC1 / skoYC1 <
     sqrt(1 - cos(skoXC1) ^ 2) / sqrt(1 - cos(skoYC1) ^ 2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 < skoYC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoYC1 < 3/2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (1/5 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (X_000036 < cos(X_000035) | ~ lgen(1, X_000036, cos(X_000035))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (~ lgen(1, X_000036,
            1 - X_000035 ^ 2 / 2 + X_000035 ^ 4 / 24 - X_000035 ^ 6 / 720 +
            X_000035 ^ 8 / 40320 - X_000035 ^ 10 / 3628800) |
     lgen(1, X_000036, cos(X_000035))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_2, plain,
    (X_000036 < cos(X_000035) |
     ~ lgen(1, X_000036,
            1 - X_000035 ^ 2 / 2 + X_000035 ^ 4 / 24 - X_000035 ^ 6 / 720 +
            X_000035 ^ 8 / 40320 - X_000035 ^ 10 / 3628800)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000036 < cos(X_000035) |
     1 +
     X_000035 *
     (X_000035 *
      (-1/2 +
       X_000035 *
       (X_000035 *
        (1/24 +
         X_000035 *
         (X_000035 *
          (-1/720 +
           X_000035 *
           (X_000035 *
            (1/40320 + X_000035 * (X_000035 * -1/3628800))))))))) <=
     X_000036), inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (0 < cos(skoXC1) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <=
     0), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) |
     skoXC1 * (skoXC1 * 1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (~ lgen(0, cos(X_000028), X_000029) | cos(X_000028) <= X_000029),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (~ lgen(0, 1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24, X_000029) |
     lgen(0, cos(X_000028), X_000029)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_8, plain,
    (~ lgen(0, 1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24, X_000029) |
     cos(X_000028) <= X_000029),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    (X_000029 <
     1 + X_000028 * (X_000028 * (-1/2 + X_000028 * (X_000028 * 1/24))) |
     cos(X_000028) <= X_000029), inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (X_000056 / X_000057 <
     1 + X_000028 * (X_000028 * (-1/2 + X_000028 * (X_000028 * 1/24))) |
     cos(X_000028) <= X_000056 / X_000057),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (X_000056 / X_000057 < cos(X_000028) | X_000057 <= 0 |
     cos(X_000028) * X_000057 <= X_000056),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_12, plain,
    (X_000056 / X_000057 <
     1 + X_000028 * (X_000028 * (-1/2 + X_000028 * (X_000028 * 1/24))) |
     X_000057 <= 0 | cos(X_000028) * X_000057 <= X_000056),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <=
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) |
     skoXC1 * (skoXC1 * 1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) <= 0 |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_15, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) <=
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_16, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))) <= 0 |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_17, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) <=
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_18, plain,
    ((-1 + cos(skoXC1) * cos(skoXC1)) * (skoYC1 * skoYC1) <
     skoXC1 * (skoXC1 * -1024/25) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) |
     skoYC1 * skoYC1 <= 0 |
     (skoXC1 * (skoXC1 * -1024/25) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * 2048/25) +
        cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))) /
     (skoYC1 * skoYC1) <= -1 + cos(skoXC1) * cos(skoXC1)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_19, plain,
    (~ lgen(0, sqrt(X_000079), X_000080) | sqrt(X_000079) <= X_000080),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_20, plain,
    (X_000080 < 0 | ~ lgen(0, X_000079, X_000080 ^ 2) |
     lgen(0, sqrt(X_000079), X_000080)), inference(subst, [], [sqrt_le])).

cnf(refute_0_21, plain,
    (X_000080 < 0 | ~ lgen(0, X_000079, X_000080 ^ 2) |
     sqrt(X_000079) <= X_000080),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (X_000080 * X_000080 < X_000079 | X_000080 < 0 |
     sqrt(X_000079) <= X_000080),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 < 0 |
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 *
     ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
      skoYC1) < 1 + cos(skoXC1) * (cos(skoXC1) * -1) |
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) <
     (1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) |
     0 <= skoXC1 * -32/5 |
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5) <= 1 + cos(skoYC1) * (cos(skoYC1) * -1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_25, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * -1 * skoYC1 <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) |
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) /
     skoYC1 <= sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * -1 |
     skoYC1 <= 0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_26, plain,
    (32/5 * skoXC1 / skoYC1 <
     sqrt(1 - cos(skoXC1) ^ 2) / sqrt(1 - cos(skoYC1) ^ 2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_27, plain,
    (skoXC1 * 32/5 / skoYC1 <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) /
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1))),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 * 32/5 / skoYC1 * sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) |
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) /
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <= skoXC1 * 32/5 / skoYC1 |
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_29, plain,
    (skoXC1 * 32/5 / skoYC1 * sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) |
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <= 0),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * -1 <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) /
     skoYC1 | cos(skoYC1) * (cos(skoYC1) * -1) <= -1),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (cos(X_000028) < X_000029 | ~ lgen(1, cos(X_000028), X_000029)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_32, plain,
    (~ lgen(1, 1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24, X_000029) |
     lgen(1, cos(X_000028), X_000029)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_33, plain,
    (cos(X_000028) < X_000029 |
     ~ lgen(1, 1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24, X_000029)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (cos(X_000028) < X_000029 |
     X_000029 <=
     1 + X_000028 * (X_000028 * (-1/2 + X_000028 * (X_000028 * 1/24)))),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (cos(skoYC1) <
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) |
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (cos(skoYC1) < 1 / cos(skoYC1) |
     1 / cos(skoYC1) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_37, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     cos(skoYC1) * (cos(skoYC1) * -1) <= -1),
    introduced(tautology, [assume])).

cnf(refute_0_38, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) | 0 <= cos(skoYC1) * -1 |
     -1 / (cos(skoYC1) * -1) <= cos(skoYC1)),
    inference(split, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     1 / cos(skoYC1) <= cos(skoYC1) | cos(skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (0 < cos(skoYC1) |
     1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <=
     0), inference(subst, [], [refute_0_3])).

cnf(refute_0_41, plain,
    (0 < cos(skoYC1) |
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <=
     -1), inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain, (skoXC1 < skoYC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_43, plain, (skoYC1 < 3/2),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_44, plain, (1/5 < skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_45, plain,
    (-1 <
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800)))))))))),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_46, plain, (0 < cos(skoYC1)),
    inference(resolve, [], [refute_0_41, refute_0_45])).

cnf(refute_0_47, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     1 / cos(skoYC1) <= cos(skoYC1)),
    inference(resolve, [], [refute_0_39, refute_0_46])).

cnf(refute_0_48, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     1 / cos(skoYC1) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_47, refute_0_36])).

cnf(refute_0_49, plain,
    (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <
     1 / cos(skoYC1) |
     1 <=
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     cos(skoYC1) | cos(skoYC1) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_50, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     1 <=
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     cos(skoYC1) | cos(skoYC1) <= 0),
    inference(resolve, [], [refute_0_48, refute_0_49])).

cnf(refute_0_51, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     cos(skoYC1) *
     (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <= -1 |
     cos(skoYC1) <= 0), inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     cos(skoYC1) *
     (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <= -1),
    inference(resolve, [], [refute_0_51, refute_0_46])).

cnf(refute_0_53, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     0 <= -1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24))) |
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <=
     cos(skoYC1)), inference(split, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= -1 |
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <=
     cos(skoYC1)), inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (-1 < skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_56, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <=
     cos(skoYC1)), inference(resolve, [], [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_56, refute_0_35])).

cnf(refute_0_58, plain,
    (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <
     -1 / (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) |
     0 <= -1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24))) |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <= -1),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_59, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     0 <= -1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24))) |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (-1 + skoYC1 * (skoYC1 * (1/2 + skoYC1 * (skoYC1 * -1/24)))) <= -1),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= -1 |
     skoYC1 *
     (skoYC1 *
      (1 +
       skoYC1 *
       (skoYC1 *
        (-1/3 +
         skoYC1 * (skoYC1 * (1/24 + skoYC1 * (skoYC1 * -1/576))))))) <= 0),
    inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (0 <
     skoYC1 *
     (skoYC1 *
      (1 +
       skoYC1 *
       (skoYC1 *
        (-1/3 +
         skoYC1 * (skoYC1 * (1/24 + skoYC1 * (skoYC1 * -1/576))))))) |
     skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= -1),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_62, plain,
    (-1 < cos(skoYC1) * (cos(skoYC1) * -1) |
     skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= -1),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain, (-1 < cos(skoYC1) * (cos(skoYC1) * -1)),
    inference(resolve, [], [refute_0_62, refute_0_55])).

cnf(refute_0_64, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * -1 <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) /
     skoYC1), inference(resolve, [], [refute_0_30, refute_0_63])).

cnf(refute_0_65, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * -1 * skoYC1 <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) |
     skoYC1 <= 0), inference(resolve, [], [refute_0_25, refute_0_64])).

cnf(refute_0_66, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) |
     skoYC1 <= 0), inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain, (0 < skoYC1),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_68, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) <
     sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5)),
    inference(resolve, [], [refute_0_66, refute_0_67])).

cnf(refute_0_69, plain,
    (sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5) | 0 <= skoXC1 * -32/5),
    inference(split, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_69])).

cnf(refute_0_71, plain, (0 < skoXC1),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_72, plain,
    (sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5)),
    inference(resolve, [], [refute_0_70, refute_0_71])).

cnf(refute_0_73, plain,
    (~ lgen(0, X_000114, sqrt(X_000113)) | X_000114 <= sqrt(X_000113)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_74, plain,
    (1 < X_000113 | X_000113 < 0 | ~ lgen(0, X_000114, X_000113) |
     lgen(0, X_000114, sqrt(X_000113))),
    inference(subst, [], [sqrt_lower_bound_id])).

cnf(refute_0_75, plain,
    (1 < X_000113 | X_000113 < 0 | ~ lgen(0, X_000114, X_000113) |
     X_000114 <= sqrt(X_000113)),
    inference(resolve, [], [refute_0_74, refute_0_73])).

cnf(refute_0_76, plain,
    (1 < X_000113 | X_000113 < X_000114 | X_000113 < 0 |
     X_000114 <= sqrt(X_000113)),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (1 + cos(skoYC1) * (cos(skoYC1) * -1) < 0 |
     1 < 1 + cos(skoYC1) * (cos(skoYC1) * -1) |
     1 + cos(skoYC1) * (cos(skoYC1) * -1) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5) |
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5) <= sqrt(1 + cos(skoYC1) * (cos(skoYC1) * -1))),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (1 + cos(skoYC1) * (cos(skoYC1) * -1) < 0 |
     1 < 1 + cos(skoYC1) * (cos(skoYC1) * -1) |
     1 + cos(skoYC1) * (cos(skoYC1) * -1) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5)),
    inference(resolve, [], [refute_0_77, refute_0_72])).

cnf(refute_0_79, plain,
    (0 < cos(skoYC1) * (cos(skoYC1) * -1) | 1 < cos(skoYC1) * cos(skoYC1) |
     1 + cos(skoYC1) * (cos(skoYC1) * -1) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5)), inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (1 <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= 0 |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_81, plain,
    (1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) <=
     1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))))),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_82, plain,
    (1 <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     cos(skoYC1) |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     1 / cos(skoYC1) | cos(skoYC1) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_83, plain,
    (1 / cos(skoYC1) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) <= 1 / cos(skoYC1)), inference(subst, [], [refute_0_9])).

cnf(refute_0_84, plain,
    (1 < cos(skoYC1) * cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1),
    introduced(tautology, [assume])).

cnf(refute_0_85, plain,
    (1 / cos(skoYC1) < cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1 |
     cos(skoYC1) <= 0), inference(split, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (1 / cos(skoYC1) < cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_85, refute_0_46])).

cnf(refute_0_87, plain,
    (1 / cos(skoYC1) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_83, refute_0_86])).

cnf(refute_0_88, plain,
    (1 <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1 | cos(skoYC1) <= 0),
    inference(resolve, [], [refute_0_82, refute_0_87])).

cnf(refute_0_89, plain,
    (1 <
     cos(skoYC1) *
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) |
     cos(skoYC1) * cos(skoYC1) <= 1 | cos(skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (1 <
     cos(skoYC1) *
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) |
     cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_89, refute_0_46])).

cnf(refute_0_91, plain,
    (1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) <
     cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1 |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= 0),
    inference(split, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) <
     cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1 |
     skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= -1),
    inference(arithmetic, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) <
     cos(skoYC1) | cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_92, refute_0_55])).

cnf(refute_0_94, plain,
    (1 / (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_81, refute_0_93])).

cnf(refute_0_95, plain,
    (1 <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) |
     cos(skoYC1) * cos(skoYC1) <= 1 |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <= 0),
    inference(resolve, [], [refute_0_80, refute_0_94])).

cnf(refute_0_96, plain,
    (0 < 1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) <= 0), inference(subst, [], [refute_0_9])).

cnf(refute_0_97, plain,
    (0 < 1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_96, refute_0_46])).

cnf(refute_0_98, plain,
    (1 <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) |
     cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_95, refute_0_97])).

cnf(refute_0_99, plain,
    (0 <
     skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (1/3 +
         skoYC1 * (skoYC1 * (-1/24 + skoYC1 * (skoYC1 * 1/576))))))) |
     cos(skoYC1) * cos(skoYC1) <= 1),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (1/3 +
         skoYC1 * (skoYC1 * (-1/24 + skoYC1 * (skoYC1 * 1/576))))))) <= 0),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_101, plain, (cos(skoYC1) * cos(skoYC1) <= 1),
    inference(resolve, [], [refute_0_100, refute_0_99])).

cnf(refute_0_102, plain,
    (0 < cos(skoYC1) * (cos(skoYC1) * -1) |
     1 + cos(skoYC1) * (cos(skoYC1) * -1) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5)),
    inference(resolve, [], [refute_0_101, refute_0_79])).

cnf(refute_0_103, plain,
    (0 < cos(skoYC1) * (cos(skoYC1) * -1) |
     cos(skoYC1) * (cos(skoYC1) * -1) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_104, plain, (0 < -1 | cos(skoYC1) * (cos(skoYC1) * -1) <= 0),
    inference(split, [], [refute_0_103])).

cnf(refute_0_105, plain, (cos(skoYC1) * (cos(skoYC1) * -1) <= 0),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (1 + cos(skoYC1) * (cos(skoYC1) * -1) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) /
     (skoXC1 * -32/5)),
    inference(resolve, [], [refute_0_105, refute_0_102])).

cnf(refute_0_107, plain,
    (sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * (skoYC1 * -1) <
     (1 + cos(skoYC1) * (cos(skoYC1) * -1)) * (skoXC1 * -32/5) |
     0 <= skoXC1 * -32/5),
    inference(resolve, [], [refute_0_24, refute_0_106])).

cnf(refute_0_108, plain,
    (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * skoYC1 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5)) <
     sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) * skoYC1),
    inference(resolve, [], [refute_0_108, refute_0_71])).

cnf(refute_0_110, plain,
    ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 < sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1)) | skoYC1 <= 0),
    inference(split, [], [refute_0_109])).

cnf(refute_0_111, plain,
    ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 < sqrt(1 + cos(skoXC1) * (cos(skoXC1) * -1))),
    inference(resolve, [], [refute_0_110, refute_0_67])).

cnf(refute_0_112, plain,
    ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 < 0 |
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 *
     ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
      skoYC1) < 1 + cos(skoXC1) * (cos(skoXC1) * -1)),
    inference(resolve, [], [refute_0_23, refute_0_111])).

cnf(refute_0_113, plain,
    ((skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 < 0 |
     -1 + cos(skoXC1) * cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * 2048/25) +
        cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))) /
     (skoYC1 * skoYC1)), inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (cos(skoYC1) <= 0 | cos(skoYC1) <= 1 / cos(skoYC1)),
    inference(split, [], [refute_0_101])).

cnf(refute_0_115, plain, (cos(skoYC1) <= 1 / cos(skoYC1)),
    inference(resolve, [], [refute_0_114, refute_0_46])).

cnf(refute_0_116, plain,
    (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5)) <
     0 * skoYC1 |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 | skoYC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_117, plain,
    (skoXC1 * 32/5 < cos(skoYC1) * (cos(skoYC1) * (skoXC1 * 32/5)) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 | skoYC1 <= 0), inference(arithmetic, [], [refute_0_116])).

cnf(refute_0_118, plain,
    (skoXC1 * 32/5 < cos(skoYC1) * (cos(skoYC1) * (skoXC1 * 32/5)) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(resolve, [], [refute_0_117, refute_0_67])).

cnf(refute_0_119, plain,
    (skoXC1 * 32/5 / (cos(skoYC1) * (skoXC1 * 32/5)) < cos(skoYC1) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 | cos(skoYC1) * (skoXC1 * 32/5) <= 0),
    inference(split, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (1 / cos(skoYC1) < cos(skoYC1) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 | cos(skoYC1) * (skoXC1 * 32/5) <= 0 | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain, (skoXC1 != 0),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_122, plain,
    (1 / cos(skoYC1) < cos(skoYC1) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1 | cos(skoYC1) * (skoXC1 * 32/5) <= 0),
    inference(resolve, [], [refute_0_120, refute_0_121])).

cnf(refute_0_123, plain,
    (cos(skoYC1) * (skoXC1 * 32/5) < 0 | skoXC1 * 32/5 < 0 |
     0 < cos(skoYC1) * (skoXC1 * 32/5) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(split, [], [refute_0_118])).

cnf(refute_0_124, plain,
    (skoXC1 < 0 | 0 < cos(skoYC1) * (skoXC1 * -32/5) |
     0 < cos(skoYC1) * (skoXC1 * 32/5) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(arithmetic, [], [refute_0_123])).

cnf(refute_0_125, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_126, plain,
    (0 < cos(skoYC1) * (skoXC1 * -32/5) |
     0 < cos(skoYC1) * (skoXC1 * 32/5) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(resolve, [], [refute_0_125, refute_0_124])).

cnf(refute_0_127, plain,
    (~ lgen(0, X_000036, cos(X_000035)) | X_000036 <= cos(X_000035)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_128, plain,
    (~ lgen(0, X_000036,
            1 - X_000035 ^ 2 / 2 + X_000035 ^ 4 / 24 - X_000035 ^ 6 / 720 +
            X_000035 ^ 8 / 40320 - X_000035 ^ 10 / 3628800) |
     lgen(0, X_000036, cos(X_000035))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_129, plain,
    (~ lgen(0, X_000036,
            1 - X_000035 ^ 2 / 2 + X_000035 ^ 4 / 24 - X_000035 ^ 6 / 720 +
            X_000035 ^ 8 / 40320 - X_000035 ^ 10 / 3628800) |
     X_000036 <= cos(X_000035)),
    inference(resolve, [], [refute_0_128, refute_0_127])).

cnf(refute_0_130, plain,
    (1 +
     X_000035 *
     (X_000035 *
      (-1/2 +
       X_000035 *
       (X_000035 *
        (1/24 +
         X_000035 *
         (X_000035 *
          (-1/720 +
           X_000035 *
           (X_000035 *
            (1/40320 + X_000035 * (X_000035 * -1/3628800))))))))) <
     X_000036 | X_000036 <= cos(X_000035)),
    inference(arithmetic, [], [refute_0_129])).

cnf(refute_0_131, plain,
    (1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <
     0 | 0 <= cos(skoYC1)), inference(subst, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (0 < cos(skoYC1) * (skoXC1 * -32/5) |
     cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_133, plain,
    (cos(skoYC1) < 0 / (skoXC1 * -32/5) | 0 <= skoXC1 * -32/5 |
     cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    inference(split, [], [refute_0_132])).

cnf(refute_0_134, plain,
    (cos(skoYC1) < 0 | cos(skoYC1) * (skoXC1 * -32/5) <= 0 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (cos(skoYC1) < 0 | cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    inference(resolve, [], [refute_0_134, refute_0_71])).

cnf(refute_0_136, plain,
    (1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <
     0 | cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    inference(resolve, [], [refute_0_131, refute_0_135])).

cnf(refute_0_137, plain,
    (1 <
     skoYC1 *
     (skoYC1 *
      (1/2 +
       skoYC1 *
       (skoYC1 *
        (-1/24 +
         skoYC1 *
         (skoYC1 *
          (1/720 +
           skoYC1 *
           (skoYC1 * (-1/40320 + skoYC1 * (skoYC1 * 1/3628800))))))))) |
     cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain,
    (skoYC1 *
     (skoYC1 *
      (1/2 +
       skoYC1 *
       (skoYC1 *
        (-1/24 +
         skoYC1 *
         (skoYC1 *
          (1/720 +
           skoYC1 *
           (skoYC1 * (-1/40320 + skoYC1 * (skoYC1 * 1/3628800))))))))) <=
     1), inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_139, plain, (cos(skoYC1) * (skoXC1 * -32/5) <= 0),
    inference(resolve, [], [refute_0_138, refute_0_137])).

cnf(refute_0_140, plain,
    (0 < cos(skoYC1) * (skoXC1 * 32/5) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(resolve, [], [refute_0_139, refute_0_126])).

cnf(refute_0_141, plain,
    (1 / cos(skoYC1) < cos(skoYC1) |
     0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(resolve, [], [refute_0_122, refute_0_140])).

cnf(refute_0_142, plain,
    (0 <=
     (skoXC1 * 32/5 + cos(skoYC1) * (cos(skoYC1) * (skoXC1 * -32/5))) /
     skoYC1), inference(resolve, [], [refute_0_115, refute_0_141])).

cnf(refute_0_143, plain,
    (-1 + cos(skoXC1) * cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * 2048/25) +
        cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))) /
     (skoYC1 * skoYC1)),
    inference(resolve, [], [refute_0_142, refute_0_113])).

cnf(refute_0_144, plain,
    ((-1 + cos(skoXC1) * cos(skoXC1)) * (skoYC1 * skoYC1) <
     skoXC1 * (skoXC1 * -1024/25) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) |
     skoYC1 * skoYC1 <= 0),
    inference(resolve, [], [refute_0_18, refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) |
     skoYC1 * skoYC1 <= 0), inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain, (0 < skoYC1 * skoYC1),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_147, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))))),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) <
     cos(skoYC1) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))) <= 0),
    inference(split, [], [refute_0_147])).

cnf(refute_0_149, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))) < 0 |
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) < 0 |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))),
    inference(split, [], [refute_0_147])).

cnf(refute_0_150, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))),
    inference(arithmetic, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (X_000022 < cos(X_000021) | ~ lgen(1, X_000022, cos(X_000021))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_152, plain,
    (~ lgen(1, X_000022, 1 - X_000021 ^ 2 / 2) |
     lgen(1, X_000022, cos(X_000021))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_153, plain,
    (X_000022 < cos(X_000021) | ~ lgen(1, X_000022, 1 - X_000021 ^ 2 / 2)),
    inference(resolve, [], [refute_0_152, refute_0_151])).

cnf(refute_0_154, plain,
    (X_000022 < cos(X_000021) |
     1 + X_000021 * (X_000021 * -1/2) <= X_000022),
    inference(arithmetic, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (0 < cos(skoXC1) | 1 + skoXC1 * (skoXC1 * -1/2) <= 0),
    inference(subst, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_157, plain,
    (~ lgen(0, X_000022, cos(X_000021)) | X_000022 <= cos(X_000021)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_158, plain,
    (~ lgen(0, X_000022, 1 - X_000021 ^ 2 / 2) |
     lgen(0, X_000022, cos(X_000021))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_159, plain,
    (~ lgen(0, X_000022, 1 - X_000021 ^ 2 / 2) |
     X_000022 <= cos(X_000021)),
    inference(resolve, [], [refute_0_158, refute_0_157])).

cnf(refute_0_160, plain,
    (1 + X_000021 * (X_000021 * -1/2) < X_000022 |
     X_000022 <= cos(X_000021)),
    inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_163, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_164, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    introduced(tautology, [assume])).

cnf(refute_0_165, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(split, [], [refute_0_164])).

cnf(refute_0_166, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_165])).

cnf(refute_0_167, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_163, refute_0_166])).

cnf(refute_0_168, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_162, refute_0_167])).

cnf(refute_0_169, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     cos(skoXC1) * (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(split, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_170])).

cnf(refute_0_172, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_161, refute_0_171])).

cnf(refute_0_173, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_174, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) <=
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))))),
    inference(subst, [], [refute_0_9])).

cnf(refute_0_175, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) < cos(skoXC1) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(split, [], [refute_0_169])).

cnf(refute_0_176, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(resolve, [], [refute_0_174, refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(resolve, [], [refute_0_173, refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 *
       (skoXC1 *
        (1 + skoXC1 * (skoXC1 * (-7/24 + skoXC1 * (skoXC1 * 1/48))))))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(arithmetic, [], [refute_0_177])).

cnf(refute_0_179, plain,
    (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) < 0 |
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) < 0 |
     0 < skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(split, [], [refute_0_169])).

cnf(refute_0_180, plain,
    (0 < skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     0 < skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) |
     skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (0 < skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     0 < skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) |
     skoYC1 * skoYC1 <= skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_182, plain,
    (0 < skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     0 < skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_181, refute_0_180])).

cnf(refute_0_183, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 *
       (skoXC1 *
        (1 + skoXC1 * (skoXC1 * (-7/24 + skoXC1 * (skoXC1 * 1/48))))))) <=
     skoXC1 * (skoXC1 * 1024/25) |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(decision, [],
              [refute_0_182, refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_184, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) <= 0),
    inference(resolve, [], [refute_0_183, refute_0_178])).

cnf(refute_0_185, plain,
    (0 < skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(decision, [],
              [refute_0_184, refute_0_182, refute_0_44, refute_0_43,
               refute_0_42])).

cnf(refute_0_186, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_172, refute_0_185])).

cnf(refute_0_187, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_156, refute_0_186])).

cnf(refute_0_188, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 * (skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/4))))) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 *
     (skoYC1 * (skoXC1 * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/4))))) <=
     skoXC1 * (skoXC1 * 1024/25) |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_184, refute_0_182, refute_0_44, refute_0_43,
               refute_0_42])).

cnf(refute_0_190, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_189, refute_0_188])).

cnf(refute_0_191, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_190, refute_0_185])).

cnf(refute_0_192, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     skoYC1 * skoYC1 <= 0 | cos(skoXC1) <= 0 / (skoYC1 * skoYC1)),
    inference(split, [], [refute_0_191])).

cnf(refute_0_193, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     skoYC1 * skoYC1 <= 0 | cos(skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     0 | 0 <= cos(skoXC1)), inference(subst, [], [refute_0_130])).

cnf(refute_0_195, plain,
    (cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) < 0 |
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) < 0 |
     0 < cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(split, [], [refute_0_164])).

cnf(refute_0_196, plain,
    (0 < cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     0 < cos(skoXC1) * (skoYC1 * skoYC1) |
     skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(arithmetic, [], [refute_0_195])).

cnf(refute_0_197, plain,
    (0 < cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(resolve, [], [refute_0_191, refute_0_196])).

cnf(refute_0_198, plain,
    (skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) < 0 / (skoYC1 * (skoYC1 * -1)) |
     0 <= skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(split, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (cos(skoXC1) < 0 | skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     skoYC1 * skoYC1 <= 0), inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    (0 < skoYC1 * skoYC1 | skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_201, plain,
    (cos(skoXC1) < 0 | skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(resolve, [], [refute_0_199, refute_0_200])).

cnf(refute_0_202, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <
     0 | skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(resolve, [], [refute_0_194, refute_0_201])).

cnf(refute_0_203, plain,
    (1 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) |
     skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 * (-1/40320 + skoXC1 * (skoXC1 * 1/3628800))))))))) <=
     1), inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_205, plain,
    (skoXC1 * (skoXC1 * 1024/25) < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(resolve, [], [refute_0_204, refute_0_203])).

cnf(refute_0_206, plain,
    (0 < skoYC1 * skoYC1 |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(decision, [],
              [refute_0_205, refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_207, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     cos(skoXC1) <= 0),
    inference(resolve, [], [refute_0_193, refute_0_206])).

cnf(refute_0_208, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     1 + skoXC1 * (skoXC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_207, refute_0_155])).

cnf(refute_0_209, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) |
     skoXC1 * (skoXC1 * -1/2) <= -1),
    inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) |
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(decision, [],
              [refute_0_205, refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_211, plain,
    (cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1)),
    inference(resolve, [], [refute_0_209, refute_0_210])).

cnf(refute_0_212, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))),
    inference(resolve, [], [refute_0_211, refute_0_150])).

cnf(refute_0_213, plain,
    (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) |
     0 <=
     -1024/25 +
     skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75))) |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     -2048/25), inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_214, plain,
    (cos(skoYC1) <
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) |
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_215, plain,
    (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <
     -2048/25 / (cos(skoYC1) * -1024/25) | 0 <= cos(skoYC1) * -1024/25 |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) * -1024/25) <= -2048/25),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_216, plain,
    (cos(skoYC1) < -2048/25 / (cos(skoYC1) * -1024/25) |
     -2048/25 / (cos(skoYC1) * -1024/25) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_217, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_218, plain,
    (cos(skoYC1) <
     0 /
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) |
     0 <=
     skoXC1 * (skoXC1 * -2048/25) +
     cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(split, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (cos(skoYC1) < 0 |
     cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))) <=
     skoXC1 * (skoXC1 * -2048/25) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (1 <
     skoYC1 *
     (skoYC1 *
      (1/2 +
       skoYC1 *
       (skoYC1 *
        (-1/24 +
         skoYC1 *
         (skoYC1 *
          (1/720 +
           skoYC1 *
           (skoYC1 * (-1/40320 + skoYC1 * (skoYC1 * 1/3628800))))))))) |
     0 <= cos(skoYC1)), inference(arithmetic, [], [refute_0_131])).

cnf(refute_0_221, plain, (0 <= cos(skoYC1)),
    inference(resolve, [], [refute_0_138, refute_0_220])).

cnf(refute_0_222, plain,
    (cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))) <=
     skoXC1 * (skoXC1 * -2048/25) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_221, refute_0_219])).

cnf(refute_0_223, plain,
    (0 <= cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     skoXC1 * (skoXC1 * -2048/25) /
     (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))) <= cos(skoYC1)),
    inference(split, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)) <= 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 / (cos(skoYC1) * -1024/25) <= cos(skoYC1) | skoXC1 = 0),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)) <= 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 / (cos(skoYC1) * -1024/25) <= cos(skoYC1)),
    inference(resolve, [], [refute_0_224, refute_0_121])).

cnf(refute_0_226, plain,
    (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) < 0 |
     0 < cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) |
     0 <= skoXC1 * (skoXC1 * -2048/25) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(split, [], [refute_0_222])).

cnf(refute_0_227, plain,
    (0 < cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) |
     0 < cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     skoXC1 * (skoXC1 * 2048/25) <= 0),
    inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain, (0 < skoXC1 * (skoXC1 * 2048/25)),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_229, plain,
    (0 < cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) |
     0 < cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_227, refute_0_228])).

cnf(refute_0_230, plain,
    (0 < cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) |
     cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_231, plain,
    (cos(skoYC1) < 0 / (skoXC1 * (skoXC1 * -1024/25)) |
     0 <= skoXC1 * (skoXC1 * -1024/25) |
     cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    inference(split, [], [refute_0_230])).

cnf(refute_0_232, plain,
    (cos(skoYC1) < 0 | cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0 |
     skoXC1 * (skoXC1 * 1024/25) <= 0),
    inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain, (0 < skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_234, plain,
    (cos(skoYC1) < 0 | cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    inference(resolve, [], [refute_0_232, refute_0_233])).

cnf(refute_0_235, plain,
    (1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <
     0 | cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    inference(resolve, [], [refute_0_131, refute_0_234])).

cnf(refute_0_236, plain,
    (1 <
     skoYC1 *
     (skoYC1 *
      (1/2 +
       skoYC1 *
       (skoYC1 *
        (-1/24 +
         skoYC1 *
         (skoYC1 *
          (1/720 +
           skoYC1 *
           (skoYC1 * (-1/40320 + skoYC1 * (skoYC1 * 1/3628800))))))))) |
     cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    inference(arithmetic, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)) <= 0),
    inference(resolve, [], [refute_0_138, refute_0_236])).

cnf(refute_0_238, plain,
    (0 < cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_237, refute_0_229])).

cnf(refute_0_239, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 / (cos(skoYC1) * -1024/25) <= cos(skoYC1)),
    inference(resolve, [], [refute_0_225, refute_0_238])).

cnf(refute_0_240, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 / (cos(skoYC1) * -1024/25) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_239, refute_0_216])).

cnf(refute_0_241, plain,
    (0 <= cos(skoYC1) * -1024/25 |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) * -1024/25) <= -2048/25 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_240, refute_0_215])).

cnf(refute_0_242, plain,
    (cos(skoYC1) *
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     -2048/25 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     cos(skoYC1) <= 0), inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain,
    (cos(skoYC1) *
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     -2048/25 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_242, refute_0_46])).

cnf(refute_0_244, plain,
    (0 <=
     -1024/25 +
     skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     cos(skoYC1)), inference(split, [], [refute_0_243])).

cnf(refute_0_245, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     skoYC1 * (skoYC1 * (-512/25 + skoYC1 * (skoYC1 * 128/75))) <=
     -1024/25 |
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     cos(skoYC1)), inference(arithmetic, [], [refute_0_244])).

cnf(refute_0_246, plain,
    (-1024/25 <
     skoYC1 * (skoYC1 * (-512/25 + skoYC1 * (skoYC1 * 128/75)))),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_247, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     cos(skoYC1)), inference(resolve, [], [refute_0_245, refute_0_246])).

cnf(refute_0_248, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     -2048/25 /
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_247, refute_0_214])).

cnf(refute_0_249, plain,
    (0 <=
     -1024/25 +
     skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75))) |
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (-1024/25 +
      skoYC1 * (skoYC1 * (512/25 + skoYC1 * (skoYC1 * -128/75)))) <=
     -2048/25 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_248, refute_0_213])).

cnf(refute_0_250, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     skoYC1 * (skoYC1 * (-512/25 + skoYC1 * (skoYC1 * 128/75))) <=
     -1024/25 |
     skoYC1 *
     (skoYC1 *
      (1024/25 +
       skoYC1 *
       (skoYC1 *
        (-1024/75 +
         skoYC1 * (skoYC1 * (128/75 + skoYC1 * (skoYC1 * -16/225))))))) <=
     -1024/25), inference(arithmetic, [], [refute_0_249])).

cnf(refute_0_251, plain,
    (-1024/25 <
     skoYC1 *
     (skoYC1 *
      (1024/25 +
       skoYC1 *
       (skoYC1 *
        (-1024/75 +
         skoYC1 * (skoYC1 * (128/75 + skoYC1 * (skoYC1 * -16/225))))))) |
     skoYC1 * (skoYC1 * (-512/25 + skoYC1 * (skoYC1 * 128/75))) <=
     -1024/25),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_252, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0 |
     skoYC1 * (skoYC1 * (-512/25 + skoYC1 * (skoYC1 * 128/75))) <=
     -1024/25), inference(resolve, [], [refute_0_250, refute_0_251])).

cnf(refute_0_253, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * 1024/25)))) <= 0),
    inference(resolve, [], [refute_0_252, refute_0_246])).

cnf(refute_0_254, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))),
    inference(resolve, [], [refute_0_253, refute_0_212])).

cnf(refute_0_255, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) <
     cos(skoYC1)), inference(resolve, [], [refute_0_148, refute_0_254])).

cnf(refute_0_256, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))),
    inference(resolve, [], [refute_0_17, refute_0_255])).

cnf(refute_0_257, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 2048/25) +
       cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_256])).

cnf(refute_0_258, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      cos(skoYC1) * (cos(skoYC1) * (skoXC1 * (skoXC1 * -1024/25)))) <= 0),
    inference(arithmetic, [], [refute_0_257])).

cnf(refute_0_259, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))))),
    inference(resolve, [], [refute_0_258, refute_0_254])).

cnf(refute_0_260, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) < 0 |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) < 0 |
     0 <
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))),
    inference(split, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))))),
    inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))))),
    inference(resolve, [], [refute_0_211, refute_0_261])).

cnf(refute_0_263, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_264, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_265, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <=
     1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_266, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_267, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * -1/2)) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     0 <=
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225)))))))))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) <=
     1 + skoYC1 * (skoYC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_268, plain,
    (1 + skoYC1 * (skoYC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) <=
     cos(skoYC1)), inference(subst, [], [refute_0_160])).

cnf(refute_0_269, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_270, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) <=
     skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_271, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <= 0),
    inference(resolve, [], [refute_0_270, refute_0_259])).

cnf(refute_0_272, plain,
    ((skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75)))))),
    inference(arithmetic, [], [refute_0_271])).

cnf(refute_0_273, plain,
    (skoXC1 * (skoXC1 * 1024/25) + skoYC1 * (skoYC1 * -1) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
      cos(skoYC1) *
      (cos(skoYC1) *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 512/25) +
          skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) +
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <= 0),
    inference(resolve, [], [refute_0_269, refute_0_272])).

cnf(refute_0_274, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75)))))),
    inference(arithmetic, [], [refute_0_273])).

cnf(refute_0_275, plain,
    (cos(skoYC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     0 <=
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225)))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75)))))),
    inference(split, [], [refute_0_274])).

cnf(refute_0_276, plain,
    (cos(skoYC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_275])).

cnf(refute_0_277, plain,
    (1 + skoYC1 * (skoYC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_268, refute_0_276])).

cnf(refute_0_278, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * -1/2)) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/25) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1024/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 128/75) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225))))))))))) |
     0 <=
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -16/225)))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_267, refute_0_277])).

cnf(refute_0_279, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) < 0 |
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) < 0 |
     0 <
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(split, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) |
     skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_280])).

cnf(refute_0_282, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) |
     skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(split, [], [refute_0_281])).

cnf(refute_0_283, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) |
     skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_282])).

cnf(refute_0_284, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) |
     skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_266, refute_0_283])).

cnf(refute_0_285, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (cos(skoXC1) * -1) |
     0 <= cos(skoXC1) * -1 |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) <= 1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_286, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_287, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) <=
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_288, plain,
    ((skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) |
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_289, plain,
    ((skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <= 0),
    inference(resolve, [], [refute_0_288, refute_0_279])).

cnf(refute_0_290, plain,
    ((skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_289])).

cnf(refute_0_291, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25) |
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <= 0),
    inference(resolve, [], [refute_0_287, refute_0_290])).

cnf(refute_0_292, plain,
    (skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2944/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1856/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 1456/225) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -64/75) +
               skoYC1 *
               (skoYC1 *
                (skoXC1 * (skoXC1 * 38/675) +
                 skoYC1 *
                 (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))))) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_291])).

cnf(refute_0_293, plain,
    (cos(skoXC1) <
     skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2944/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1856/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 1456/225) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -64/75) +
               skoYC1 *
               (skoYC1 *
                (skoXC1 * (skoXC1 * 38/675) +
                 skoYC1 *
                 (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)), inference(split, [], [refute_0_292])).

cnf(refute_0_294, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25) | skoYC1 = 0),
    inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain,
    (skoYC1 != 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_296, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_294, refute_0_295])).

cnf(refute_0_297, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_286, refute_0_296])).

cnf(refute_0_298, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (cos(skoXC1) * -1) |
     0 <= cos(skoXC1) * -1 | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_285, refute_0_297])).

cnf(refute_0_299, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     cos(skoXC1) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25) | cos(skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (0 < cos(skoXC1) |
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <=
     -1), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_301, plain,
    (-1 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800)))))))))),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_302, plain, (0 < cos(skoXC1)),
    inference(resolve, [], [refute_0_300, refute_0_301])).

cnf(refute_0_303, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     cos(skoXC1) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_299, refute_0_302])).

cnf(refute_0_304, plain,
    (-1 + skoXC1 * (skoXC1 * 1/2) < 0 |
     -1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     0 | 0 < -1 + skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)), inference(split, [], [refute_0_303])).

cnf(refute_0_305, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | 1 < skoXC1 * (skoXC1 * 1/2) |
     -1 <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_304])).

cnf(refute_0_306, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | 1 < skoXC1 * (skoXC1 * 1/2) |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) <=
     -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_307, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | 1 < skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_306, refute_0_305])).

cnf(refute_0_308, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | 1 < skoXC1 * (skoXC1 * 1/2) |
     skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_309, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | 1 < skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_307, refute_0_308])).

cnf(refute_0_310, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (-1 + skoXC1 * (skoXC1 * 1/2)) | -1 + skoXC1 * (skoXC1 * 1/2) <= 0 |
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) <=
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_311, plain,
    ((-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) * (-1 + skoXC1 * (skoXC1 * 1/2)) <=
     -1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) |
     -1 + skoXC1 * (skoXC1 * 1/2) <= 0),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_312, plain,
    ((-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25) | -1 + skoXC1 * (skoXC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_311, refute_0_303])).

cnf(refute_0_313, plain,
    ((-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_312])).

cnf(refute_0_314, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     skoXC1 * (skoXC1 * 1/2) <= 1),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_315, plain,
    ((-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <
     1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1),
    inference(resolve, [], [refute_0_313, refute_0_314])).

cnf(refute_0_316, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/24)))) *
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1 | -1 + skoXC1 * (skoXC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_310, refute_0_315])).

cnf(refute_0_317, plain,
    (skoXC1 *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (7/24 + skoXC1 * (skoXC1 * -1/48))))) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1),
    inference(arithmetic, [], [refute_0_316])).

cnf(refute_0_318, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) <=
     skoXC1 *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (7/24 + skoXC1 * (skoXC1 * -1/48)))))),
    inference(decision, [],
              [refute_0_309, refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_319, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * 1/2) <= 1),
    inference(resolve, [], [refute_0_318, refute_0_317])).

cnf(refute_0_320, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <= 1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_321, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_322, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)), inference(split, [], [refute_0_303])).

cnf(refute_0_323, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_322])).

cnf(refute_0_324, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_319, refute_0_309])).

cnf(refute_0_325, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_323, refute_0_324])).

cnf(refute_0_326, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2944/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1856/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1456/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -64/75) +
              skoYC1 *
              (skoYC1 *
               (skoXC1 * (skoXC1 * 38/675) +
                skoYC1 *
                (skoYC1 * (skoXC1 * (skoXC1 * -1/675)))))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_321, refute_0_325])).

cnf(refute_0_327, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * 38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -1/675))))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_320, refute_0_326])).

cnf(refute_0_328, plain,
    (skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/4))) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(arithmetic, [], [refute_0_327])).

cnf(refute_0_329, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2944/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1856/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1456/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 64/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -38/675) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 1/675))))))))))))) <=
     skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/4))) |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [],
              [refute_0_319, refute_0_309, refute_0_44, refute_0_43,
               refute_0_42])).

cnf(refute_0_330, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_329, refute_0_328])).

cnf(refute_0_331, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [],
              [refute_0_319, refute_0_309, refute_0_44, refute_0_43,
               refute_0_42])).

cnf(refute_0_332, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(resolve, [], [refute_0_330, refute_0_331])).

cnf(refute_0_333, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <= 1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_334, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_335, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (cos(skoXC1) * -1) |
     0 <= cos(skoXC1) * -1 |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) <= 1 + skoXC1 * (skoXC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_336, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) |
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) <= cos(skoXC1)),
    inference(subst, [], [refute_0_160])).

cnf(refute_0_337, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * -1/2)) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     0 <=
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <=
     1 + skoYC1 * (skoYC1 * -1/2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_338, plain,
    (1 + skoYC1 * (skoYC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) <=
     cos(skoYC1)), inference(subst, [], [refute_0_160])).

cnf(refute_0_339, plain,
    (cos(skoYC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     0 <=
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(split, [], [refute_0_279])).

cnf(refute_0_340, plain,
    (cos(skoYC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_339])).

cnf(refute_0_341, plain,
    (1 + skoYC1 * (skoYC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
      cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1))) /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_338, refute_0_340])).

cnf(refute_0_342, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) +
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * skoYC1)) <
     (1 + skoYC1 * (skoYC1 * -1/2)) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 1536/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -512/15) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 128/15) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -208/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225)))))))))))) |
     0 <=
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_337, refute_0_341])).

cnf(refute_0_343, plain,
    (skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2816/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1664/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 1136/225) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -112/225) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))))) <
     cos(skoXC1) * (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_342])).

cnf(refute_0_344, plain,
    (cos(skoXC1) <
     skoYC1 *
     (skoYC1 *
      (-1 +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2816/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1664/75) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 1136/225) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -112/225) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)), inference(split, [], [refute_0_343])).

cnf(refute_0_345, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) | skoYC1 = 0),
    inference(arithmetic, [], [refute_0_344])).

cnf(refute_0_346, plain,
    (skoYC1 != 0 | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(decision, [],
              [refute_0_332, refute_0_319, refute_0_309, refute_0_44,
               refute_0_43, refute_0_42])).

cnf(refute_0_347, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_345, refute_0_346])).

cnf(refute_0_348, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (cos(skoXC1) * -1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_336, refute_0_347])).

cnf(refute_0_349, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (cos(skoXC1) * -1) |
     0 <= cos(skoXC1) * -1 | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_335, refute_0_348])).

cnf(refute_0_350, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     cos(skoXC1) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) | cos(skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_349])).

cnf(refute_0_351, plain,
    (0 < cos(skoXC1) | skoXC1 * (skoXC1 * -1/2) <= -1),
    inference(arithmetic, [], [refute_0_155])).

cnf(refute_0_352, plain,
    (0 < cos(skoXC1) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_351, refute_0_324])).

cnf(refute_0_353, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     cos(skoXC1) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_350, refute_0_352])).

cnf(refute_0_354, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)), inference(split, [], [refute_0_353])).

cnf(refute_0_355, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_354])).

cnf(refute_0_356, plain,
    (cos(skoXC1) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_355, refute_0_324])).

cnf(refute_0_357, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (-1 +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 2816/75) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -1664/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 1136/225) +
            skoYC1 *
            (skoYC1 *
             (skoXC1 * (skoXC1 * -112/225) +
              skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225)))))))))))) /
     (-1 + skoXC1 * (skoXC1 * 1/2)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_334, refute_0_356])).

cnf(refute_0_358, plain,
    (-1 +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 4/225))))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) * (-1 + skoXC1 * (skoXC1 * 1/2)) |
     0 <= -1 + skoXC1 * (skoXC1 * 1/2) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_333, refute_0_357])).

cnf(refute_0_359, plain,
    (skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/4))) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -4/225))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -2816/75) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1664/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -1136/225) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 112/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -4/225))))))))))) <=
     skoXC1 * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/4)))),
    inference(decision, [],
              [refute_0_332, refute_0_319, refute_0_309, refute_0_44,
               refute_0_43, refute_0_42])).

cnf(refute_0_361, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoXC1 * (skoXC1 * -1/2) <= -1 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_360, refute_0_359])).

cnf(refute_0_362, plain,
    (-1 < skoXC1 * (skoXC1 * -1/2) | cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(decision, [],
              [refute_0_332, refute_0_319, refute_0_309, refute_0_44,
               refute_0_43, refute_0_42])).

cnf(refute_0_363, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_361, refute_0_362])).

cnf(refute_0_364, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -8/225))))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(decision, [],
              [refute_0_42, refute_0_43, refute_0_44, refute_0_309,
               refute_0_319, refute_0_332, refute_0_363])).

cnf(refute_0_365, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1536/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/15) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 128/15) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * -208/225) +
             skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 8/225))))))))))) |
     1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_364, refute_0_284])).

cnf(refute_0_366, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_332, refute_0_365])).

cnf(refute_0_367, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (cos(skoXC1) * (skoYC1 * (skoYC1 * -1))) |
     0 <= cos(skoXC1) * (skoYC1 * (skoYC1 * -1)) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_265, refute_0_366])).

cnf(refute_0_368, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     cos(skoXC1) * (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(arithmetic, [], [refute_0_367])).

cnf(refute_0_369, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(split, [], [refute_0_368])).

cnf(refute_0_370, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (0 < skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(decision, [],
              [refute_0_363, refute_0_332, refute_0_319, refute_0_309,
               refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_372, plain,
    (cos(skoXC1) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_370, refute_0_371])).

cnf(refute_0_373, plain,
    (1 + skoXC1 * (skoXC1 * -1/2) <
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (-1 + skoXC1 * (skoXC1 * 2048/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * -2048/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * 256/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225)))))))))) /
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_264, refute_0_372])).

cnf(refute_0_374, plain,
    (skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (-1 + skoXC1 * (skoXC1 * 2048/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * 256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -32/225))))))))) <
     (1 + skoXC1 * (skoXC1 * -1/2)) *
     (skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2)))) |
     0 <= skoYC1 * (skoYC1 * (-1 + skoXC1 * (skoXC1 * 1/2))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_263, refute_0_373])).

cnf(refute_0_375, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * (-2023/25 + skoXC1 * (skoXC1 * -1/4))) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 32/225))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * (-2023/25 + skoXC1 * (skoXC1 * -1/4))) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2048/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -256/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 32/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(decision, [],
              [refute_0_363, refute_0_332, refute_0_319, refute_0_309,
               refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_377, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0 |
     skoYC1 * (skoYC1 * (1 + skoXC1 * (skoXC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_376, refute_0_375])).

cnf(refute_0_378, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <= 0),
    inference(resolve, [], [refute_0_377, refute_0_371])).

cnf(refute_0_379, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     skoXC1 * (skoXC1 * 1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <= 0 |
     cos(skoYC1) <=
     0 /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))))),
    inference(split, [], [refute_0_378])).

cnf(refute_0_380, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) | cos(skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_379])).

cnf(refute_0_381, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))),
    inference(decision, [], [refute_0_42, refute_0_43, refute_0_44])).

cnf(refute_0_382, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoYC1) <= 0),
    inference(resolve, [], [refute_0_380, refute_0_381])).

cnf(refute_0_383, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <=
     0), inference(resolve, [], [refute_0_382, refute_0_40])).

cnf(refute_0_384, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <=
     -1), inference(arithmetic, [], [refute_0_383])).

cnf(refute_0_385, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))))) <=
     skoXC1 * (skoXC1 * -2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75)))))),
    inference(resolve, [], [refute_0_384, refute_0_45])).

cnf(refute_0_386, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     cos(skoYC1) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_385, refute_0_262])).

cnf(refute_0_387, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) <= 0),
    inference(resolve, [], [refute_0_15, refute_0_386])).

cnf(refute_0_388, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) < 0 |
     skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) < 0 |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(split, [], [refute_0_386])).

cnf(refute_0_389, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) |
     skoXC1 * (skoXC1 * 2048/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(arithmetic, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -256/75))))) <=
     skoXC1 * (skoXC1 * 2048/25)),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_391, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) |
     0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_390, refute_0_389])).

cnf(refute_0_392, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_393, plain,
    (cos(skoYC1) <
     0 /
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) |
     0 <=
     skoXC1 * (skoXC1 * -1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * 512/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    inference(split, [], [refute_0_392])).

cnf(refute_0_394, plain,
    (cos(skoYC1) < 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -512/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_393])).

cnf(refute_0_395, plain,
    (skoXC1 * (skoXC1 * -1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -512/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_396, plain,
    (cos(skoYC1) < 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    inference(resolve, [], [refute_0_394, refute_0_395])).

cnf(refute_0_397, plain,
    (1 +
     skoYC1 *
     (skoYC1 *
      (-1/2 +
       skoYC1 *
       (skoYC1 *
        (1/24 +
         skoYC1 *
         (skoYC1 *
          (-1/720 +
           skoYC1 *
           (skoYC1 * (1/40320 + skoYC1 * (skoYC1 * -1/3628800))))))))) <
     0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    inference(resolve, [], [refute_0_131, refute_0_396])).

cnf(refute_0_398, plain,
    (1 <
     skoYC1 *
     (skoYC1 *
      (1/2 +
       skoYC1 *
       (skoYC1 *
        (-1/24 +
         skoYC1 *
         (skoYC1 *
          (1/720 +
           skoYC1 *
           (skoYC1 * (-1/40320 + skoYC1 * (skoYC1 * 1/3628800))))))))) |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    inference(arithmetic, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (cos(skoYC1) *
     (skoXC1 * (skoXC1 * -1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * 512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * -128/75)))))) <= 0),
    inference(resolve, [], [refute_0_138, refute_0_398])).

cnf(refute_0_400, plain,
    (0 <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_399, refute_0_391])).

cnf(refute_0_401, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_387, refute_0_400])).

cnf(refute_0_402, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (cos(skoYC1) *
      (skoXC1 * (skoXC1 * 1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * -512/25) +
         skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) <= 0),
    inference(resolve, [], [refute_0_14, refute_0_401])).

cnf(refute_0_403, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -512/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 128/75)))))) <= 0),
    inference(arithmetic, [], [refute_0_402])).
SZS status Timeout for sqrt-1mcosq-8-weak.tptp
Wall time limit exceeded.

cnf(refute_0_404, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     cos(skoYC1) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_403, refute_0_400])).

cnf(refute_0_405, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoXC1 * (skoXC1 * 1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <= 0),
    inference(resolve, [], [refute_0_13, refute_0_404])).

cnf(refute_0_406, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_405])).

cnf(refute_0_407, plain,
    ((skoXC1 * (skoXC1 * 2048/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75)))))) /
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) <
     1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_406, refute_0_381])).

cnf(refute_0_408, plain,
    (skoXC1 * (skoXC1 * 2048/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 256/75))))) <
     (1 + skoYC1 * (skoYC1 * (-1/2 + skoYC1 * (skoYC1 * 1/24)))) *
     (skoXC1 * (skoXC1 * 1024/25) +
      skoYC1 *
      (skoYC1 *
       (skoXC1 * (skoXC1 * -1024/25) +
        skoYC1 *
        (skoYC1 *
         (skoXC1 * (skoXC1 * 1024/75) +
          skoYC1 *
          (skoYC1 *
           (skoXC1 * (skoXC1 * -128/75) +
            skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225)))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoXC1 * (skoXC1 * 1024/25) +
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <= 0),
    inference(resolve, [], [refute_0_5, refute_0_407])).

cnf(refute_0_409, plain,
    (skoXC1 * (skoXC1 * 1024/25) <
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -512/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2432/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -256/25) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 112/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -8/75) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 2/675))))))))))))) |
     cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(arithmetic, [], [refute_0_408])).

cnf(refute_0_410, plain,
    (skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25) |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -512/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 2432/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -256/25) +
           skoYC1 *
           (skoYC1 *
            (skoXC1 * (skoXC1 * 112/75) +
             skoYC1 *
             (skoYC1 *
              (skoXC1 * (skoXC1 * -8/75) +
               skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 2/675))))))))))))) <=
     skoXC1 * (skoXC1 * 1024/25)),
    inference(decision, [], [refute_0_44, refute_0_43, refute_0_42])).

cnf(refute_0_411, plain,
    (cos(skoXC1) * (skoYC1 * skoYC1) <= 0 |
     skoYC1 *
     (skoYC1 *
      (skoXC1 * (skoXC1 * -1024/25) +
       skoYC1 *
       (skoYC1 *
        (skoXC1 * (skoXC1 * 1024/75) +
         skoYC1 *
         (skoYC1 *
          (skoXC1 * (skoXC1 * -128/75) +
           skoYC1 * (skoYC1 * (skoXC1 * (skoXC1 * 16/225))))))))) <=
     skoXC1 * (skoXC1 * -1024/25)),
    inference(resolve, [], [refute_0_410, refute_0_409])).

cnf(refute_0_412, plain, (cos(skoXC1) * (skoYC1 * skoYC1) <= 0),
    inference(resolve, [], [refute_0_411, refute_0_381])).

cnf(refute_0_413, plain,
    (skoYC1 * skoYC1 <= 0 | cos(skoXC1) <= 0 / (skoYC1 * skoYC1)),
    inference(split, [], [refute_0_412])).

cnf(refute_0_414, plain, (skoYC1 * skoYC1 <= 0 | cos(skoXC1) <= 0),
    inference(arithmetic, [], [refute_0_413])).

cnf(refute_0_415, plain, (cos(skoXC1) <= 0),
    inference(resolve, [], [refute_0_414, refute_0_146])).

cnf(refute_0_416, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <=
     0), inference(resolve, [], [refute_0_415, refute_0_4])).

cnf(refute_0_417, plain,
    (skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 * (1/40320 + skoXC1 * (skoXC1 * -1/3628800))))))))) <=
     -1), inference(arithmetic, [], [refute_0_416])).

cnf(refute_0_418, plain, ($false),
    inference(resolve, [], [refute_0_417, refute_0_301])).
% SZS output end CNFRefutation for sqrt-1mcosq-8-weak.tptp
