% SZS output start CNFRefutation for sqrt-1mcosq-4-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(sqrt_ge, axiom, (~ lgen(R, Y ^ 2, X) | lgen(R, Y, sqrt(X)))).

cnf(cos_upper_bound_0, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_2, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600, Y) |
     lgen(R, cos(X), Y))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof(sqrt_1mcosq_4, conjecture,
    (! [X] :
       ((0 < X & X < 19/20) =>
        pi <=
        10 ^ -6 +
        sqrt(1 - cos(15707963/5000000 * X) ^ 2) / (X * (1 - X))))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 19/20) =>
        pi <=
        10 ^ -6 +
        sqrt(1 - cos(15707963/5000000 * X) ^ 2) / (X * (1 - X)))),
    inference(strip, [], [sqrt_1mcosq_4])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 19/20) =>
          pi <=
          10 ^ -6 +
          sqrt(1 - cos(15707963/5000000 * X) ^ 2) / (X * (1 - X)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 19/20 &
        10 ^ -6 + sqrt(1 - cos(15707963/5000000 * X) ^ 2) / (X * (1 - X)) <
        pi)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 19/20 &
     10 ^ -6 +
     sqrt(1 - cos(15707963/5000000 * skoXC1) ^ 2) /
     (skoXC1 * (1 - skoXC1)) < pi & 0 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (10 ^ -6 +
     sqrt(1 - cos(15707963/5000000 * skoXC1) ^ 2) /
     (skoXC1 * (1 - skoXC1)) < pi),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_4, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_5, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoXC1 < 19/20),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, cos(X_000027), X_000028) | cos(X_000027) <= X_000028),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, 1 - X_000027 ^ 2 / 2 + X_000027 ^ 4 / 24, X_000028) |
     lgen(0, cos(X_000027), X_000028)),
    inference(subst, [], [cos_upper_bound_0])).

cnf(refute_0_3, plain,
    (~ lgen(0, 1 - X_000027 ^ 2 / 2 + X_000027 ^ 4 / 24, X_000028) |
     cos(X_000027) <= X_000028),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000028 <
     1 + X_000027 * (X_000027 * (-1/2 + X_000027 * (X_000027 * 1/24))) |
     cos(X_000027) <= X_000028), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24))) |
     cos(skoXC1 * 15707963/5000000) <=
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) *
     cos(skoXC1 * 15707963/5000000) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | cos(skoXC1 * 15707963/5000000) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24))) |
     cos(skoXC1 * 15707963/5000000) <=
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000)), inference(subst, [], [refute_0_4])).

cnf(refute_0_8, plain,
    (~ lgen(0, X_000014, sqrt(X_000013)) | X_000014 <= sqrt(X_000013)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000014 ^ 2, X_000013) |
     lgen(0, X_000014, sqrt(X_000013))), inference(subst, [], [sqrt_ge])).

cnf(refute_0_10, plain,
    (~ lgen(0, X_000014 ^ 2, X_000013) | X_000014 <= sqrt(X_000013)),
    inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (X_000013 < X_000014 * X_000014 | X_000014 <= sqrt(X_000013)),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 +
     cos(skoXC1 * 15707963/5000000) *
     (cos(skoXC1 * 15707963/5000000) * -1) <
     skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1)) *
     (skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1))) |
     skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1)) <=
     sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (10 ^ -6 +
     sqrt(1 - cos(15707963/5000000 * skoXC1) ^ 2) /
     (skoXC1 * (1 - skoXC1)) < pi),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain,
    (sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) /
     (skoXC1 * (1 + skoXC1 * -1)) < -1/1000000 + pi),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) <
     (-1/1000000 + pi) * (skoXC1 * (1 + skoXC1 * -1)) |
     skoXC1 * (1 + skoXC1 * -1) <= 0 |
     -1/1000000 + pi <=
     sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) /
     (skoXC1 * (1 + skoXC1 * -1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_16, plain,
    (sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) <
     (-1/1000000 + pi) * (skoXC1 * (1 + skoXC1 * -1)) |
     skoXC1 * (1 + skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    (sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) <
     skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1)) |
     skoXC1 * (1 + skoXC1 * -1) <= 0),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_19, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_20, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_21, plain, (skoXC1 < 19/20),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_22, plain, (0 < skoXC1 * (1 + skoXC1 * -1)),
    inference(decision, [],
              [refute_0_18, refute_0_19, refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (sqrt(1 +
          cos(skoXC1 * 15707963/5000000) *
          (cos(skoXC1 * 15707963/5000000) * -1)) <
     skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1))),
    inference(resolve, [], [refute_0_17, refute_0_22])).

cnf(refute_0_24, plain,
    (1 +
     cos(skoXC1 * 15707963/5000000) *
     (cos(skoXC1 * 15707963/5000000) * -1) <
     skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1)) *
     (skoXC1 * (-1/1000000 + pi + skoXC1 * (1/1000000 + pi * -1)))),
    inference(resolve, [], [refute_0_12, refute_0_23])).

cnf(refute_0_25, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     cos(skoXC1 * 15707963/5000000) * cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) < cos(skoXC1 * 15707963/5000000) |
     cos(skoXC1 * 15707963/5000000) <= 0),
    inference(split, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) < 0 |
     cos(skoXC1 * 15707963/5000000) < 0 |
     0 < cos(skoXC1 * 15707963/5000000)),
    inference(split, [], [refute_0_25])).

cnf(refute_0_28, plain,
    (cos(skoXC1 * 15707963/5000000) < 0 |
     1 <
     skoXC1 *
     (skoXC1 *
      (1/1000000000000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 * (1/1000000000000 + pi * (-1/500000 + pi))))) |
     0 < cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (skoXC1 *
      (1/1000000000000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 * (1/1000000000000 + pi * (-1/500000 + pi))))) <= 1),
    inference(decision, [],
              [refute_0_18, refute_0_19, refute_0_20, refute_0_21])).

cnf(refute_0_30, plain,
    (cos(skoXC1 * 15707963/5000000) < 0 |
     0 < cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_32, plain,
    (~ lgen(0, X_000035, cos(X_000034)) | X_000035 <= cos(X_000034)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_33, plain,
    (~ lgen(0, X_000035,
            1 - X_000034 ^ 2 / 2 + X_000034 ^ 4 / 24 - X_000034 ^ 6 / 720 +
            X_000034 ^ 8 / 40320 - X_000034 ^ 10 / 3628800) |
     lgen(0, X_000035, cos(X_000034))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_34, plain,
    (~ lgen(0, X_000035,
            1 - X_000034 ^ 2 / 2 + X_000034 ^ 4 / 24 - X_000034 ^ 6 / 720 +
            X_000034 ^ 8 / 40320 - X_000034 ^ 10 / 3628800) |
     X_000035 <= cos(X_000034)),
    inference(resolve, [], [refute_0_33, refute_0_32])).

cnf(refute_0_35, plain,
    (1 +
     X_000034 *
     (X_000034 *
      (-1/2 +
       X_000034 *
       (X_000034 *
        (1/24 +
         X_000034 *
         (X_000034 *
          (-1/720 +
           X_000034 *
           (X_000034 *
            (1/40320 + X_000034 * (X_000034 * -1/3628800))))))))) <
     X_000035 | X_000035 <= cos(X_000034)),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 * -1/3628800))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) <=
     cos(skoXC1 * 15707963/5000000)), inference(subst, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 * -1/3628800))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <= cos(skoXC1 * 15707963/5000000)),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_38, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(split, [], [refute_0_25])).

cnf(refute_0_39, plain,
    (1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 * -1/3628800))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) *
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_42, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) *
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_41, refute_0_40])).

cnf(refute_0_43, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     cos(skoXC1 * 15707963/5000000) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(split, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 *
     (skoXC1 *
      (246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (-60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <=
     1), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 * -1/3628800))))))))) < 0 |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(subst, [], [refute_0_35])).

cnf(refute_0_47, plain,
    (1 <
     skoXC1 *
     (skoXC1 *
      (246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (-60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_45, refute_0_47])).

cnf(refute_0_49, plain,
    (1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 * -1/3628800))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_36, refute_0_48])).

cnf(refute_0_50, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_49])).

cnf(refute_0_51, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         (60880677742201738583868578161/15000000000000000000000000000 +
          skoXC1 *
          (skoXC1 *
           (-15021704612158106655647061185284945866390409/
            11250000000000000000000000000000000000000000 +
            skoXC1 *
            (skoXC1 *
             (3706456922349818181989301391855381284268604199747766141921/
              15750000000000000000000000000000000000000000000000000000000 +
              skoXC1 *
              (skoXC1 *
               -914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000)))))))))) |
     0 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              -914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_31, refute_0_50])).

cnf(refute_0_52, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742203613583868578161/1875000000000000000000000000 +
         pi * (-1/500000 + pi) +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           351562500000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             123046875000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-914531557631343243871062962086475955327466949188263064978064380157257849/
               69213867187500000000000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (6995199893102896138520698929928302436952577935703330224841955151200791635337875022805711/
                 1771875000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (skoXC1 *
                  (-55677301045227605170042772954687775880119514186102150639208366413757127997267666138279483880287235689/
                   66445312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (skoXC1 *
                    (1387520114640723498955651427310376000431870000386874894195743200505759663054752418753596386852584700811703128470594341/
                     11162812500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                     skoXC1 *
                     (skoXC1 *
                      (-3389671822490053857266275896341214725774575846747482581524470875881388402899503370784832606620812849993139588611654920013525187929/
                       279070312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                       skoXC1 *
                       (skoXC1 *
                        836367969903610889036262127505590128577855823889161882817755391107517531716617924243757516177495291973323216988607662713979673189367271072106801/
                        1255816406250000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))))))))))))))))))) |
     0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 *
     (skoXC1 *
      (246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (-60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <=
     1), inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (~ lgen(0, X_000021, cos(X_000020)) | X_000021 <= cos(X_000020)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_54, plain,
    (~ lgen(0, X_000021, 1 - X_000020 ^ 2 / 2) |
     lgen(0, X_000021, cos(X_000020))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_55, plain,
    (~ lgen(0, X_000021, 1 - X_000020 ^ 2 / 2) |
     X_000021 <= cos(X_000020)),
    inference(resolve, [], [refute_0_54, refute_0_53])).

cnf(refute_0_56, plain,
    (1 + X_000020 * (X_000020 * -1/2) < X_000021 |
     X_000021 <= cos(X_000020)), inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (1 + skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * -1/2) <
     0 | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(subst, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (1 < skoXC1 * (skoXC1 * 246740101609369/50000000000000) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_57])).

cnf(refute_0_59, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) *
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= 1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) <=
     1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_60, plain,
    (1 + skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * -1/2) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) <=
     cos(skoXC1 * 15707963/5000000)), inference(subst, [], [refute_0_56])).

cnf(refute_0_61, plain,
    (1 + skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * -1/2) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <= cos(skoXC1 * 15707963/5000000)),
    inference(subst, [], [refute_0_56])).

cnf(refute_0_62, plain,
    (1 + skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * -1/2) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_61, refute_0_38])).

cnf(refute_0_63, plain,
    (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) *
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000) |
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <=
     1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_65, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) *
     cos(skoXC1 * 15707963/5000000) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_64, refute_0_63])).

cnf(refute_0_66, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     cos(skoXC1 * 15707963/5000000) *
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= 1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(split, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 * (skoXC1 * 246740101609369/50000000000000) <= 1),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (cos(skoXC1 * 15707963/5000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_68, refute_0_58])).

cnf(refute_0_70, plain,
    (1 + skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * -1/2) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_60, refute_0_69])).

cnf(refute_0_71, plain,
    (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) *
     (1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000)) |
     0 <= 1 + skoXC1 * (skoXC1 * -246740101609369/50000000000000) |
     0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_59, refute_0_71])).

cnf(refute_0_73, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742204238583868578161/2500000000000000000000000000 +
         pi * (-1/500000 + pi))))) | 0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 * (skoXC1 * 246740101609369/50000000000000) <= 1),
    inference(arithmetic, [], [refute_0_72])).

cnf(refute_0_74, plain,
    (1 < skoXC1 * (skoXC1 * 246740101609369/50000000000000) |
     0 <
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742204238583868578161/2500000000000000000000000000 +
         pi * (-1/500000 + pi))))) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(decision, [],
              [refute_0_47, refute_0_58, refute_0_18, refute_0_19,
               refute_0_20, refute_0_21])).

cnf(refute_0_75, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742204238583868578161/2500000000000000000000000000 +
         pi * (-1/500000 + pi))))) | 0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_73, refute_0_74])).

cnf(refute_0_76, plain,
    (0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 *
     (skoXC1 *
      (246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (-60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <=
     1 |
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742203613583868578161/1875000000000000000000000000 +
         pi * (-1/500000 + pi) +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           351562500000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             123046875000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-914531557631343243871062962086475955327466949188263064978064380157257849/
               69213867187500000000000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                (6995199893102896138520698929928302436952577935703330224841955151200791635337875022805711/
                 1771875000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 (skoXC1 *
                  (-55677301045227605170042772954687775880119514186102150639208366413757127997267666138279483880287235689/
                   66445312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                   skoXC1 *
                   (skoXC1 *
                    (1387520114640723498955651427310376000431870000386874894195743200505759663054752418753596386852584700811703128470594341/
                     11162812500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                     skoXC1 *
                     (skoXC1 *
                      (-3389671822490053857266275896341214725774575846747482581524470875881388402899503370784832606620812849993139588611654920013525187929/
                       279070312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
                       skoXC1 *
                       (skoXC1 *
                        836367969903610889036262127505590128577855823889161882817755391107517531716617924243757516177495291973323216988607662713979673189367271072106801/
                        1255816406250000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))))))))))))))))))) <=
     0),
    inference(decision, [],
              [refute_0_21, refute_0_20, refute_0_19, refute_0_18,
               refute_0_58, refute_0_47, refute_0_75])).

cnf(refute_0_77, plain,
    (0 <= cos(skoXC1 * 15707963/5000000) |
     skoXC1 *
     (skoXC1 *
      (246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (-60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (-3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              914531557631343243871062962086475955327466949188263064978064380157257849/
              35437500000000000000000000000000000000000000000000000000000000000000000000))))))))) <=
     1), inference(resolve, [], [refute_0_76, refute_0_52])).

cnf(refute_0_78, plain, (0 <= cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_77, refute_0_47])).

cnf(refute_0_79, plain, (0 < cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_78, refute_0_30])).

cnf(refute_0_80, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) < cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_26, refute_0_79])).

cnf(refute_0_81, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24)))),
    inference(resolve, [], [refute_0_7, refute_0_80])).

cnf(refute_0_82, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     cos(skoXC1 * 15707963/5000000) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000)))),
    inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) *
     cos(skoXC1 * 15707963/5000000) | cos(skoXC1 * 15707963/5000000) <= 0),
    inference(resolve, [], [refute_0_6, refute_0_82])).

cnf(refute_0_84, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     cos(skoXC1 * 15707963/5000000) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) |
     cos(skoXC1 * 15707963/5000000) <= 0),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     cos(skoXC1 * 15707963/5000000) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000))))),
    inference(resolve, [], [refute_0_84, refute_0_79])).

cnf(refute_0_86, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     cos(skoXC1 * 15707963/5000000) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     0), inference(split, [], [refute_0_85])).

cnf(refute_0_87, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     cos(skoXC1 * 15707963/5000000) |
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     -1), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (0 <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24))) |
     cos(skoXC1 * 15707963/5000000) <= 0),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_89, plain,
    (0 <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24)))),
    inference(resolve, [], [refute_0_88, refute_0_79])).

cnf(refute_0_90, plain,
    (-1 <
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000)))),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     cos(skoXC1 * 15707963/5000000)),
    inference(resolve, [], [refute_0_87, refute_0_90])).

cnf(refute_0_92, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 * (skoXC1 * 15707963/5000000 * 1/24)))),
    inference(resolve, [], [refute_0_5, refute_0_91])).

cnf(refute_0_93, plain,
    ((1 +
      skoXC1 *
      (skoXC1 *
       (-1/1000000000000 + pi * (1/500000 + pi * -1) +
        skoXC1 *
        (1/500000000000 + pi * (-1/250000 + pi * 2) +
         skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1)))))) /
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000)))),
    inference(arithmetic, [], [refute_0_92])).

cnf(refute_0_94, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/1000000000000 + pi * (1/500000 + pi * -1) +
       skoXC1 *
       (1/500000000000 + pi * (-1/250000 + pi * 2) +
        skoXC1 * (-1/1000000000000 + pi * (1/500000 + pi * -1))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-246740101609369/50000000000000 +
        skoXC1 *
        (skoXC1 *
         60880677742201738583868578161/15000000000000000000000000000)))) |
     1 +
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_93])).

cnf(refute_0_95, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742203613583868578161/1875000000000000000000000000 +
         pi * (-1/500000 + pi) +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           375000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            3706456922349818181989301391855381284268604199747766141921/
            225000000000000000000000000000000000000000000000000000000))))))) |
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     -1), inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (~ lgen(0, cos(X_000041), X_000042) | cos(X_000041) <= X_000042),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_97, plain,
    (~ lgen(0,
            1 - X_000041 ^ 2 / 2 + X_000041 ^ 4 / 24 - X_000041 ^ 6 / 720 +
            X_000041 ^ 8 / 40320 - X_000041 ^ 10 / 3628800 +
            X_000041 ^ 12 / 479001600, X_000042) |
     lgen(0, cos(X_000041), X_000042)),
    inference(subst, [], [cos_upper_bound_2])).

cnf(refute_0_98, plain,
    (~ lgen(0,
            1 - X_000041 ^ 2 / 2 + X_000041 ^ 4 / 24 - X_000041 ^ 6 / 720 +
            X_000041 ^ 8 / 40320 - X_000041 ^ 10 / 3628800 +
            X_000041 ^ 12 / 479001600, X_000042) |
     cos(X_000041) <= X_000042),
    inference(resolve, [], [refute_0_97, refute_0_96])).

cnf(refute_0_99, plain,
    (X_000042 <
     1 +
     X_000041 *
     (X_000041 *
      (-1/2 +
       X_000041 *
       (X_000041 *
        (1/24 +
         X_000041 *
         (X_000041 *
          (-1/720 +
           X_000041 *
           (X_000041 *
            (1/40320 +
             X_000041 *
             (X_000041 *
              (-1/3628800 + X_000041 * (X_000041 * 1/479001600))))))))))) |
     cos(X_000041) <= X_000042), inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (0 <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 *
              (-1/3628800 +
               skoXC1 * 15707963/5000000 *
               (skoXC1 * 15707963/5000000 * 1/479001600))))))))))) |
     cos(skoXC1 * 15707963/5000000) <= 0),
    inference(subst, [], [refute_0_99])).

cnf(refute_0_101, plain,
    (0 <
     1 +
     skoXC1 * 15707963/5000000 *
     (skoXC1 * 15707963/5000000 *
      (-1/2 +
       skoXC1 * 15707963/5000000 *
       (skoXC1 * 15707963/5000000 *
        (1/24 +
         skoXC1 * 15707963/5000000 *
         (skoXC1 * 15707963/5000000 *
          (-1/720 +
           skoXC1 * 15707963/5000000 *
           (skoXC1 * 15707963/5000000 *
            (1/40320 +
             skoXC1 * 15707963/5000000 *
             (skoXC1 * 15707963/5000000 *
              (-1/3628800 +
               skoXC1 * 15707963/5000000 *
               (skoXC1 * 15707963/5000000 * 1/479001600)))))))))))),
    inference(resolve, [], [refute_0_100, refute_0_79])).

cnf(refute_0_102, plain,
    (-1 <
     skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        (60880677742201738583868578161/15000000000000000000000000000 +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           11250000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            (3706456922349818181989301391855381284268604199747766141921/
             15750000000000000000000000000000000000000000000000000000000 +
             skoXC1 *
             (skoXC1 *
              (-914531557631343243871062962086475955327466949188263064978064380157257849/
               35437500000000000000000000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (skoXC1 *
                225651609454932133500667707417042014095244449538817104027159843587122310817350807187281/
                116943750000000000000000000000000000000000000000000000000000000000000000000000000000000000)))))))))))),
    inference(arithmetic, [], [refute_0_101])).

cnf(refute_0_103, plain,
    (skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     -1 |
     skoXC1 *
     (skoXC1 *
      (-1927657043823/195312500000 + pi * (-1/500000 + pi) +
       skoXC1 *
       (-1/500000000000 + pi * (1/250000 + pi * -2) +
        skoXC1 *
        (60880677742203613583868578161/1875000000000000000000000000 +
         pi * (-1/500000 + pi) +
         skoXC1 *
         (skoXC1 *
          (-15021704612158106655647061185284945866390409/
           375000000000000000000000000000000000000000 +
           skoXC1 *
           (skoXC1 *
            3706456922349818181989301391855381284268604199747766141921/
            225000000000000000000000000000000000000000000000000000000))))))) <=
     0),
    inference(decision, [],
              [refute_0_102, refute_0_90, refute_0_18, refute_0_19,
               refute_0_20, refute_0_21])).

cnf(refute_0_104, plain,
    (skoXC1 *
     (skoXC1 *
      (-246740101609369/50000000000000 +
       skoXC1 *
       (skoXC1 *
        60880677742201738583868578161/15000000000000000000000000000))) <=
     -1), inference(resolve, [], [refute_0_103, refute_0_95])).

cnf(refute_0_105, plain, ($false),
    inference(resolve, [], [refute_0_104, refute_0_90])).
% SZS output end CNFRefutation for sqrt-1mcosq-4-weak.tptp
