% SZS output start CNFRefutation for tan-problem-1-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(cos_lower_bound_1, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_3, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_5, axiom,
    (~ lgen(R, Y,
            1 - 1/2 * X ^ 2 + 1/24 * X ^ 4 - 1/720 * X ^ 6 +
            1/40320 * X ^ 8 - 1/3628800 * X ^ 10 + 1/479001600 * X ^ 12 -
            1/87178291200 * X ^ 14 + 1/20922789888000 * X ^ 16 -
            1/6402373705728000 * X ^ 18 + 1/2432902008176640000 * X ^ 20 -
            1/1124000727777607680000 * X ^ 22) | lgen(R, Y, cos(X)))).

cnf(tan_lower_bound_pos0, axiom,
    (X < 0 | 157/100 < X | ~ lgen(R, Y, X) | lgen(R, Y, tan(X)))).

cnf(tan_lower_bound_pos1, axiom,
    (X < 0 | 157/100 < X |
     ~ lgen(R, Y, 1/3 * (X ^ 2 - 15) * X / (2 * X ^ 2 - 5)) |
     lgen(R, Y, tan(X)))).

cnf(tan_lower_bound_pos2, axiom,
    (X < 0 | 314159/200000 < X |
     ~ lgen(R, Y,
            -(21 * (X ^ 4 - 60 * X ^ 2 + 495)) * X /
            (X ^ 6 - 210 * X ^ 4 + 4725 * X ^ 2 - 10395)) |
     lgen(R, Y, tan(X)))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof(tan_problem_1, conjecture,
    (! [X] :
       ((0 <= X & X < pi / 2 - 1/80000) => 1 - cos(X) <= tan(X / 2)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X < pi / 2 - 1/80000) => 1 - cos(X) <= tan(X / 2))),
    inference(strip, [], [tan_problem_1])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X < pi / 2 - 1/80000) => 1 - cos(X) <= tan(X / 2))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (tan(X / 2) < 1 - cos(X) & X < pi / 2 - 1/80000 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (tan(skoXC1 / 2) < 1 - cos(skoXC1) & skoXC1 < pi / 2 - 1/80000 &
     0 <= skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (tan(skoXC1 / 2) < 1 - cos(skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 < pi / 2 - 1/80000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_5, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

fof(normalize_0_6, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (-10395 +
     skoXC1 *
     (10395/2 +
      skoXC1 *
      (4725/4 +
       skoXC1 *
       (-315/2 + skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/4 + skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))) |
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000038, cos(X_000037)) | X_000038 <= cos(X_000037)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000038,
            1 - 1/2 * X_000037 ^ 2 + 1/24 * X_000037 ^ 4 -
            1/720 * X_000037 ^ 6 + 1/40320 * X_000037 ^ 8 -
            1/3628800 * X_000037 ^ 10 + 1/479001600 * X_000037 ^ 12 -
            1/87178291200 * X_000037 ^ 14 +
            1/20922789888000 * X_000037 ^ 16 -
            1/6402373705728000 * X_000037 ^ 18 +
            1/2432902008176640000 * X_000037 ^ 20 -
            1/1124000727777607680000 * X_000037 ^ 22) |
     lgen(0, X_000038, cos(X_000037))),
    inference(subst, [], [cos_lower_bound_5])).

cnf(refute_0_3, plain,
    (~ lgen(0, X_000038,
            1 - 1/2 * X_000037 ^ 2 + 1/24 * X_000037 ^ 4 -
            1/720 * X_000037 ^ 6 + 1/40320 * X_000037 ^ 8 -
            1/3628800 * X_000037 ^ 10 + 1/479001600 * X_000037 ^ 12 -
            1/87178291200 * X_000037 ^ 14 +
            1/20922789888000 * X_000037 ^ 16 -
            1/6402373705728000 * X_000037 ^ 18 +
            1/2432902008176640000 * X_000037 ^ 20 -
            1/1124000727777607680000 * X_000037 ^ 22) |
     X_000038 <= cos(X_000037)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (1 +
     X_000037 *
     (X_000037 *
      (-1/2 +
       X_000037 *
       (X_000037 *
        (1/24 +
         X_000037 *
         (X_000037 *
          (-1/720 +
           X_000037 *
           (X_000037 *
            (1/40320 +
             X_000037 *
             (X_000037 *
              (-1/3628800 +
               X_000037 *
               (X_000037 *
                (1/479001600 +
                 X_000037 *
                 (X_000037 *
                  (-1/87178291200 +
                   X_000037 *
                   (X_000037 *
                    (1/20922789888000 +
                     X_000037 *
                     (X_000037 *
                      (-1/6402373705728000 +
                       X_000037 *
                       (X_000037 *
                        (1/2432902008176640000 +
                         X_000037 *
                         (X_000037 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     X_000038 | X_000038 <= cos(X_000037)),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((1 + cos(skoXC1) * -1) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <
     skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/4 + skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))) |
     1 + cos(skoXC1) * -1 <=
     skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_7, plain, (tan(skoXC1 / 2) < 1 - cos(skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain, (tan(skoXC1 * 1/2) < 1 + cos(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000116, tan(X_000115)) | X_000116 <= tan(X_000115)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (314159/200000 < X_000115 | X_000115 < 0 |
     ~ lgen(0, X_000116,
            -(21 * (X_000115 ^ 4 - 60 * X_000115 ^ 2 + 495)) * X_000115 /
            (X_000115 ^ 6 - 210 * X_000115 ^ 4 + 4725 * X_000115 ^ 2 -
             10395)) | lgen(0, X_000116, tan(X_000115))),
    inference(subst, [], [tan_lower_bound_pos2])).

cnf(refute_0_11, plain,
    (314159/200000 < X_000115 | X_000115 < 0 |
     ~ lgen(0, X_000116,
            -(21 * (X_000115 ^ 4 - 60 * X_000115 ^ 2 + 495)) * X_000115 /
            (X_000115 ^ 6 - 210 * X_000115 ^ 4 + 4725 * X_000115 ^ 2 -
             10395)) | X_000116 <= tan(X_000115)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (314159/200000 < X_000115 |
     X_000115 *
     (-10395 +
      X_000115 * (X_000115 * (1260 + X_000115 * (X_000115 * -21)))) /
     (-10395 +
      X_000115 *
      (X_000115 *
       (4725 + X_000115 * (X_000115 * (-210 + X_000115 * X_000115))))) <
     X_000116 | X_000115 < 0 | X_000116 <= tan(X_000115)),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (skoXC1 * 1/2 < 0 | 314159/200000 < skoXC1 * 1/2 |
     skoXC1 * 1/2 *
     (-10395 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (1260 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -21)))) /
     (-10395 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (4725 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 * (-210 + skoXC1 * 1/2 * (skoXC1 * 1/2)))))) <
     1 + cos(skoXC1) * -1 | 1 + cos(skoXC1) * -1 <= tan(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 * 1/2 < 0 | 314159/200000 < skoXC1 * 1/2 |
     skoXC1 * 1/2 *
     (-10395 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (1260 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -21)))) /
     (-10395 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (4725 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 * (-210 + skoXC1 * 1/2 * (skoXC1 * 1/2)))))) <
     1 + cos(skoXC1) * -1),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <
     1 + cos(skoXC1) * -1 | 314159/100000 < skoXC1),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (-5 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -5 + skoXC1 * (skoXC1 * 1/2) |
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200)))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_17, plain,
    (~ lgen(0, X_000024, cos(X_000023)) | X_000024 <= cos(X_000023)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800 +
            X_000023 ^ 12 / 479001600 - X_000023 ^ 14 / 87178291200) |
     lgen(0, X_000024, cos(X_000023))),
    inference(subst, [], [cos_lower_bound_3])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000024,
            1 - X_000023 ^ 2 / 2 + X_000023 ^ 4 / 24 - X_000023 ^ 6 / 720 +
            X_000023 ^ 8 / 40320 - X_000023 ^ 10 / 3628800 +
            X_000023 ^ 12 / 479001600 - X_000023 ^ 14 / 87178291200) |
     X_000024 <= cos(X_000023)),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (1 +
     X_000023 *
     (X_000023 *
      (-1/2 +
       X_000023 *
       (X_000023 *
        (1/24 +
         X_000023 *
         (X_000023 *
          (-1/720 +
           X_000023 *
           (X_000023 *
            (1/40320 +
             X_000023 *
             (X_000023 *
              (-1/3628800 +
               X_000023 *
               (X_000023 *
                (1/479001600 +
                 X_000023 * (X_000023 * -1/87178291200))))))))))))) <
     X_000024 | X_000024 <= cos(X_000023)),
    inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) |
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) <= cos(skoXC1)),
    inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    ((1 + cos(skoXC1) * -1) * (-5 + skoXC1 * (skoXC1 * 1/2)) <
     skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) |
     0 <= -5 + skoXC1 * (skoXC1 * 1/2) |
     1 + cos(skoXC1) * -1 <=
     skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) /
     (-5 + skoXC1 * (skoXC1 * 1/2))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000102, tan(X_000101)) | X_000102 <= tan(X_000101)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_24, plain,
    (157/100 < X_000101 | X_000101 < 0 |
     ~ lgen(0, X_000102,
            1/3 * (X_000101 ^ 2 - 15) * X_000101 /
            (2 * X_000101 ^ 2 - 5)) | lgen(0, X_000102, tan(X_000101))),
    inference(subst, [], [tan_lower_bound_pos1])).

cnf(refute_0_25, plain,
    (157/100 < X_000101 | X_000101 < 0 |
     ~ lgen(0, X_000102,
            1/3 * (X_000101 ^ 2 - 15) * X_000101 /
            (2 * X_000101 ^ 2 - 5)) | X_000102 <= tan(X_000101)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (157/100 < X_000101 |
     X_000101 * (-5 + X_000101 * (X_000101 * 1/3)) /
     (-5 + X_000101 * (X_000101 * 2)) < X_000102 | X_000101 < 0 |
     X_000102 <= tan(X_000101)), inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (skoXC1 * 1/2 < 0 | 157/100 < skoXC1 * 1/2 |
     skoXC1 * 1/2 * (-5 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/3)) /
     (-5 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 2)) < 1 + cos(skoXC1) * -1 |
     1 + cos(skoXC1) * -1 <= tan(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 * 1/2 < 0 | 157/100 < skoXC1 * 1/2 |
     skoXC1 * 1/2 * (-5 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/3)) /
     (-5 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 2)) < 1 + cos(skoXC1) * -1),
    inference(resolve, [], [refute_0_27, refute_0_8])).

cnf(refute_0_29, plain,
    (skoXC1 < 0 |
     skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) < 1 + cos(skoXC1) * -1 |
     157/50 < skoXC1), inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain, (skoXC1 < pi / 2 - 1/80000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_31, plain, (skoXC1 < -1/80000 + pi * 1/2),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_33, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_34, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_35, plain, (157/50 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) < 1 + cos(skoXC1) * -1 |
     157/50 < skoXC1), inference(resolve, [], [refute_0_35, refute_0_29])).

cnf(refute_0_37, plain, (skoXC1 <= 157/50),
    inference(decision, [],
              [refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_38, plain,
    (skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) < 1 + cos(skoXC1) * -1),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    ((1 + cos(skoXC1) * -1) * (-5 + skoXC1 * (skoXC1 * 1/2)) <
     skoXC1 * (-5/2 + skoXC1 * (skoXC1 * 1/24)) |
     0 <= -5 + skoXC1 * (skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_22, refute_0_38])).

cnf(refute_0_40, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     cos(skoXC1) * (-5 + skoXC1 * (skoXC1 * 1/2)) |
     skoXC1 * (skoXC1 * -1/2) <= -5),
    inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     1 + skoXC1 * -1/2 | 1 + skoXC1 * -1/2 <= cos(skoXC1)),
    inference(subst, [], [refute_0_20])).

cnf(refute_0_42, plain,
    (~ lgen(0, X_000086, tan(X_000085)) | X_000086 <= tan(X_000085)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_43, plain,
    (157/100 < X_000085 | X_000085 < 0 | ~ lgen(0, X_000086, X_000085) |
     lgen(0, X_000086, tan(X_000085))),
    inference(subst, [], [tan_lower_bound_pos0])).

cnf(refute_0_44, plain,
    (157/100 < X_000085 | X_000085 < 0 | ~ lgen(0, X_000086, X_000085) |
     X_000086 <= tan(X_000085)),
    inference(resolve, [], [refute_0_43, refute_0_42])).

cnf(refute_0_45, plain,
    (157/100 < X_000085 | X_000085 < X_000086 | X_000085 < 0 |
     X_000086 <= tan(X_000085)), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (skoXC1 * 1/2 < 0 | 157/100 < skoXC1 * 1/2 |
     skoXC1 * 1/2 < 1 + cos(skoXC1) * -1 |
     1 + cos(skoXC1) * -1 <= tan(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (skoXC1 * 1/2 < 0 | 157/100 < skoXC1 * 1/2 |
     skoXC1 * 1/2 < 1 + cos(skoXC1) * -1),
    inference(resolve, [], [refute_0_46, refute_0_8])).

cnf(refute_0_48, plain,
    (skoXC1 < 0 | cos(skoXC1) < 1 + skoXC1 * -1/2 | 157/50 < skoXC1),
    inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain,
    (cos(skoXC1) < 1 + skoXC1 * -1/2 | 157/50 < skoXC1),
    inference(resolve, [], [refute_0_35, refute_0_48])).

cnf(refute_0_50, plain, (cos(skoXC1) < 1 + skoXC1 * -1/2),
    inference(resolve, [], [refute_0_37, refute_0_49])).

cnf(refute_0_51, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     1 + skoXC1 * -1/2),
    inference(resolve, [], [refute_0_41, refute_0_50])).

cnf(refute_0_52, plain,
    (0 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200)))))))))))))),
    inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (~ lgen(0, X_000010, cos(X_000009)) | X_000010 <= cos(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_54, plain,
    (~ lgen(0, X_000010,
            1 - X_000009 ^ 2 / 2 + X_000009 ^ 4 / 24 -
            X_000009 ^ 6 / 720) | lgen(0, X_000010, cos(X_000009))),
    inference(subst, [], [cos_lower_bound_1])).

cnf(refute_0_55, plain,
    (~ lgen(0, X_000010,
            1 - X_000009 ^ 2 / 2 + X_000009 ^ 4 / 24 -
            X_000009 ^ 6 / 720) | X_000010 <= cos(X_000009)),
    inference(resolve, [], [refute_0_54, refute_0_53])).

cnf(refute_0_56, plain,
    (1 +
     X_000009 *
     (X_000009 *
      (-1/2 +
       X_000009 * (X_000009 * (1/24 + X_000009 * (X_000009 * -1/720))))) <
     X_000010 | X_000010 <= cos(X_000009)),
    inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     1 + skoXC1 * -1/2 | 1 + skoXC1 * -1/2 <= cos(skoXC1)),
    inference(subst, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     1 + skoXC1 * -1/2),
    inference(resolve, [], [refute_0_57, refute_0_50])).

cnf(refute_0_59, plain,
    (0 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/720)))))),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain, (-5 < skoXC1 * (skoXC1 * -1/2)),
    inference(decision, [],
              [refute_0_52, refute_0_59, refute_0_31, refute_0_32,
               refute_0_33, refute_0_34])).

cnf(refute_0_61, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     cos(skoXC1) * (-5 + skoXC1 * (skoXC1 * 1/2))),
    inference(resolve, [], [refute_0_40, refute_0_60])).

cnf(refute_0_62, plain,
    (cos(skoXC1) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -5 + skoXC1 * (skoXC1 * 1/2)),
    inference(split, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (cos(skoXC1) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) | skoXC1 * (skoXC1 * -1/2) <= -5),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (cos(skoXC1) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2))),
    inference(resolve, [], [refute_0_63, refute_0_60])).

cnf(refute_0_65, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2))),
    inference(resolve, [], [refute_0_21, refute_0_64])).

cnf(refute_0_66, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (-5 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -5 + skoXC1 * (skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_16, refute_0_65])).

cnf(refute_0_67, plain,
    (0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 *
         (skoXC1 *
          (1/36 +
           skoXC1 *
           (skoXC1 *
            (-11/13440 +
             skoXC1 *
             (skoXC1 *
              (1/72576 +
               skoXC1 *
               (skoXC1 *
                (-71/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (1/908107200 +
                   skoXC1 * (skoXC1 * -1/174356582400))))))))))))))) |
     skoXC1 * (skoXC1 * -1/2) <= -5),
    inference(arithmetic, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (-5 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -5 + skoXC1 * (skoXC1 * 1/2) |
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_69, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) |
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2)) <= cos(skoXC1)),
    inference(subst, [], [refute_0_56])).

cnf(refute_0_70, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24))) /
     (-5 + skoXC1 * (skoXC1 * 1/2))),
    inference(resolve, [], [refute_0_69, refute_0_64])).

cnf(refute_0_71, plain,
    (-5 + skoXC1 * (5/2 + skoXC1 * (1/2 + skoXC1 * -1/24)) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (-5 + skoXC1 * (skoXC1 * 1/2)) | 0 <= -5 + skoXC1 * (skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_68, refute_0_70])).

cnf(refute_0_72, plain,
    (0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 * (skoXC1 * (1/36 + skoXC1 * (skoXC1 * -1/1440))))))) |
     skoXC1 * (skoXC1 * -1/2) <= -5),
    inference(arithmetic, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (-5 < skoXC1 * (skoXC1 * -1/2) |
     0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 * (skoXC1 * (1/36 + skoXC1 * (skoXC1 * -1/1440)))))))),
    inference(decision, [],
              [refute_0_52, refute_0_59, refute_0_31, refute_0_32,
               refute_0_33, refute_0_34])).

cnf(refute_0_74, plain,
    (0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 * (skoXC1 * (1/36 + skoXC1 * (skoXC1 * -1/1440)))))))),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain,
    (-5 < skoXC1 * (skoXC1 * -1/2) |
     0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 *
         (skoXC1 *
          (1/36 +
           skoXC1 *
           (skoXC1 *
            (-11/13440 +
             skoXC1 *
             (skoXC1 *
              (1/72576 +
               skoXC1 *
               (skoXC1 *
                (-71/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (1/908107200 +
                   skoXC1 * (skoXC1 * -1/174356582400)))))))))))))))),
    inference(decision, [],
              [refute_0_74, refute_0_52, refute_0_59, refute_0_31,
               refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_76, plain,
    (0 <
     skoXC1 *
     (-5/2 +
      skoXC1 *
      (5/2 +
       skoXC1 *
       (1/24 +
        skoXC1 *
        (-11/24 +
         skoXC1 *
         (skoXC1 *
          (1/36 +
           skoXC1 *
           (skoXC1 *
            (-11/13440 +
             skoXC1 *
             (skoXC1 *
              (1/72576 +
               skoXC1 *
               (skoXC1 *
                (-71/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (1/908107200 +
                   skoXC1 * (skoXC1 * -1/174356582400)))))))))))))))),
    inference(resolve, [], [refute_0_67, refute_0_75])).

cnf(refute_0_77, plain, (314159/100000 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_76, refute_0_74, refute_0_52, refute_0_59,
               refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_78, plain,
    (skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <
     1 + cos(skoXC1) * -1 | 314159/100000 < skoXC1),
    inference(resolve, [], [refute_0_77, refute_0_15])).

cnf(refute_0_79, plain, (skoXC1 <= 314159/100000),
    inference(decision, [],
              [refute_0_76, refute_0_74, refute_0_52, refute_0_59,
               refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_80, plain,
    (skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <
     1 + cos(skoXC1) * -1),
    inference(resolve, [], [refute_0_79, refute_0_78])).

cnf(refute_0_81, plain,
    ((1 + cos(skoXC1) * -1) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) <
     skoXC1 *
     (-10395/2 +
      skoXC1 * (skoXC1 * (315/2 + skoXC1 * (skoXC1 * -21/32)))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/4 + skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))),
    inference(resolve, [], [refute_0_6, refute_0_80])).

cnf(refute_0_82, plain,
    (-10395 +
     skoXC1 *
     (10395/2 +
      skoXC1 *
      (4725/4 +
       skoXC1 *
       (-315/2 + skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64))))) <
     cos(skoXC1) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -10395), inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (-10395 <
     skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64)))))),
    inference(decision, [],
              [refute_0_76, refute_0_74, refute_0_52, refute_0_59,
               refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_84, plain,
    (-10395 +
     skoXC1 *
     (10395/2 +
      skoXC1 *
      (4725/4 +
       skoXC1 *
       (-315/2 + skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64))))) <
     cos(skoXC1) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))))),
    inference(resolve, [], [refute_0_82, refute_0_83])).

cnf(refute_0_85, plain,
    (cos(skoXC1) <
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/4 + skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))),
    inference(split, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (cos(skoXC1) <
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -10395), inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (cos(skoXC1) <
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))))),
    inference(resolve, [], [refute_0_86, refute_0_83])).

cnf(refute_0_88, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (-10395 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (4725/4 +
        skoXC1 *
        (-315/2 +
         skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64)))))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64))))))),
    inference(resolve, [], [refute_0_5, refute_0_87])).

cnf(refute_0_89, plain,
    (-10395 +
     skoXC1 *
     (10395/2 +
      skoXC1 *
      (4725/4 +
       skoXC1 *
       (-315/2 + skoXC1 * (-105/8 + skoXC1 * (21/32 + skoXC1 * 1/64))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/4 +
        skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/4 + skoXC1 * (skoXC1 * (-105/8 + skoXC1 * (skoXC1 * 1/64)))))),
    inference(resolve, [], [refute_0_0, refute_0_88])).

cnf(refute_0_90, plain,
    (0 <
     skoXC1 *
     (-10395/2 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (315/2 +
        skoXC1 *
        (-4095/4 +
         skoXC1 *
         (-21/32 +
          skoXC1 *
          (2247/32 +
           skoXC1 *
           (skoXC1 *
            (-157/64 +
             skoXC1 *
             (skoXC1 *
              (49/960 +
               skoXC1 *
               (skoXC1 *
                (-1/1440 +
                 skoXC1 *
                 (skoXC1 *
                  (1459/221405184 +
                   skoXC1 *
                   (skoXC1 *
                    (-1013/22140518400 +
                     skoXC1 *
                     (skoXC1 *
                      (467/1935713894400 +
                       skoXC1 *
                       (skoXC1 *
                        (-1993/2002388484096000 +
                         skoXC1 *
                         (skoXC1 *
                          (4271/1297547737694208000 +
                           skoXC1 *
                           (skoXC1 *
                            (-761/85638150687817728000 +
                             skoXC1 *
                             (skoXC1 *
                              (31/1712763013756354560000 +
                               skoXC1 *
                               (skoXC1 *
                                -1/
                                71936046577766891520000))))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -10395), inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -10395 |
     skoXC1 *
     (-10395/2 +
      skoXC1 *
      (10395/2 +
       skoXC1 *
       (315/2 +
        skoXC1 *
        (-4095/4 +
         skoXC1 *
         (-21/32 +
          skoXC1 *
          (2247/32 +
           skoXC1 *
           (skoXC1 *
            (-157/64 +
             skoXC1 *
             (skoXC1 *
              (49/960 +
               skoXC1 *
               (skoXC1 *
                (-1/1440 +
                 skoXC1 *
                 (skoXC1 *
                  (1459/221405184 +
                   skoXC1 *
                   (skoXC1 *
                    (-1013/22140518400 +
                     skoXC1 *
                     (skoXC1 *
                      (467/1935713894400 +
                       skoXC1 *
                       (skoXC1 *
                        (-1993/2002388484096000 +
                         skoXC1 *
                         (skoXC1 *
                          (4271/1297547737694208000 +
                           skoXC1 *
                           (skoXC1 *
                            (-761/85638150687817728000 +
                             skoXC1 *
                             (skoXC1 *
                              (31/1712763013756354560000 +
                               skoXC1 *
                               (skoXC1 *
                                -1/
                                71936046577766891520000))))))))))))))))))))))))))) <=
     0),
    inference(decision, [],
              [refute_0_76, refute_0_74, refute_0_52, refute_0_59,
               refute_0_31, refute_0_32, refute_0_33, refute_0_34])).

cnf(refute_0_92, plain,
    (skoXC1 *
     (skoXC1 *
      (-4725/4 +
       skoXC1 * (skoXC1 * (105/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -10395), inference(resolve, [], [refute_0_91, refute_0_90])).

cnf(refute_0_93, plain, ($false),
    inference(resolve, [], [refute_0_92, refute_0_83])).
% SZS output end CNFRefutation for tan-problem-1-weak.tptp
