% SZS output start CNFRefutation for tan-problem-1-sin-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(tan_upper_bound_pos1, axiom,
    (X < 0 | 1 < X | ~ lgen(R, 987/100 * X / (-4 * X ^ 2 + 493/50), Y) |
     lgen(R, tan(X), Y))).

cnf(tan_upper_bound_pos2, axiom,
    (X < 0 | 39/25 < X |
     ~ lgen(R,
            -6 * X * (120 - 20 * X ^ 2 + X ^ 4) /
            (-720 + 360 * X ^ 2 - 30 * X ^ 4 + X ^ 6), Y) |
     lgen(R, tan(X), Y))).

cnf(tan_upper_bound_pos3, axiom,
    (X < 0 | 314159/200000 < X |
     ~ lgen(R,
            -10 * X *
            (362880 - 60480 * X ^ 2 + 3024 * X ^ 4 - 72 * X ^ 6 + X ^ 8) /
            (-3628800 + 1814400 * X ^ 2 - 151200 * X ^ 4 + 5040 * X ^ 6 -
             90 * X ^ 8 + X ^ 10), Y) | lgen(R, tan(X), Y))).

cnf(tan_upper_bound_neg0, axiom,
    (0 < X | X < -157/100 | ~ lgen(R, X, Y) | lgen(R, tan(X), Y))).

fof(pi_upper_bound, axiom, (pi < 31415927/10000000)).

fof(pi_lower_bound, axiom, (15707963/5000000 < pi)).

fof(tan_problem_1, conjecture,
    (! [X] : ((0 <= X & X < pi / 2 - 10 ^ -7) => tan(X / 2) <= sin(X)))).

fof(subgoal_0, plain,
    (! [X] : ((0 <= X & X < pi / 2 - 10 ^ -7) => tan(X / 2) <= sin(X))),
    inference(strip, [], [tan_problem_1])).

fof(negate_0_0, plain,
    (~ ! [X] : ((0 <= X & X < pi / 2 - 10 ^ -7) => tan(X / 2) <= sin(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (X < pi / 2 - 10 ^ -7 & sin(X) < tan(X / 2) & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < pi / 2 - 10 ^ -7 & sin(skoXC1) < tan(skoXC1 / 2) &
     0 <= skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (sin(skoXC1) < tan(skoXC1 / 2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < pi / 2 - 10 ^ -7),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [pi_lower_bound])).

fof(normalize_0_6, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [pi_upper_bound])).

cnf(refute_0_0, plain,
    (skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))) |
     0 <=
     -3628800 +
     skoXC1 *
     (skoXC1 *
      (453600 +
       skoXC1 *
       (skoXC1 *
        (-9450 +
         skoXC1 *
         (skoXC1 *
          (315/4 +
           skoXC1 *
           (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024))))))))) |
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) /
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain, (sin(skoXC1) < tan(skoXC1 / 2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain, (sin(skoXC1) < tan(skoXC1 * 1/2)),
    inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (~ lgen(0, X_000083, sin(X_000082)) | X_000083 <= sin(X_000082)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_4, plain,
    (X_000082 < 0 |
     ~ lgen(0, X_000083,
            X_000082 - X_000082 ^ 3 / 6 + X_000082 ^ 5 / 120 -
            X_000082 ^ 7 / 5040 + X_000082 ^ 9 / 362880 -
            X_000082 ^ 11 / 39916800) | lgen(0, X_000083, sin(X_000082))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_5, plain,
    (X_000082 < 0 |
     ~ lgen(0, X_000083,
            X_000082 - X_000082 ^ 3 / 6 + X_000082 ^ 5 / 120 -
            X_000082 ^ 7 / 5040 + X_000082 ^ 9 / 362880 -
            X_000082 ^ 11 / 39916800) | X_000083 <= sin(X_000082)),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (X_000082 *
     (1 +
      X_000082 *
      (X_000082 *
       (-1/6 +
        X_000082 *
        (X_000082 *
         (1/120 +
          X_000082 *
          (X_000082 *
           (-1/5040 +
            X_000082 *
            (X_000082 *
             (1/362880 + X_000082 * (X_000082 * -1/39916800)))))))))) <
     X_000083 | X_000082 < 0 | X_000083 <= sin(X_000082)),
    inference(arithmetic, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     tan(skoXC1 * 1/2) | tan(skoXC1 * 1/2) <= sin(skoXC1)),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     tan(skoXC1 * 1/2)), inference(resolve, [], [refute_0_7, refute_0_2])).

cnf(refute_0_9, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_10, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     tan(skoXC1 * 1/2)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (~ lgen(0, tan(X_000134), X_000135) | tan(X_000134) <= X_000135),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_12, plain,
    (314159/200000 < X_000134 | X_000134 < 0 |
     ~ lgen(0,
            -10 * X_000134 *
            (362880 - 60480 * X_000134 ^ 2 + 3024 * X_000134 ^ 4 -
             72 * X_000134 ^ 6 + X_000134 ^ 8) /
            (-3628800 + 1814400 * X_000134 ^ 2 - 151200 * X_000134 ^ 4 +
             5040 * X_000134 ^ 6 - 90 * X_000134 ^ 8 + X_000134 ^ 10),
            X_000135) | lgen(0, tan(X_000134), X_000135)),
    inference(subst, [], [tan_upper_bound_pos3])).

cnf(refute_0_13, plain,
    (314159/200000 < X_000134 | X_000134 < 0 |
     ~ lgen(0,
            -10 * X_000134 *
            (362880 - 60480 * X_000134 ^ 2 + 3024 * X_000134 ^ 4 -
             72 * X_000134 ^ 6 + X_000134 ^ 8) /
            (-3628800 + 1814400 * X_000134 ^ 2 - 151200 * X_000134 ^ 4 +
             5040 * X_000134 ^ 6 - 90 * X_000134 ^ 8 + X_000134 ^ 10),
            X_000135) | tan(X_000134) <= X_000135),
    inference(resolve, [], [refute_0_12, refute_0_11])).

cnf(refute_0_14, plain,
    (314159/200000 < X_000134 | X_000134 < 0 |
     X_000135 <
     X_000134 *
     (-3628800 +
      X_000134 *
      (X_000134 *
       (604800 +
        X_000134 *
        (X_000134 *
         (-30240 +
          X_000134 *
          (X_000134 * (720 + X_000134 * (X_000134 * -10)))))))) /
     (-3628800 +
      X_000134 *
      (X_000134 *
       (1814400 +
        X_000134 *
        (X_000134 *
         (-151200 +
          X_000134 *
          (X_000134 *
           (5040 +
            X_000134 * (X_000134 * (-90 + X_000134 * X_000134))))))))) |
     tan(X_000134) <= X_000135), inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (skoXC1 * 1/2 < 0 | 314159/200000 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 *
     (-3628800 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (604800 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (-30240 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (720 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -10)))))))) /
     (-3628800 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (1814400 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (-151200 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 * (-90 + skoXC1 * 1/2 * (skoXC1 * 1/2)))))))))) |
     tan(skoXC1 * 1/2) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (skoXC1 * 1/2 < 0 | 314159/200000 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 *
     (-3628800 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (604800 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (-30240 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (720 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -10)))))))) /
     (-3628800 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (1814400 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (-151200 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 * (-90 + skoXC1 * 1/2 * (skoXC1 * 1/2))))))))))),
    inference(resolve, [], [refute_0_15, refute_0_10])).

cnf(refute_0_17, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) /
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))) |
     314159/100000 < skoXC1), inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     0 <=
     -720 +
     skoXC1 *
     (skoXC1 *
      (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64))))) |
     skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) /
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_19, plain,
    (~ lgen(0, tan(X_000110), X_000111) | tan(X_000110) <= X_000111),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_20, plain,
    (39/25 < X_000110 | X_000110 < 0 |
     ~ lgen(0,
            -6 * X_000110 * (120 - 20 * X_000110 ^ 2 + X_000110 ^ 4) /
            (-720 + 360 * X_000110 ^ 2 - 30 * X_000110 ^ 4 + X_000110 ^ 6),
            X_000111) | lgen(0, tan(X_000110), X_000111)),
    inference(subst, [], [tan_upper_bound_pos2])).

cnf(refute_0_21, plain,
    (39/25 < X_000110 | X_000110 < 0 |
     ~ lgen(0,
            -6 * X_000110 * (120 - 20 * X_000110 ^ 2 + X_000110 ^ 4) /
            (-720 + 360 * X_000110 ^ 2 - 30 * X_000110 ^ 4 + X_000110 ^ 6),
            X_000111) | tan(X_000110) <= X_000111),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain,
    (39/25 < X_000110 | X_000110 < 0 |
     X_000111 <
     X_000110 *
     (-720 + X_000110 * (X_000110 * (120 + X_000110 * (X_000110 * -6)))) /
     (-720 +
      X_000110 *
      (X_000110 *
       (360 + X_000110 * (X_000110 * (-30 + X_000110 * X_000110))))) |
     tan(X_000110) <= X_000111), inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 * 1/2 < 0 | 39/25 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 *
     (-720 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -6)))) /
     (-720 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (360 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 * (-30 + skoXC1 * 1/2 * (skoXC1 * 1/2)))))) |
     tan(skoXC1 * 1/2) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (skoXC1 * 1/2 < 0 | 39/25 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 *
     (-720 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -6)))) /
     (-720 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (360 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 * (-30 + skoXC1 * 1/2 * (skoXC1 * 1/2))))))),
    inference(resolve, [], [refute_0_23, refute_0_10])).

cnf(refute_0_25, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) /
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     78/25 < skoXC1), inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (493/50 + skoXC1 * (skoXC1 * -1)) < skoXC1 * 987/200 |
     493/50 + skoXC1 * (skoXC1 * -1) <= 0 |
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_27, plain,
    (~ lgen(0, tan(X_000096), X_000097) | tan(X_000096) <= X_000097),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (1 < X_000096 | X_000096 < 0 |
     ~ lgen(0, 987/100 * X_000096 / (-4 * X_000096 ^ 2 + 493/50),
            X_000097) | lgen(0, tan(X_000096), X_000097)),
    inference(subst, [], [tan_upper_bound_pos1])).

cnf(refute_0_29, plain,
    (1 < X_000096 | X_000096 < 0 |
     ~ lgen(0, 987/100 * X_000096 / (-4 * X_000096 ^ 2 + 493/50),
            X_000097) | tan(X_000096) <= X_000097),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (1 < X_000096 | X_000096 < 0 |
     X_000097 <
     X_000096 * 987/100 / (493/50 + X_000096 * (X_000096 * -4)) |
     tan(X_000096) <= X_000097), inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * 1/2 < 0 | 1 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 * 987/100 /
     (493/50 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -4)) |
     tan(skoXC1 * 1/2) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (skoXC1 * 1/2 < 0 | 1 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 * 987/100 /
     (493/50 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -4))),
    inference(resolve, [], [refute_0_31, refute_0_10])).

cnf(refute_0_33, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) | 2 < skoXC1),
    inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (~ lgen(0, tan(X_000089), X_000090) | tan(X_000089) <= X_000090),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_35, plain,
    (0 < X_000089 | X_000089 < -157/100 | ~ lgen(0, X_000089, X_000090) |
     lgen(0, tan(X_000089), X_000090)),
    inference(subst, [], [tan_upper_bound_neg0])).

cnf(refute_0_36, plain,
    (0 < X_000089 | X_000089 < -157/100 | ~ lgen(0, X_000089, X_000090) |
     tan(X_000089) <= X_000090),
    inference(resolve, [], [refute_0_35, refute_0_34])).

cnf(refute_0_37, plain,
    (X_000090 < X_000089 | 0 < X_000089 | X_000089 < -157/100 |
     tan(X_000089) <= X_000090), inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (skoXC1 * 1/2 < -157/100 | 0 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2 |
     tan(skoXC1 * 1/2) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800))))))))))),
    inference(subst, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (skoXC1 * 1/2 < -157/100 | 0 < skoXC1 * 1/2 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 1/2), inference(resolve, [], [refute_0_38, refute_0_10])).

cnf(refute_0_40, plain,
    (skoXC1 < -157/50 |
     0 <
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (-1/120 +
          skoXC1 *
          (skoXC1 *
           (1/5040 +
            skoXC1 *
            (skoXC1 *
             (-1/362880 + skoXC1 * (skoXC1 * 1/39916800)))))))))) |
     0 < skoXC1), inference(arithmetic, [], [refute_0_39])).

cnf(refute_0_41, plain, (skoXC1 < pi / 2 - 10 ^ -7),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_42, plain, (skoXC1 < -1/10000000 + pi * 1/2),
    inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain, (15707963/5000000 < pi),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_44, plain, (pi < 31415927/10000000),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_45, plain,
    (skoXC1 < -157/50 | 0 < skoXC1 |
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (-1/120 +
          skoXC1 *
          (skoXC1 *
           (1/5040 +
            skoXC1 *
            (skoXC1 *
             (-1/362880 + skoXC1 * (skoXC1 * 1/39916800)))))))))) <= 0),
    inference(decision, [],
              [refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_46, plain, (skoXC1 < -157/50 | 0 < skoXC1),
    inference(resolve, [], [refute_0_45, refute_0_40])).

cnf(refute_0_47, plain, (0 < skoXC1 | -157/50 <= skoXC1),
    inference(decision, [],
              [refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_48, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_47, refute_0_46])).

cnf(refute_0_49, plain, (2 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_48, refute_0_42, refute_0_43, refute_0_44,
               refute_0_9])).

cnf(refute_0_50, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) | 2 < skoXC1),
    inference(resolve, [], [refute_0_49, refute_0_33])).

cnf(refute_0_51, plain, (skoXC1 <= 2),
    inference(decision, [],
              [refute_0_48, refute_0_42, refute_0_43, refute_0_44,
               refute_0_9])).

cnf(refute_0_52, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_51, refute_0_50])).

cnf(refute_0_53, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (493/50 + skoXC1 * (skoXC1 * -1)) < skoXC1 * 987/200 |
     493/50 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_26, refute_0_52])).

cnf(refute_0_54, plain,
    (0 <
     skoXC1 *
     (-197/40 +
      skoXC1 *
      (skoXC1 *
       (793/300 +
        skoXC1 *
        (skoXC1 *
         (-1493/6000 +
          skoXC1 *
          (skoXC1 *
           (2593/252000 +
            skoXC1 *
            (skoXC1 *
             (-4093/18144000 +
              skoXC1 *
              (skoXC1 *
               (5993/1995840000 +
                skoXC1 * (skoXC1 * -1/39916800)))))))))))) |
     skoXC1 * (skoXC1 * -1) <= -493/50),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (493/50 + skoXC1 * (skoXC1 * -1)) < skoXC1 * 987/200 |
     493/50 + skoXC1 * (skoXC1 * -1) <= 0 |
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) <=
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_56, plain,
    (~ lgen(0, X_000055, sin(X_000054)) | X_000055 <= sin(X_000054)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_57, plain,
    (X_000054 < 0 | ~ lgen(0, X_000055, X_000054 - X_000054 ^ 3 / 6) |
     lgen(0, X_000055, sin(X_000054))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_58, plain,
    (X_000054 < 0 | ~ lgen(0, X_000055, X_000054 - X_000054 ^ 3 / 6) |
     X_000055 <= sin(X_000054)),
    inference(resolve, [], [refute_0_57, refute_0_56])).

cnf(refute_0_59, plain,
    (X_000054 * (1 + X_000054 * (X_000054 * -1/6)) < X_000055 |
     X_000054 < 0 | X_000055 <= sin(X_000054)),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) < tan(skoXC1 * 1/2) |
     tan(skoXC1 * 1/2) <= sin(skoXC1)),
    inference(subst, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) < tan(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_60, refute_0_2])).

cnf(refute_0_62, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) < tan(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_9, refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 * 1/2 < 0 | 1 < skoXC1 * 1/2 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 * 1/2 * 987/100 /
     (493/50 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -4)) |
     tan(skoXC1 * 1/2) <= skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6))),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_64, plain,
    (skoXC1 * 1/2 < 0 | 1 < skoXC1 * 1/2 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 * 1/2 * 987/100 /
     (493/50 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -4))),
    inference(resolve, [], [refute_0_63, refute_0_62])).

cnf(refute_0_65, plain,
    (skoXC1 < 0 |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) | 2 < skoXC1),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1)) | 2 < skoXC1),
    inference(resolve, [], [refute_0_49, refute_0_65])).

cnf(refute_0_67, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) <
     skoXC1 * 987/200 / (493/50 + skoXC1 * (skoXC1 * -1))),
    inference(resolve, [], [refute_0_51, refute_0_66])).

cnf(refute_0_68, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/6)) *
     (493/50 + skoXC1 * (skoXC1 * -1)) < skoXC1 * 987/200 |
     493/50 + skoXC1 * (skoXC1 * -1) <= 0),
    inference(resolve, [], [refute_0_55, refute_0_67])).

cnf(refute_0_69, plain,
    (0 <
     skoXC1 *
     (-197/40 + skoXC1 * (skoXC1 * (793/300 + skoXC1 * (skoXC1 * -1/6)))) |
     skoXC1 * (skoXC1 * -1) <= -493/50),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (-493/50 < skoXC1 * (skoXC1 * -1) |
     0 <
     skoXC1 *
     (-197/40 + skoXC1 * (skoXC1 * (793/300 + skoXC1 * (skoXC1 * -1/6))))),
    inference(decision, [],
              [refute_0_48, refute_0_42, refute_0_43, refute_0_44,
               refute_0_9])).

cnf(refute_0_71, plain,
    (0 <
     skoXC1 *
     (-197/40 + skoXC1 * (skoXC1 * (793/300 + skoXC1 * (skoXC1 * -1/6))))),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (-493/50 < skoXC1 * (skoXC1 * -1) |
     0 <
     skoXC1 *
     (-197/40 +
      skoXC1 *
      (skoXC1 *
       (793/300 +
        skoXC1 *
        (skoXC1 *
         (-1493/6000 +
          skoXC1 *
          (skoXC1 *
           (2593/252000 +
            skoXC1 *
            (skoXC1 *
             (-4093/18144000 +
              skoXC1 *
              (skoXC1 *
               (5993/1995840000 +
                skoXC1 * (skoXC1 * -1/39916800))))))))))))),
    inference(decision, [],
              [refute_0_71, refute_0_48, refute_0_42, refute_0_43,
               refute_0_44, refute_0_9])).

cnf(refute_0_73, plain,
    (0 <
     skoXC1 *
     (-197/40 +
      skoXC1 *
      (skoXC1 *
       (793/300 +
        skoXC1 *
        (skoXC1 *
         (-1493/6000 +
          skoXC1 *
          (skoXC1 *
           (2593/252000 +
            skoXC1 *
            (skoXC1 *
             (-4093/18144000 +
              skoXC1 *
              (skoXC1 *
               (5993/1995840000 +
                skoXC1 * (skoXC1 * -1/39916800))))))))))))),
    inference(resolve, [], [refute_0_54, refute_0_72])).

cnf(refute_0_74, plain, (78/25 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_73, refute_0_71, refute_0_48, refute_0_42,
               refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_75, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) /
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     78/25 < skoXC1), inference(resolve, [], [refute_0_74, refute_0_25])).

cnf(refute_0_76, plain, (skoXC1 <= 78/25),
    inference(decision, [],
              [refute_0_73, refute_0_71, refute_0_48, refute_0_42,
               refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_77, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) /
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64))))))),
    inference(resolve, [], [refute_0_76, refute_0_75])).

cnf(refute_0_78, plain,
    (skoXC1 *
     (-360 + skoXC1 * (skoXC1 * (15 + skoXC1 * (skoXC1 * -3/16)))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-720 +
      skoXC1 *
      (skoXC1 *
       (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))) |
     0 <=
     -720 +
     skoXC1 *
     (skoXC1 *
      (90 + skoXC1 * (skoXC1 * (-15/8 + skoXC1 * (skoXC1 * 1/64)))))),
    inference(resolve, [], [refute_0_18, refute_0_77])).

cnf(refute_0_79, plain,
    (0 <
     skoXC1 *
     (-360 +
      skoXC1 *
      (skoXC1 *
       (195 +
        skoXC1 *
        (skoXC1 *
         (-363/16 +
          skoXC1 *
          (skoXC1 *
           (547/448 +
            skoXC1 *
            (skoXC1 *
             (-307/8064 +
              skoXC1 *
              (skoXC1 *
               (1363/1774080 +
                skoXC1 *
                (skoXC1 *
                 (-1/95040 +
                  skoXC1 *
                  (skoXC1 *
                   (23/255467520 +
                    skoXC1 * (skoXC1 * -1/2554675200)))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-90 + skoXC1 * (skoXC1 * (15/8 + skoXC1 * (skoXC1 * -1/64))))) <=
     -720), inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (-720 <
     skoXC1 *
     (skoXC1 *
      (-90 + skoXC1 * (skoXC1 * (15/8 + skoXC1 * (skoXC1 * -1/64))))) |
     0 <
     skoXC1 *
     (-360 +
      skoXC1 *
      (skoXC1 *
       (195 +
        skoXC1 *
        (skoXC1 *
         (-363/16 +
          skoXC1 *
          (skoXC1 *
           (547/448 +
            skoXC1 *
            (skoXC1 *
             (-307/8064 +
              skoXC1 *
              (skoXC1 *
               (1363/1774080 +
                skoXC1 *
                (skoXC1 *
                 (-1/95040 +
                  skoXC1 *
                  (skoXC1 *
                   (23/255467520 +
                    skoXC1 * (skoXC1 * -1/2554675200))))))))))))))))),
    inference(decision, [],
              [refute_0_73, refute_0_71, refute_0_48, refute_0_42,
               refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_81, plain,
    (0 <
     skoXC1 *
     (-360 +
      skoXC1 *
      (skoXC1 *
       (195 +
        skoXC1 *
        (skoXC1 *
         (-363/16 +
          skoXC1 *
          (skoXC1 *
           (547/448 +
            skoXC1 *
            (skoXC1 *
             (-307/8064 +
              skoXC1 *
              (skoXC1 *
               (1363/1774080 +
                skoXC1 *
                (skoXC1 *
                 (-1/95040 +
                  skoXC1 *
                  (skoXC1 *
                   (23/255467520 +
                    skoXC1 * (skoXC1 * -1/2554675200))))))))))))))))),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain, (314159/100000 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_81, refute_0_73, refute_0_71, refute_0_48,
               refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_83, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) /
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))) |
     314159/100000 < skoXC1),
    inference(resolve, [], [refute_0_82, refute_0_17])).

cnf(refute_0_84, plain, (skoXC1 <= 314159/100000),
    inference(decision, [],
              [refute_0_81, refute_0_73, refute_0_71, refute_0_48,
               refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_85, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) <
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) /
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024))))))))))),
    inference(resolve, [], [refute_0_84, refute_0_83])).

cnf(refute_0_86, plain,
    (skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (75600 +
        skoXC1 *
        (skoXC1 *
         (-945 +
          skoXC1 * (skoXC1 * (45/8 + skoXC1 * (skoXC1 * -5/256)))))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 + skoXC1 * (skoXC1 * -1/39916800)))))))))) *
     (-3628800 +
      skoXC1 *
      (skoXC1 *
       (453600 +
        skoXC1 *
        (skoXC1 *
         (-9450 +
          skoXC1 *
          (skoXC1 *
           (315/4 +
            skoXC1 *
            (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))) |
     0 <=
     -3628800 +
     skoXC1 *
     (skoXC1 *
      (453600 +
       skoXC1 *
       (skoXC1 *
        (-9450 +
         skoXC1 *
         (skoXC1 *
          (315/4 +
           skoXC1 *
           (skoXC1 * (-45/128 + skoXC1 * (skoXC1 * 1/1024)))))))))),
    inference(resolve, [], [refute_0_0, refute_0_85])).

cnf(refute_0_87, plain,
    (0 <
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (982800 +
        skoXC1 *
        (skoXC1 *
         (-114345 +
          skoXC1 *
          (skoXC1 *
           (49185/8 +
            skoXC1 *
            (skoXC1 *
             (-49205/256 +
              skoXC1 *
              (skoXC1 *
               (44287/11264 +
                skoXC1 *
                (skoXC1 *
                 (-3793/67584 +
                  skoXC1 *
                  (skoXC1 *
                   (15091/28385280 +
                    skoXC1 *
                    (skoXC1 *
                     (-89/28385280 +
                      skoXC1 *
                      (skoXC1 *
                       (47/4087480320 +
                        skoXC1 *
                        (skoXC1 * -1/40874803200)))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-453600 +
       skoXC1 *
       (skoXC1 *
        (9450 +
         skoXC1 *
         (skoXC1 *
          (-315/4 +
           skoXC1 *
           (skoXC1 * (45/128 + skoXC1 * (skoXC1 * -1/1024))))))))) <=
     -3628800), inference(arithmetic, [], [refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 *
     (skoXC1 *
      (-453600 +
       skoXC1 *
       (skoXC1 *
        (9450 +
         skoXC1 *
         (skoXC1 *
          (-315/4 +
           skoXC1 *
           (skoXC1 * (45/128 + skoXC1 * (skoXC1 * -1/1024))))))))) <=
     -3628800 |
     skoXC1 *
     (-1814400 +
      skoXC1 *
      (skoXC1 *
       (982800 +
        skoXC1 *
        (skoXC1 *
         (-114345 +
          skoXC1 *
          (skoXC1 *
           (49185/8 +
            skoXC1 *
            (skoXC1 *
             (-49205/256 +
              skoXC1 *
              (skoXC1 *
               (44287/11264 +
                skoXC1 *
                (skoXC1 *
                 (-3793/67584 +
                  skoXC1 *
                  (skoXC1 *
                   (15091/28385280 +
                    skoXC1 *
                    (skoXC1 *
                     (-89/28385280 +
                      skoXC1 *
                      (skoXC1 *
                       (47/4087480320 +
                        skoXC1 *
                        (skoXC1 * -1/40874803200)))))))))))))))))))) <= 0),
    inference(decision, [],
              [refute_0_81, refute_0_73, refute_0_71, refute_0_48,
               refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_89, plain,
    (skoXC1 *
     (skoXC1 *
      (-453600 +
       skoXC1 *
       (skoXC1 *
        (9450 +
         skoXC1 *
         (skoXC1 *
          (-315/4 +
           skoXC1 *
           (skoXC1 * (45/128 + skoXC1 * (skoXC1 * -1/1024))))))))) <=
     -3628800), inference(resolve, [], [refute_0_88, refute_0_87])).

cnf(refute_0_90, plain,
    (-3628800 <
     skoXC1 *
     (skoXC1 *
      (-453600 +
       skoXC1 *
       (skoXC1 *
        (9450 +
         skoXC1 *
         (skoXC1 *
          (-315/4 +
           skoXC1 *
           (skoXC1 * (45/128 + skoXC1 * (skoXC1 * -1/1024)))))))))),
    inference(decision, [],
              [refute_0_81, refute_0_73, refute_0_71, refute_0_48,
               refute_0_42, refute_0_43, refute_0_44, refute_0_9])).

cnf(refute_0_91, plain, ($false),
    inference(resolve, [], [refute_0_89, refute_0_90])).
% SZS output end CNFRefutation for tan-problem-1-sin-weak.tptp
