% SZS output start CNFRefutation for tanh-8-weak2.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(tanh_lower_bound_pos1, axiom,
    (X < 0 | ~ lgen(R, Y, 3 * X / (X ^ 2 + 3)) | lgen(R, Y, tanh(X)))).

cnf(tanh_lower_bound_pos2, axiom,
    (X < 0 |
     ~ lgen(R, Y, 5 * (2 * X ^ 2 + 21) * X / (X ^ 4 + 45 * X ^ 2 + 105)) |
     lgen(R, Y, tanh(X)))).

cnf(sinh_upper_bound_pos1, axiom,
    (X < 0 | 3/2 < X | ~ lgen(R, -4 * X / ((X - 2) * (X + 2)), Y) |
     lgen(R, sinh(X), Y))).

cnf(sinh_upper_bound_pos2, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -24 * X * (X ^ 4 + 70 * X ^ 2 + 600) /
            ((X ^ 3 - 12 * X ^ 2 + 60 * X - 120) *
             (X ^ 3 + 12 * X ^ 2 + 60 * X + 120)), Y) |
     lgen(R, sinh(X), Y))).

cnf(sinh_taylor_lower_bound_pos1, axiom,
    (X < 1 |
     ~ lgen(R, Y,
            1/2 *
            (1 + 1/3 * X + 1/18 * X ^ 2 + 1/162 * X ^ 3 + 1/1944 * X ^ 4 +
             1/29160 * X ^ 5) ^ 3 -
            1152 / (X ^ 3 + 6 * X ^ 2 + 24 * X + 48) ^ 2) |
     lgen(R, Y, sinh(X)))).

cnf(sinh_lower_bound_pos1, axiom,
    (X < 0 | 5/2 < X |
     ~ lgen(R, Y,
            12 * X * (X ^ 2 + 12) /
            ((X ^ 2 - 6 * X + 12) * (X ^ 2 + 6 * X + 12))) |
     lgen(R, Y, sinh(X)))).

fof(tanh_8, conjecture,
    (! [X] :
       ((0 <= X & X <= 6/5) => (1 - tanh(X) ^ 2) * sinh(2 * X) <= 2))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 6/5) => (1 - tanh(X) ^ 2) * sinh(2 * X) <= 2)),
    inference(strip, [], [tanh_8])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 6/5) => (1 - tanh(X) ^ 2) * sinh(2 * X) <= 2)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] : (2 < (1 - tanh(X) ^ 2) * sinh(2 * X) & X <= 6/5 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (2 < (1 - tanh(skoXC1) ^ 2) * sinh(2 * skoXC1) & 0 <= skoXC1 &
     skoXC1 <= 6/5), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (2 < (1 - tanh(skoXC1) ^ 2) * sinh(2 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 6/5),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    ((86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-259200 +
        skoXC1 * (skoXC1 * (-120960 + skoXC1 * (skoXC1 * -6912)))))) |
     3 + skoXC1 * skoXC1 <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-259200 +
        skoXC1 * (skoXC1 * (-120960 + skoXC1 * (skoXC1 * -6912)))))) /
     (3 + skoXC1 * skoXC1) <=
     86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-86400 +
       skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304))))) |
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000041, tanh(X_000040)) | X_000041 <= tanh(X_000040)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (X_000040 < 0 |
     ~ lgen(0, X_000041, 3 * X_000040 / (X_000040 ^ 2 + 3)) |
     lgen(0, X_000041, tanh(X_000040))),
    inference(subst, [], [tanh_lower_bound_pos1])).

cnf(refute_0_4, plain,
    (X_000040 < 0 |
     ~ lgen(0, X_000041, 3 * X_000040 / (X_000040 ^ 2 + 3)) |
     X_000041 <= tanh(X_000040)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (X_000040 * 3 / (3 + X_000040 * X_000040) < X_000041 | X_000040 < 0 |
     X_000041 <= tanh(X_000040)), inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    ((28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     3 + skoXC1 * skoXC1 <= 0 |
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) /
     (3 + skoXC1 * skoXC1) <=
     28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_8, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768))))) |
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_9, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_10, plain,
    (skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) +
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))))) <
     2 *
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) |
     0 <=
     -14400 +
     skoXC1 *
     (skoXC1 *
      (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64))))) |
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) +
      tanh(skoXC1) *
      (tanh(skoXC1) *
       (skoXC1 *
        (28800 +
         skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) <=
     2), inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_11, plain,
    (2 <
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     1 + tanh(skoXC1) * (tanh(skoXC1) * -1) <= 0 |
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) <=
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_12, plain, (2 < (1 - tanh(skoXC1) ^ 2) * sinh(2 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_13, plain,
    (2 < sinh(skoXC1 * 2) * (1 + tanh(skoXC1) * (tanh(skoXC1) * -1))),
    inference(arithmetic, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) < sinh(skoXC1 * 2) |
     1 + tanh(skoXC1) * (tanh(skoXC1) * -1) <= 0),
    inference(split, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) < sinh(skoXC1 * 2) |
     tanh(skoXC1) * (tanh(skoXC1) * -1) <= -1),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (2 < 0 | 1 + tanh(skoXC1) * (tanh(skoXC1) * -1) < 0 |
     0 < 1 + tanh(skoXC1) * (tanh(skoXC1) * -1)),
    inference(split, [], [refute_0_13])).

cnf(refute_0_17, plain,
    (-1 < tanh(skoXC1) * (tanh(skoXC1) * -1) |
     1 < tanh(skoXC1) * tanh(skoXC1)),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    ((864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2592 + skoXC1 * (skoXC1 * -864)))) |
     3 + skoXC1 * skoXC1 <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (-2592 + skoXC1 * (skoXC1 * -864)))) /
     (3 + skoXC1 * skoXC1) <=
     864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_19, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     0 <= skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288))) |
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_20, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_21, plain,
    ((288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     3 + skoXC1 * skoXC1 <= 0 |
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) /
     (3 + skoXC1 * skoXC1) <=
     288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_22, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     0 <= tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96))) |
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_23, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_24, plain,
    (2 * (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) +
     tanh(skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16))) <= 0 |
     (skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) +
      tanh(skoXC1) *
      (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96))))) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <= 2),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_25, plain,
    (2 <
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     0 <= 1 + tanh(skoXC1) * (tanh(skoXC1) * -1) |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <=
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_26, plain,
    (sinh(skoXC1 * 2) < 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     0 <= 1 + tanh(skoXC1) * (tanh(skoXC1) * -1)),
    inference(split, [], [refute_0_13])).

cnf(refute_0_27, plain,
    (sinh(skoXC1 * 2) < 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (~ lgen(0, X_000104, sinh(X_000103)) | X_000104 <= sinh(X_000103)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_29, plain,
    (5/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, X_000104,
            12 * X_000103 * (X_000103 ^ 2 + 12) /
            ((X_000103 ^ 2 - 6 * X_000103 + 12) *
             (X_000103 ^ 2 + 6 * X_000103 + 12))) |
     lgen(0, X_000104, sinh(X_000103))),
    inference(subst, [], [sinh_lower_bound_pos1])).

cnf(refute_0_30, plain,
    (5/2 < X_000103 | X_000103 < 0 |
     ~ lgen(0, X_000104,
            12 * X_000103 * (X_000103 ^ 2 + 12) /
            ((X_000103 ^ 2 - 6 * X_000103 + 12) *
             (X_000103 ^ 2 + 6 * X_000103 + 12))) |
     X_000104 <= sinh(X_000103)),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain,
    (5/2 < X_000103 |
     X_000103 * (144 + X_000103 * (X_000103 * 12)) /
     (144 + X_000103 * (X_000103 * (-12 + X_000103 * X_000103))) <
     X_000104 | X_000103 < 0 | X_000104 <= sinh(X_000103)),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (skoXC1 * 2 < 0 | 5/2 < skoXC1 * 2 |
     skoXC1 * 2 * (144 + skoXC1 * 2 * (skoXC1 * 2 * 12)) /
     (144 +
      skoXC1 * 2 * (skoXC1 * 2 * (-12 + skoXC1 * 2 * (skoXC1 * 2)))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <= sinh(skoXC1 * 2)),
    inference(subst, [], [refute_0_31])).

cnf(refute_0_33, plain,
    (skoXC1 * 2 < 0 | 5/2 < skoXC1 * 2 |
     skoXC1 * 2 * (144 + skoXC1 * 2 * (skoXC1 * 2 * 12)) /
     (144 +
      skoXC1 * 2 * (skoXC1 * 2 * (-12 + skoXC1 * 2 * (skoXC1 * 2)))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_32, refute_0_27])).

cnf(refute_0_34, plain,
    (skoXC1 < 0 |
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) | 5/4 < skoXC1 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    ((13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-41472 +
        skoXC1 *
        (-82944 +
         skoXC1 *
         (-110592 +
          skoXC1 *
          (-109728 +
           skoXC1 *
           (-428544/5 +
            skoXC1 *
            (-1472672/27 +
             skoXC1 *
             (-260416/9 +
              skoXC1 *
              (-1053952/81 +
               skoXC1 *
               (-164844928/32805 +
                skoXC1 *
                (-276173504/164025 +
                 skoXC1 *
                 (-26964608/54675 +
                  skoXC1 *
                  (-561814144/4428675 +
                   skoXC1 *
                   (-127113728/4428675 +
                    skoXC1 *
                    (-75796864/13286025 +
                     skoXC1 *
                     (-592884992/597871125 +
                      skoXC1 *
                      (-89319424/597871125 +
                       skoXC1 *
                       (-11482624/597871125 +
                        skoXC1 *
                        (-3684608/1793613375 +
                         skoXC1 *
                         (-313856/1793613375 +
                          skoXC1 *
                          (-19456/1793613375 +
                           skoXC1 *
                           -2048/5380840125)))))))))))))))))))))) |
     3 + skoXC1 * skoXC1 <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-41472 +
        skoXC1 *
        (-82944 +
         skoXC1 *
         (-110592 +
          skoXC1 *
          (-109728 +
           skoXC1 *
           (-428544/5 +
            skoXC1 *
            (-1472672/27 +
             skoXC1 *
             (-260416/9 +
              skoXC1 *
              (-1053952/81 +
               skoXC1 *
               (-164844928/32805 +
                skoXC1 *
                (-276173504/164025 +
                 skoXC1 *
                 (-26964608/54675 +
                  skoXC1 *
                  (-561814144/4428675 +
                   skoXC1 *
                   (-127113728/4428675 +
                    skoXC1 *
                    (-75796864/13286025 +
                     skoXC1 *
                     (-592884992/597871125 +
                      skoXC1 *
                      (-89319424/597871125 +
                       skoXC1 *
                       (-11482624/597871125 +
                        skoXC1 *
                        (-3684608/1793613375 +
                         skoXC1 *
                         (-313856/1793613375 +
                          skoXC1 *
                          (-19456/1793613375 +
                           skoXC1 *
                           -2048/5380840125)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) <=
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125))))))))))))))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_36, plain,
    (13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-13824 +
       skoXC1 *
       (-27648 +
        skoXC1 *
        (-36864 +
         skoXC1 *
         (-36576 +
          skoXC1 *
          (-142848/5 +
           skoXC1 *
           (-1472672/81 +
            skoXC1 *
            (-260416/27 +
             skoXC1 *
             (-1053952/243 +
              skoXC1 *
              (-164844928/98415 +
               skoXC1 *
               (-276173504/492075 +
                skoXC1 *
                (-26964608/164025 +
                 skoXC1 *
                 (-561814144/13286025 +
                  skoXC1 *
                  (-127113728/13286025 +
                   skoXC1 *
                   (-75796864/39858075 +
                    skoXC1 *
                    (-592884992/1793613375 +
                     skoXC1 *
                     (-89319424/1793613375 +
                      skoXC1 *
                      (-11482624/1793613375 +
                       skoXC1 *
                       (-3684608/5380840125 +
                        skoXC1 *
                        (-313856/5380840125 +
                         skoXC1 *
                         (-19456/5380840125 +
                          skoXC1 * -2048/16142520375))))))))))))))))))))) |
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_37, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_38, plain,
    ((4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     3 + skoXC1 * skoXC1 <= 0 |
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) <=
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125))))))))))))))))))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_39, plain,
    (4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-4608 +
       skoXC1 *
       (-9216 +
        skoXC1 *
        (-12288 +
         skoXC1 *
         (-12192 +
          skoXC1 *
          (-47616/5 +
           skoXC1 *
           (-1472672/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) |
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_40, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_41, plain,
    (2 *
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) <
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) +
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     2304 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) <= 0 |
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125)))))))))))))))))))) +
      tanh(skoXC1) *
      (tanh(skoXC1) *
       (skoXC1 *
        (-4608 +
         skoXC1 *
         (-9216 +
          skoXC1 *
          (-12288 +
           skoXC1 *
           (-12192 +
            skoXC1 *
            (-47616/5 +
             skoXC1 *
             (-1472672/243 +
              skoXC1 *
              (-260416/81 +
               skoXC1 *
               (-1053952/729 +
                skoXC1 *
                (-164844928/295245 +
                 skoXC1 *
                 (-276173504/1476225 +
                  skoXC1 *
                  (-26964608/492075 +
                   skoXC1 *
                   (-561814144/39858075 +
                    skoXC1 *
                    (-127113728/39858075 +
                     skoXC1 *
                     (-75796864/119574225 +
                      skoXC1 *
                      (-592884992/5380840125 +
                       skoXC1 *
                       (-89319424/5380840125 +
                        skoXC1 *
                        (-11482624/5380840125 +
                         skoXC1 *
                         (-3684608/16142520375 +
                          skoXC1 *
                          (-313856/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) <= 2),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_42, plain,
    (2 <
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     0 <= 1 + tanh(skoXC1) * (tanh(skoXC1) * -1) |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <=
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_43, plain,
    (~ lgen(0, X_000069, sinh(X_000068)) | X_000069 <= sinh(X_000068)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_44, plain,
    (X_000068 < 1 |
     ~ lgen(0, X_000069,
            1/2 *
            (1 + 1/3 * X_000068 + 1/18 * X_000068 ^ 2 +
             1/162 * X_000068 ^ 3 + 1/1944 * X_000068 ^ 4 +
             1/29160 * X_000068 ^ 5) ^ 3 -
            1152 /
            (X_000068 ^ 3 + 6 * X_000068 ^ 2 + 24 * X_000068 + 48) ^ 2) |
     lgen(0, X_000069, sinh(X_000068))),
    inference(subst, [], [sinh_taylor_lower_bound_pos1])).

cnf(refute_0_45, plain,
    (X_000068 < 1 |
     ~ lgen(0, X_000069,
            1/2 *
            (1 + 1/3 * X_000068 + 1/18 * X_000068 ^ 2 +
             1/162 * X_000068 ^ 3 + 1/1944 * X_000068 ^ 4 +
             1/29160 * X_000068 ^ 5) ^ 3 -
            1152 /
            (X_000068 ^ 3 + 6 * X_000068 ^ 2 + 24 * X_000068 + 48) ^ 2) |
     X_000069 <= sinh(X_000068)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (1/2 +
     X_000068 *
     (1/2 +
      X_000068 *
      (1/4 +
       X_000068 *
       (1/12 +
        X_000068 *
        (1/48 +
         X_000068 *
         (1/240 +
          X_000068 *
          (121/174960 +
           X_000068 *
           (17/174960 +
            X_000068 *
            (49/4199040 +
             X_000068 *
             (409/340122240 +
              X_000068 *
              (361/3401222400 +
               X_000068 *
               (1/125971200 +
                X_000068 *
                (181/367332019200 +
                 X_000068 *
                 (1/40814668800 +
                  X_000068 *
                  (1/1101996057600 +
                   X_000068 * 1/49589822592000)))))))))))))) +
     1152 /
     (2304 +
      X_000068 *
      (2304 +
       X_000068 *
       (1152 +
        X_000068 *
        (384 + X_000068 * (84 + X_000068 * (12 + X_000068)))))) * -1 <
     X_000069 | X_000068 < 1 | X_000069 <= sinh(X_000068)),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (skoXC1 * 2 < 1 |
     1/2 +
     skoXC1 * 2 *
     (1/2 +
      skoXC1 * 2 *
      (1/4 +
       skoXC1 * 2 *
       (1/12 +
        skoXC1 * 2 *
        (1/48 +
         skoXC1 * 2 *
         (1/240 +
          skoXC1 * 2 *
          (121/174960 +
           skoXC1 * 2 *
           (17/174960 +
            skoXC1 * 2 *
            (49/4199040 +
             skoXC1 * 2 *
             (409/340122240 +
              skoXC1 * 2 *
              (361/3401222400 +
               skoXC1 * 2 *
               (1/125971200 +
                skoXC1 * 2 *
                (181/367332019200 +
                 skoXC1 * 2 *
                 (1/40814668800 +
                  skoXC1 * 2 *
                  (1/1101996057600 +
                   skoXC1 * 2 * 1/49589822592000)))))))))))))) +
     1152 /
     (2304 +
      skoXC1 * 2 *
      (2304 +
       skoXC1 * 2 *
       (1152 +
        skoXC1 * 2 *
        (384 + skoXC1 * 2 * (84 + skoXC1 * 2 * (12 + skoXC1 * 2)))))) *
     -1 < 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <= sinh(skoXC1 * 2)),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (skoXC1 * 2 < 1 |
     1/2 +
     skoXC1 * 2 *
     (1/2 +
      skoXC1 * 2 *
      (1/4 +
       skoXC1 * 2 *
       (1/12 +
        skoXC1 * 2 *
        (1/48 +
         skoXC1 * 2 *
         (1/240 +
          skoXC1 * 2 *
          (121/174960 +
           skoXC1 * 2 *
           (17/174960 +
            skoXC1 * 2 *
            (49/4199040 +
             skoXC1 * 2 *
             (409/340122240 +
              skoXC1 * 2 *
              (361/3401222400 +
               skoXC1 * 2 *
               (1/125971200 +
                skoXC1 * 2 *
                (181/367332019200 +
                 skoXC1 * 2 *
                 (1/40814668800 +
                  skoXC1 * 2 *
                  (1/1101996057600 +
                   skoXC1 * 2 * 1/49589822592000)))))))))))))) +
     1152 /
     (2304 +
      skoXC1 * 2 *
      (2304 +
       skoXC1 * 2 *
       (1152 +
        skoXC1 * 2 *
        (384 + skoXC1 * 2 * (84 + skoXC1 * 2 * (12 + skoXC1 * 2)))))) *
     -1 < 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_47, refute_0_27])).

cnf(refute_0_49, plain,
    (skoXC1 < 1/2 |
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 |
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 * (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) =
     -2304), inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_51, plain, (skoXC1 <= 6/5),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_52, plain,
    (1 < tanh(skoXC1) * tanh(skoXC1) | tanh(skoXC1) * tanh(skoXC1) <= 1),
    introduced(tautology, [assume])).

cnf(refute_0_53, plain,
    (1 < 0 | tanh(skoXC1) < 0 | 0 < tanh(skoXC1) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(split, [], [refute_0_52])).

cnf(refute_0_54, plain,
    (skoXC1 < 0 | 0 < skoXC1 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_53])).

cnf(refute_0_55, plain, (0 < skoXC1 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_50, refute_0_54])).

cnf(refute_0_56, plain,
    (skoXC1 < 1/2 |
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) != -2304 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 < 1/2 |
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_49, refute_0_56])).

cnf(refute_0_58, plain,
    (skoXC1 < 1/2 |
     2 <
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     0 <= 1 + tanh(skoXC1) * (tanh(skoXC1) * -1) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_42, refute_0_57])).

cnf(refute_0_59, plain,
    (skoXC1 < 1/2 |
     2 <
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125)))))))))))))))))))) +
      tanh(skoXC1) *
      (tanh(skoXC1) *
       (skoXC1 *
        (-4608 +
         skoXC1 *
         (-9216 +
          skoXC1 *
          (-12288 +
           skoXC1 *
           (-12192 +
            skoXC1 *
            (-47616/5 +
             skoXC1 *
             (-1472672/243 +
              skoXC1 *
              (-260416/81 +
               skoXC1 *
               (-1053952/729 +
                skoXC1 *
                (-164844928/295245 +
                 skoXC1 *
                 (-276173504/1476225 +
                  skoXC1 *
                  (-26964608/492075 +
                   skoXC1 *
                   (-561814144/39858075 +
                    skoXC1 *
                    (-127113728/39858075 +
                     skoXC1 *
                     (-75796864/119574225 +
                      skoXC1 *
                      (-592884992/5380840125 +
                       skoXC1 *
                       (-89319424/5380840125 +
                        skoXC1 *
                        (-11482624/5380840125 +
                         skoXC1 *
                         (-3684608/16142520375 +
                          skoXC1 *
                          (-313856/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_58])).

cnf(refute_0_60, plain,
    (skoXC1 < 1/2 |
     2 *
     (2304 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64)))))) <
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) +
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 |
     2304 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) <= 0),
    inference(resolve, [], [refute_0_41, refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 < 1/2 |
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) <= -2304 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 < 1/2 |
     -2304 <
     skoXC1 *
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 * (3072 + skoXC1 * (1344 + skoXC1 * (384 + skoXC1 * 64))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_55, refute_0_51, refute_0_50])).

cnf(refute_0_63, plain,
    (skoXC1 < 1/2 |
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_61, refute_0_62])).

cnf(refute_0_64, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-4608 +
       skoXC1 *
       (-9216 +
        skoXC1 *
        (-12288 +
         skoXC1 *
         (-12192 +
          skoXC1 *
          (-47616/5 +
           skoXC1 *
           (-1472672/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(split, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) <=
     0 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_64])).

cnf(refute_0_66, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) |
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_67, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) |
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) <=
     0 |
     tanh(skoXC1) <=
     0 /
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125)))))))))))))))))))))),
    inference(split, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) |
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) <=
     0 | skoXC1 <= 0), inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) |
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) <=
     0 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_68, refute_0_55])).

cnf(refute_0_70, plain,
    (0 <
     skoXC1 *
     (4608 +
      skoXC1 *
      (9216 +
       skoXC1 *
       (12288 +
        skoXC1 *
        (12192 +
         skoXC1 *
         (47616/5 +
          skoXC1 *
          (1472672/243 +
           skoXC1 *
           (260416/81 +
            skoXC1 *
            (1053952/729 +
             skoXC1 *
             (164844928/295245 +
              skoXC1 *
              (276173504/1476225 +
               skoXC1 *
               (26964608/492075 +
                skoXC1 *
                (561814144/39858075 +
                 skoXC1 *
                 (127113728/39858075 +
                  skoXC1 *
                  (75796864/119574225 +
                   skoXC1 *
                   (592884992/5380840125 +
                    skoXC1 *
                    (89319424/5380840125 +
                     skoXC1 *
                     (11482624/5380840125 +
                      skoXC1 *
                      (3684608/16142520375 +
                       skoXC1 *
                       (313856/16142520375 +
                        skoXC1 *
                        (19456/16142520375 +
                         skoXC1 * 2048/48427561125)))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_71, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_65, refute_0_71])).

cnf(refute_0_73, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_40, refute_0_72])).

cnf(refute_0_74, plain,
    (skoXC1 < 1/2 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_50, refute_0_73])).

cnf(refute_0_75, plain,
    (skoXC1 < 1/2 |
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-4608 +
        skoXC1 *
        (-9216 +
         skoXC1 *
         (-12288 +
          skoXC1 *
          (-12192 +
           skoXC1 *
           (-47616/5 +
            skoXC1 *
            (-1472672/243 +
             skoXC1 *
             (-260416/81 +
              skoXC1 *
              (-1053952/729 +
               skoXC1 *
               (-164844928/295245 +
                skoXC1 *
                (-276173504/1476225 +
                 skoXC1 *
                 (-26964608/492075 +
                  skoXC1 *
                  (-561814144/39858075 +
                   skoXC1 *
                   (-127113728/39858075 +
                    skoXC1 *
                    (-75796864/119574225 +
                     skoXC1 *
                     (-592884992/5380840125 +
                      skoXC1 *
                      (-89319424/5380840125 +
                       skoXC1 *
                       (-11482624/5380840125 +
                        skoXC1 *
                        (-3684608/16142520375 +
                         skoXC1 *
                         (-313856/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-4608 +
       skoXC1 *
       (-9216 +
        skoXC1 *
        (-12288 +
         skoXC1 *
         (-12192 +
          skoXC1 *
          (-47616/5 +
           skoXC1 *
           (-1472672/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_39, refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 < 1/2 |
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) |
     tanh(skoXC1) *
     (skoXC1 *
      (4608 +
       skoXC1 *
       (9216 +
        skoXC1 *
        (12288 +
         skoXC1 *
         (12192 +
          skoXC1 *
          (47616/5 +
           skoXC1 *
           (1472672/243 +
            skoXC1 *
            (260416/81 +
             skoXC1 *
             (1053952/729 +
              skoXC1 *
              (164844928/295245 +
               skoXC1 *
               (276173504/1476225 +
                skoXC1 *
                (26964608/492075 +
                 skoXC1 *
                 (561814144/39858075 +
                  skoXC1 *
                  (127113728/39858075 +
                   skoXC1 *
                   (75796864/119574225 +
                    skoXC1 *
                    (592884992/5380840125 +
                     skoXC1 *
                     (89319424/5380840125 +
                      skoXC1 *
                      (11482624/5380840125 +
                       skoXC1 *
                       (3684608/16142520375 +
                        skoXC1 *
                        (313856/16142520375 +
                         skoXC1 *
                         (19456/16142520375 +
                          skoXC1 * 2048/48427561125))))))))))))))))))))) <=
     0 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 < 1/2 |
     4608 +
     skoXC1 *
     (4608 +
      skoXC1 *
      (skoXC1 *
       (-6144 +
        skoXC1 *
        (-9504 +
         skoXC1 *
         (-43776/5 +
          skoXC1 *
          (-1441568/243 +
           skoXC1 *
           (-260416/81 +
            skoXC1 *
            (-1053952/729 +
             skoXC1 *
             (-164844928/295245 +
              skoXC1 *
              (-276173504/1476225 +
               skoXC1 *
               (-26964608/492075 +
                skoXC1 *
                (-561814144/39858075 +
                 skoXC1 *
                 (-127113728/39858075 +
                  skoXC1 *
                  (-75796864/119574225 +
                   skoXC1 *
                   (-592884992/5380840125 +
                    skoXC1 *
                    (-89319424/5380840125 +
                     skoXC1 *
                     (-11482624/5380840125 +
                      skoXC1 *
                      (-3684608/16142520375 +
                       skoXC1 *
                       (-313856/16142520375 +
                        skoXC1 *
                        (-19456/16142520375 +
                         skoXC1 * -2048/48427561125)))))))))))))))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_76, refute_0_71])).

cnf(refute_0_78, plain,
    (skoXC1 < 1/2 |
     (4608 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (skoXC1 *
        (-6144 +
         skoXC1 *
         (-9504 +
          skoXC1 *
          (-43776/5 +
           skoXC1 *
           (-1441568/243 +
            skoXC1 *
            (-260416/81 +
             skoXC1 *
             (-1053952/729 +
              skoXC1 *
              (-164844928/295245 +
               skoXC1 *
               (-276173504/1476225 +
                skoXC1 *
                (-26964608/492075 +
                 skoXC1 *
                 (-561814144/39858075 +
                  skoXC1 *
                  (-127113728/39858075 +
                   skoXC1 *
                   (-75796864/119574225 +
                    skoXC1 *
                    (-592884992/5380840125 +
                     skoXC1 *
                     (-89319424/5380840125 +
                      skoXC1 *
                      (-11482624/5380840125 +
                       skoXC1 *
                       (-3684608/16142520375 +
                        skoXC1 *
                        (-313856/16142520375 +
                         skoXC1 *
                         (-19456/16142520375 +
                          skoXC1 * -2048/48427561125))))))))))))))))))))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_38, refute_0_77])).

cnf(refute_0_79, plain,
    (skoXC1 < 1/2 |
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_78])).

cnf(refute_0_80, plain,
    (skoXC1 < 1/2 | -3 < skoXC1 * skoXC1 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_81, plain,
    (skoXC1 < 1/2 |
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-13824 +
       skoXC1 *
       (-27648 +
        skoXC1 *
        (-36864 +
         skoXC1 *
         (-36576 +
          skoXC1 *
          (-142848/5 +
           skoXC1 *
           (-1472672/81 +
            skoXC1 *
            (-260416/27 +
             skoXC1 *
             (-1053952/243 +
              skoXC1 *
              (-164844928/98415 +
               skoXC1 *
               (-276173504/492075 +
                skoXC1 *
                (-26964608/164025 +
                 skoXC1 *
                 (-561814144/13286025 +
                  skoXC1 *
                  (-127113728/13286025 +
                   skoXC1 *
                   (-75796864/39858075 +
                    skoXC1 *
                    (-592884992/1793613375 +
                     skoXC1 *
                     (-89319424/1793613375 +
                      skoXC1 *
                      (-11482624/1793613375 +
                       skoXC1 *
                       (-3684608/5380840125 +
                        skoXC1 *
                        (-313856/5380840125 +
                         skoXC1 *
                         (-19456/5380840125 +
                          skoXC1 * -2048/16142520375))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(split, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (13824 +
       skoXC1 *
       (27648 +
        skoXC1 *
        (36864 +
         skoXC1 *
         (36576 +
          skoXC1 *
          (142848/5 +
           skoXC1 *
           (1472672/81 +
            skoXC1 *
            (260416/27 +
             skoXC1 *
             (1053952/243 +
              skoXC1 *
              (164844928/98415 +
               skoXC1 *
               (276173504/492075 +
                skoXC1 *
                (26964608/164025 +
                 skoXC1 *
                 (561814144/13286025 +
                  skoXC1 *
                  (127113728/13286025 +
                   skoXC1 *
                   (75796864/39858075 +
                    skoXC1 *
                    (592884992/1793613375 +
                     skoXC1 *
                     (89319424/1793613375 +
                      skoXC1 *
                      (11482624/1793613375 +
                       skoXC1 *
                       (3684608/5380840125 +
                        skoXC1 *
                        (313856/5380840125 +
                         skoXC1 *
                         (19456/5380840125 +
                          skoXC1 * 2048/16142520375))))))))))))))))))))) <=
     0 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 < 1/2 |
     0 <
     skoXC1 *
     (skoXC1 *
      (13824 +
       skoXC1 *
       (27648 +
        skoXC1 *
        (36864 +
         skoXC1 *
         (36576 +
          skoXC1 *
          (142848/5 +
           skoXC1 *
           (1472672/81 +
            skoXC1 *
            (260416/27 +
             skoXC1 *
             (1053952/243 +
              skoXC1 *
              (164844928/98415 +
               skoXC1 *
               (276173504/492075 +
                skoXC1 *
                (26964608/164025 +
                 skoXC1 *
                 (561814144/13286025 +
                  skoXC1 *
                  (127113728/13286025 +
                   skoXC1 *
                   (75796864/39858075 +
                    skoXC1 *
                    (592884992/1793613375 +
                     skoXC1 *
                     (89319424/1793613375 +
                      skoXC1 *
                      (11482624/1793613375 +
                       skoXC1 *
                       (3684608/5380840125 +
                        skoXC1 *
                        (313856/5380840125 +
                         skoXC1 *
                         (19456/5380840125 +
                          skoXC1 * 2048/16142520375))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_85, plain,
    (skoXC1 < 1/2 |
     tanh(skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (skoXC1 < 0 | skoXC1 < 1/2 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_37, refute_0_85])).

cnf(refute_0_87, plain,
    (skoXC1 < 1/2 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) /
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_50, refute_0_86])).

cnf(refute_0_88, plain,
    (skoXC1 < 1/2 |
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-13824 +
        skoXC1 *
        (-27648 +
         skoXC1 *
         (-36864 +
          skoXC1 *
          (-36576 +
           skoXC1 *
           (-142848/5 +
            skoXC1 *
            (-1472672/81 +
             skoXC1 *
             (-260416/27 +
              skoXC1 *
              (-1053952/243 +
               skoXC1 *
               (-164844928/98415 +
                skoXC1 *
                (-276173504/492075 +
                 skoXC1 *
                 (-26964608/164025 +
                  skoXC1 *
                  (-561814144/13286025 +
                   skoXC1 *
                   (-127113728/13286025 +
                    skoXC1 *
                    (-75796864/39858075 +
                     skoXC1 *
                     (-592884992/1793613375 +
                      skoXC1 *
                      (-89319424/1793613375 +
                       skoXC1 *
                       (-11482624/1793613375 +
                        skoXC1 *
                        (-3684608/5380840125 +
                         skoXC1 *
                         (-313856/5380840125 +
                          skoXC1 *
                          (-19456/5380840125 +
                           skoXC1 *
                           -2048/16142520375)))))))))))))))))))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-13824 +
       skoXC1 *
       (-27648 +
        skoXC1 *
        (-36864 +
         skoXC1 *
         (-36576 +
          skoXC1 *
          (-142848/5 +
           skoXC1 *
           (-1472672/81 +
            skoXC1 *
            (-260416/27 +
             skoXC1 *
             (-1053952/243 +
              skoXC1 *
              (-164844928/98415 +
               skoXC1 *
               (-276173504/492075 +
                skoXC1 *
                (-26964608/164025 +
                 skoXC1 *
                 (-561814144/13286025 +
                  skoXC1 *
                  (-127113728/13286025 +
                   skoXC1 *
                   (-75796864/39858075 +
                    skoXC1 *
                    (-592884992/1793613375 +
                     skoXC1 *
                     (-89319424/1793613375 +
                      skoXC1 *
                      (-11482624/1793613375 +
                       skoXC1 *
                       (-3684608/5380840125 +
                        skoXC1 *
                        (-313856/5380840125 +
                         skoXC1 *
                         (-19456/5380840125 +
                          skoXC1 * -2048/16142520375))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_36, refute_0_87])).

cnf(refute_0_89, plain,
    (skoXC1 < 1/2 |
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-41472 +
        skoXC1 *
        (-82944 +
         skoXC1 *
         (-110592 +
          skoXC1 *
          (-109728 +
           skoXC1 *
           (-428544/5 +
            skoXC1 *
            (-1472672/27 +
             skoXC1 *
             (-260416/9 +
              skoXC1 *
              (-1053952/81 +
               skoXC1 *
               (-164844928/32805 +
                skoXC1 *
                (-276173504/164025 +
                 skoXC1 *
                 (-26964608/54675 +
                  skoXC1 *
                  (-561814144/4428675 +
                   skoXC1 *
                   (-127113728/4428675 +
                    skoXC1 *
                    (-75796864/13286025 +
                     skoXC1 *
                     (-592884992/597871125 +
                      skoXC1 *
                      (-89319424/597871125 +
                       skoXC1 *
                       (-11482624/597871125 +
                        skoXC1 *
                        (-3684608/1793613375 +
                         skoXC1 *
                         (-313856/1793613375 +
                          skoXC1 *
                          (-19456/1793613375 +
                           skoXC1 *
                           -2048/5380840125)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (13824 +
       skoXC1 *
       (27648 +
        skoXC1 *
        (36864 +
         skoXC1 *
         (36576 +
          skoXC1 *
          (142848/5 +
           skoXC1 *
           (1472672/81 +
            skoXC1 *
            (260416/27 +
             skoXC1 *
             (1053952/243 +
              skoXC1 *
              (164844928/98415 +
               skoXC1 *
               (276173504/492075 +
                skoXC1 *
                (26964608/164025 +
                 skoXC1 *
                 (561814144/13286025 +
                  skoXC1 *
                  (127113728/13286025 +
                   skoXC1 *
                   (75796864/39858075 +
                    skoXC1 *
                    (592884992/1793613375 +
                     skoXC1 *
                     (89319424/1793613375 +
                      skoXC1 *
                      (11482624/1793613375 +
                       skoXC1 *
                       (3684608/5380840125 +
                        skoXC1 *
                        (313856/5380840125 +
                         skoXC1 *
                         (19456/5380840125 +
                          skoXC1 * 2048/16142520375))))))))))))))))))))) <=
     0 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (skoXC1 < 1/2 |
     13824 +
     skoXC1 *
     (13824 +
      skoXC1 *
      (4608 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-28512 +
         skoXC1 *
         (-162048/5 +
          skoXC1 *
          (-2211392/81 +
           skoXC1 *
           (-2484032/135 +
            skoXC1 *
            (-277280/27 +
             skoXC1 *
             (-481250368/98415 +
              skoXC1 *
              (-987591104/492075 +
               skoXC1 *
               (-1066906112/1476225 +
                skoXC1 *
                (-609475136/2657205 +
                 skoXC1 *
                 (-855158144/13286025 +
                  skoXC1 *
                  (-637611008/39858075 +
                   skoXC1 *
                   (-6313002752/1793613375 +
                    skoXC1 *
                    (-1226272384/1793613375 +
                     skoXC1 *
                     (-627332864/5380840125 +
                      skoXC1 *
                      (-31001344/1793613375 +
                       skoXC1 *
                       (-262144/119574225 +
                        skoXC1 *
                        (-3742976/16142520375 +
                         skoXC1 *
                         (-315904/16142520375 +
                          skoXC1 *
                          (-19456/16142520375 +
                           skoXC1 *
                           -2048/48427561125)))))))))))))))))))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-41472 +
        skoXC1 *
        (-82944 +
         skoXC1 *
         (-110592 +
          skoXC1 *
          (-109728 +
           skoXC1 *
           (-428544/5 +
            skoXC1 *
            (-1472672/27 +
             skoXC1 *
             (-260416/9 +
              skoXC1 *
              (-1053952/81 +
               skoXC1 *
               (-164844928/32805 +
                skoXC1 *
                (-276173504/164025 +
                 skoXC1 *
                 (-26964608/54675 +
                  skoXC1 *
                  (-561814144/4428675 +
                   skoXC1 *
                   (-127113728/4428675 +
                    skoXC1 *
                    (-75796864/13286025 +
                     skoXC1 *
                     (-592884992/597871125 +
                      skoXC1 *
                      (-89319424/597871125 +
                       skoXC1 *
                       (-11482624/597871125 +
                        skoXC1 *
                        (-3684608/1793613375 +
                         skoXC1 *
                         (-313856/1793613375 +
                          skoXC1 *
                          (-19456/1793613375 +
                           skoXC1 *
                           -2048/5380840125)))))))))))))))))))))) /
     (3 + skoXC1 * skoXC1) | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_89, refute_0_84])).

cnf(refute_0_91, plain,
    (skoXC1 < 1/2 |
     (13824 +
      skoXC1 *
      (13824 +
       skoXC1 *
       (4608 +
        skoXC1 *
        (-13824 +
         skoXC1 *
         (-28512 +
          skoXC1 *
          (-162048/5 +
           skoXC1 *
           (-2211392/81 +
            skoXC1 *
            (-2484032/135 +
             skoXC1 *
             (-277280/27 +
              skoXC1 *
              (-481250368/98415 +
               skoXC1 *
               (-987591104/492075 +
                skoXC1 *
                (-1066906112/1476225 +
                 skoXC1 *
                 (-609475136/2657205 +
                  skoXC1 *
                  (-855158144/13286025 +
                   skoXC1 *
                   (-637611008/39858075 +
                    skoXC1 *
                    (-6313002752/1793613375 +
                     skoXC1 *
                     (-1226272384/1793613375 +
                      skoXC1 *
                      (-627332864/5380840125 +
                       skoXC1 *
                       (-31001344/1793613375 +
                        skoXC1 *
                        (-262144/119574225 +
                         skoXC1 *
                         (-3742976/16142520375 +
                          skoXC1 *
                          (-315904/16142520375 +
                           skoXC1 *
                           (-19456/16142520375 +
                            skoXC1 *
                            -2048/48427561125))))))))))))))))))))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-41472 +
        skoXC1 *
        (-82944 +
         skoXC1 *
         (-110592 +
          skoXC1 *
          (-109728 +
           skoXC1 *
           (-428544/5 +
            skoXC1 *
            (-1472672/27 +
             skoXC1 *
             (-260416/9 +
              skoXC1 *
              (-1053952/81 +
               skoXC1 *
               (-164844928/32805 +
                skoXC1 *
                (-276173504/164025 +
                 skoXC1 *
                 (-26964608/54675 +
                  skoXC1 *
                  (-561814144/4428675 +
                   skoXC1 *
                   (-127113728/4428675 +
                    skoXC1 *
                    (-75796864/13286025 +
                     skoXC1 *
                     (-592884992/597871125 +
                      skoXC1 *
                      (-89319424/597871125 +
                       skoXC1 *
                       (-11482624/597871125 +
                        skoXC1 *
                        (-3684608/1793613375 +
                         skoXC1 *
                         (-313856/1793613375 +
                          skoXC1 *
                          (-19456/1793613375 +
                           skoXC1 *
                           -2048/5380840125)))))))))))))))))))))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_35, refute_0_90])).

cnf(refute_0_92, plain,
    (skoXC1 < 1/2 |
     41472 <
     skoXC1 *
     (-41472 +
      skoXC1 *
      (-27648 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-2016 +
         skoXC1 *
         (2304/5 +
          skoXC1 *
          (18560/27 +
           skoXC1 *
           (85568/45 +
            skoXC1 *
            (288896/81 +
             skoXC1 *
             (135653824/32805 +
              skoXC1 *
              (537814304/164025 +
               skoXC1 *
               (1000484032/492075 +
                skoXC1 *
                (4479011008/4428675 +
                 skoXC1 *
                 (1871743232/4428675 +
                  skoXC1 *
                  (222171584/1476225 +
                   skoXC1 *
                   (27634765952/597871125 +
                    skoXC1 *
                    (7379578624/597871125 +
                     skoXC1 *
                     (1032336128/358722675 +
                      skoXC1 *
                      (1051318144/1793613375 +
                       skoXC1 *
                       (559378688/5380840125 +
                        skoXC1 *
                        (85693184/5380840125 +
                         skoXC1 *
                         (11170816/5380840125 +
                          skoXC1 *
                          (725248/3228504075 +
                           skoXC1 *
                           (103936/5380840125 +
                            skoXC1 *
                            (19456/16142520375 +
                             skoXC1 *
                             2048/48427561125)))))))))))))))))))))))) |
     skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_91])).

cnf(refute_0_93, plain,
    (skoXC1 < 1/2 |
     skoXC1 *
     (-41472 +
      skoXC1 *
      (-27648 +
       skoXC1 *
       (-13824 +
        skoXC1 *
        (-2016 +
         skoXC1 *
         (2304/5 +
          skoXC1 *
          (18560/27 +
           skoXC1 *
           (85568/45 +
            skoXC1 *
            (288896/81 +
             skoXC1 *
             (135653824/32805 +
              skoXC1 *
              (537814304/164025 +
               skoXC1 *
               (1000484032/492075 +
                skoXC1 *
                (4479011008/4428675 +
                 skoXC1 *
                 (1871743232/4428675 +
                  skoXC1 *
                  (222171584/1476225 +
                   skoXC1 *
                   (27634765952/597871125 +
                    skoXC1 *
                    (7379578624/597871125 +
                     skoXC1 *
                     (1032336128/358722675 +
                      skoXC1 *
                      (1051318144/1793613375 +
                       skoXC1 *
                       (559378688/5380840125 +
                        skoXC1 *
                        (85693184/5380840125 +
                         skoXC1 *
                         (11170816/5380840125 +
                          skoXC1 *
                          (725248/3228504075 +
                           skoXC1 *
                           (103936/5380840125 +
                            skoXC1 *
                            (19456/16142520375 +
                             skoXC1 *
                             2048/48427561125)))))))))))))))))))))))) <=
     41472 | skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [], [refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_94, plain,
    (skoXC1 < 1/2 | skoXC1 * skoXC1 <= -3 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_93, refute_0_92])).

cnf(refute_0_95, plain, (skoXC1 < 1/2 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_94, refute_0_80])).

cnf(refute_0_96, plain,
    (5/4 < skoXC1 | 0 <= skoXC1 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_97, plain,
    (skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) | 5/4 < skoXC1 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_96, refute_0_34])).

cnf(refute_0_98, plain, (tanh(skoXC1) * tanh(skoXC1) <= 1 | skoXC1 <= 5/4),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_99, plain,
    (skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_98, refute_0_97])).

cnf(refute_0_100, plain,
    (2 <
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     0 <= 1 + tanh(skoXC1) * (tanh(skoXC1) * -1) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_25, refute_0_99])).

cnf(refute_0_101, plain,
    (2 <
     (skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) +
      tanh(skoXC1) *
      (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96))))) /
     (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (2 * (144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16)))) <
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) +
     tanh(skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 |
     144 + skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16))) <= 0),
    inference(resolve, [], [refute_0_24, refute_0_101])).

cnf(refute_0_103, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     tanh(skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16))) <= -144 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (-144 < skoXC1 * (skoXC1 * (-48 + skoXC1 * (skoXC1 * 16))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_55, refute_0_51, refute_0_50, refute_0_95])).

cnf(refute_0_105, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     tanh(skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (tanh(skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     0 <= tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(split, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (tanh(skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) <= 0 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (0 < tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) |
     tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_109, plain,
    (0 < tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) |
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) <= 0 |
     tanh(skoXC1) <= 0 / (skoXC1 * (288 + skoXC1 * (skoXC1 * 96)))),
    inference(split, [], [refute_0_108])).

cnf(refute_0_110, plain,
    (0 < tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) |
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) <= 0 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (0 < tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) |
     skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) <= 0 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_110, refute_0_55])).

cnf(refute_0_112, plain,
    (0 < skoXC1 * (288 + skoXC1 * (skoXC1 * 96)) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_113, plain,
    (0 < tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (tanh(skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_107, refute_0_113])).

cnf(refute_0_115, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_23, refute_0_114])).

cnf(refute_0_116, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) /
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_50, refute_0_115])).

cnf(refute_0_117, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96)))) |
     0 <= tanh(skoXC1) * (skoXC1 * (-288 + skoXC1 * (skoXC1 * -96))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_22, refute_0_116])).

cnf(refute_0_118, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) /
     (3 + skoXC1 * skoXC1) |
     tanh(skoXC1) * (skoXC1 * (288 + skoXC1 * (skoXC1 * 96))) <= 0 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (288 +
     skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) /
     (3 + skoXC1 * skoXC1) | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_118, refute_0_113])).

cnf(refute_0_120, plain,
    ((288 +
      skoXC1 * (-288 + skoXC1 * (-96 + skoXC1 * (-96 + skoXC1 * 32)))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_21, refute_0_119])).

cnf(refute_0_121, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (-3 < skoXC1 * skoXC1 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_123, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_121, refute_0_122])).

cnf(refute_0_124, plain,
    (tanh(skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     0 <= skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(split, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (tanh(skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     skoXC1 * (skoXC1 * (864 + skoXC1 * (skoXC1 * 288))) <= 0 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (0 < skoXC1 * (skoXC1 * (864 + skoXC1 * (skoXC1 * 288))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_127, plain,
    (tanh(skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_20, refute_0_127])).

cnf(refute_0_129, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) /
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_50, refute_0_128])).

cnf(refute_0_130, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288)))) |
     0 <= skoXC1 * (skoXC1 * (-864 + skoXC1 * (skoXC1 * -288))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_19, refute_0_129])).

cnf(refute_0_131, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2592 + skoXC1 * (skoXC1 * -864)))) /
     (3 + skoXC1 * skoXC1) |
     skoXC1 * (skoXC1 * (864 + skoXC1 * (skoXC1 * 288))) <= 0 |
     tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (864 +
     skoXC1 *
     (-864 +
      skoXC1 *
      (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32))))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2592 + skoXC1 * (skoXC1 * -864)))) /
     (3 + skoXC1 * skoXC1) | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_131, refute_0_126])).

cnf(refute_0_133, plain,
    ((864 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 * (-576 + skoXC1 * (skoXC1 * (-96 + skoXC1 * 32)))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2592 + skoXC1 * (skoXC1 * -864)))) |
     tanh(skoXC1) * tanh(skoXC1) <= 1 | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_18, refute_0_132])).

cnf(refute_0_134, plain,
    (2592 <
     skoXC1 *
     (2592 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (-96 + skoXC1 * (96 + skoXC1 * -32))))))) |
     skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (skoXC1 *
     (2592 +
      skoXC1 *
      (-864 +
       skoXC1 *
       (skoXC1 *
        (skoXC1 * (skoXC1 * (-96 + skoXC1 * (96 + skoXC1 * -32))))))) <=
     2592 | skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(decision, [],
              [refute_0_95, refute_0_50, refute_0_51, refute_0_55])).

cnf(refute_0_136, plain,
    (skoXC1 * skoXC1 <= -3 | tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_135, refute_0_134])).

cnf(refute_0_137, plain, (tanh(skoXC1) * tanh(skoXC1) <= 1),
    inference(resolve, [], [refute_0_136, refute_0_122])).

cnf(refute_0_138, plain, (-1 < tanh(skoXC1) * (tanh(skoXC1) * -1)),
    inference(resolve, [], [refute_0_137, refute_0_17])).

cnf(refute_0_139, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) < sinh(skoXC1 * 2)),
    inference(resolve, [], [refute_0_15, refute_0_138])).

cnf(refute_0_140, plain,
    (~ lgen(0, sinh(X_000120), X_000121) | sinh(X_000120) <= X_000121),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_141, plain,
    (4 < X_000120 | X_000120 < 0 |
     ~ lgen(0,
            -24 * X_000120 * (X_000120 ^ 4 + 70 * X_000120 ^ 2 + 600) /
            ((X_000120 ^ 3 - 12 * X_000120 ^ 2 + 60 * X_000120 - 120) *
             (X_000120 ^ 3 + 12 * X_000120 ^ 2 + 60 * X_000120 + 120)),
            X_000121) | lgen(0, sinh(X_000120), X_000121)),
    inference(subst, [], [sinh_upper_bound_pos2])).

cnf(refute_0_142, plain,
    (4 < X_000120 | X_000120 < 0 |
     ~ lgen(0,
            -24 * X_000120 * (X_000120 ^ 4 + 70 * X_000120 ^ 2 + 600) /
            ((X_000120 ^ 3 - 12 * X_000120 ^ 2 + 60 * X_000120 - 120) *
             (X_000120 ^ 3 + 12 * X_000120 ^ 2 + 60 * X_000120 + 120)),
            X_000121) | sinh(X_000120) <= X_000121),
    inference(resolve, [], [refute_0_141, refute_0_140])).

cnf(refute_0_143, plain,
    (4 < X_000120 | X_000120 < 0 |
     X_000121 <
     X_000120 *
     (-14400 +
      X_000120 * (X_000120 * (-1680 + X_000120 * (X_000120 * -24)))) /
     (-14400 +
      X_000120 *
      (X_000120 *
       (720 + X_000120 * (X_000120 * (-24 + X_000120 * X_000120))))) |
     sinh(X_000120) <= X_000121),
    inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (skoXC1 * 2 < 0 | 4 < skoXC1 * 2 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * 2 *
     (-14400 +
      skoXC1 * 2 *
      (skoXC1 * 2 * (-1680 + skoXC1 * 2 * (skoXC1 * 2 * -24)))) /
     (-14400 +
      skoXC1 * 2 *
      (skoXC1 * 2 *
       (720 +
        skoXC1 * 2 * (skoXC1 * 2 * (-24 + skoXC1 * 2 * (skoXC1 * 2)))))) |
     sinh(skoXC1 * 2) <= 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1))),
    inference(subst, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 * 2 < 0 | 4 < skoXC1 * 2 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * 2 *
     (-14400 +
      skoXC1 * 2 *
      (skoXC1 * 2 * (-1680 + skoXC1 * 2 * (skoXC1 * 2 * -24)))) /
     (-14400 +
      skoXC1 * 2 *
      (skoXC1 * 2 *
       (720 +
        skoXC1 * 2 * (skoXC1 * 2 * (-24 + skoXC1 * 2 * (skoXC1 * 2))))))),
    inference(resolve, [], [refute_0_144, refute_0_139])).

cnf(refute_0_146, plain,
    (skoXC1 < 0 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) |
     2 < skoXC1), inference(arithmetic, [], [refute_0_145])).

cnf(refute_0_147, plain,
    ((840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-88200 +
        skoXC1 * (skoXC1 * (-16800 + skoXC1 * (skoXC1 * -800)))))) |
     105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0 |
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-88200 +
        skoXC1 * (skoXC1 * (-16800 + skoXC1 * (skoXC1 * -800)))))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <=
     840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_148, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     0 <= skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80))) |
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) <=
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_149, plain,
    (~ lgen(0, X_000055, tanh(X_000054)) | X_000055 <= tanh(X_000054)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_150, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            5 * (2 * X_000054 ^ 2 + 21) * X_000054 /
            (X_000054 ^ 4 + 45 * X_000054 ^ 2 + 105)) |
     lgen(0, X_000055, tanh(X_000054))),
    inference(subst, [], [tanh_lower_bound_pos2])).

cnf(refute_0_151, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            5 * (2 * X_000054 ^ 2 + 21) * X_000054 /
            (X_000054 ^ 4 + 45 * X_000054 ^ 2 + 105)) |
     X_000055 <= tanh(X_000054)),
    inference(resolve, [], [refute_0_150, refute_0_149])).

cnf(refute_0_152, plain,
    (X_000054 * (105 + X_000054 * (X_000054 * 10)) /
     (105 + X_000054 * (X_000054 * (45 + X_000054 * X_000054))) <
     X_000055 | X_000054 < 0 | X_000055 <= tanh(X_000054)),
    inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_152])).

cnf(refute_0_154, plain,
    ((8 + skoXC1 * (-8 + skoXC1 * -8)) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0 |
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <=
     8 + skoXC1 * (-8 + skoXC1 * -8)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_155, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (tanh(skoXC1) * (skoXC1 * -8)) | 0 <= tanh(skoXC1) * (skoXC1 * -8) |
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) <=
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_156, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_152])).

cnf(refute_0_157, plain,
    (skoXC1 * -8 + tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * 8)) <
     2 * (-4 + skoXC1 * (skoXC1 * 4)) | 0 <= -4 + skoXC1 * (skoXC1 * 4) |
     (skoXC1 * -8 + tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * 8))) /
     (-4 + skoXC1 * (skoXC1 * 4)) <= 2),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_158, plain,
    (2 <
     skoXC1 * -8 / (-4 + skoXC1 * (skoXC1 * 4)) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     1 + tanh(skoXC1) * (tanh(skoXC1) * -1) <= 0 |
     skoXC1 * -8 / (-4 + skoXC1 * (skoXC1 * 4)) <=
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_159, plain,
    (~ lgen(0, sinh(X_000082), X_000083) | sinh(X_000082) <= X_000083),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_160, plain,
    (3/2 < X_000082 | X_000082 < 0 |
     ~ lgen(0, -4 * X_000082 / ((X_000082 - 2) * (X_000082 + 2)),
            X_000083) | lgen(0, sinh(X_000082), X_000083)),
    inference(subst, [], [sinh_upper_bound_pos1])).

cnf(refute_0_161, plain,
    (3/2 < X_000082 | X_000082 < 0 |
     ~ lgen(0, -4 * X_000082 / ((X_000082 - 2) * (X_000082 + 2)),
            X_000083) | sinh(X_000082) <= X_000083),
    inference(resolve, [], [refute_0_160, refute_0_159])).

cnf(refute_0_162, plain,
    (3/2 < X_000082 | X_000082 < 0 |
     X_000083 < X_000082 * -4 / (-4 + X_000082 * X_000082) |
     sinh(X_000082) <= X_000083),
    inference(arithmetic, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (skoXC1 * 2 < 0 | 3/2 < skoXC1 * 2 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * 2 * -4 / (-4 + skoXC1 * 2 * (skoXC1 * 2)) |
     sinh(skoXC1 * 2) <= 2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1))),
    inference(subst, [], [refute_0_162])).

cnf(refute_0_164, plain,
    (skoXC1 * 2 < 0 | 3/2 < skoXC1 * 2 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * 2 * -4 / (-4 + skoXC1 * 2 * (skoXC1 * 2))),
    inference(resolve, [], [refute_0_163, refute_0_139])).

cnf(refute_0_165, plain,
    (skoXC1 < 0 |
     2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * -8 / (-4 + skoXC1 * (skoXC1 * 4)) | 3/4 < skoXC1),
    inference(arithmetic, [], [refute_0_164])).

cnf(refute_0_166, plain, (3/4 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_50, refute_0_51])).

cnf(refute_0_167, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 * -8 / (-4 + skoXC1 * (skoXC1 * 4)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_166, refute_0_165])).

cnf(refute_0_168, plain,
    (2 <
     skoXC1 * -8 / (-4 + skoXC1 * (skoXC1 * 4)) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) | 3/4 < skoXC1 |
     1 + tanh(skoXC1) * (tanh(skoXC1) * -1) <= 0),
    inference(resolve, [], [refute_0_158, refute_0_167])).

cnf(refute_0_169, plain,
    (2 <
     (skoXC1 * -8 + tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * 8))) /
     (-4 + skoXC1 * (skoXC1 * 4)) | 3/4 < skoXC1 |
     tanh(skoXC1) * (tanh(skoXC1) * -1) <= -1),
    inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (2 <
     (skoXC1 * -8 + tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * 8))) /
     (-4 + skoXC1 * (skoXC1 * 4)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_169, refute_0_138])).

cnf(refute_0_171, plain,
    (skoXC1 * -8 + tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * 8)) <
     2 * (-4 + skoXC1 * (skoXC1 * 4)) | 3/4 < skoXC1 |
     0 <= -4 + skoXC1 * (skoXC1 * 4)),
    inference(resolve, [], [refute_0_157, refute_0_170])).

cnf(refute_0_172, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * -8)) | 3/4 < skoXC1 |
     skoXC1 * (skoXC1 * -4) <= -4),
    inference(arithmetic, [], [refute_0_171])).

cnf(refute_0_173, plain, (-4 < skoXC1 * (skoXC1 * -4) | 3/4 < skoXC1),
    inference(decision, [], [refute_0_51, refute_0_50])).

cnf(refute_0_174, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) * (tanh(skoXC1) * (skoXC1 * -8)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (tanh(skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1 | 0 <= tanh(skoXC1) * (skoXC1 * -8)),
    inference(split, [], [refute_0_174])).

cnf(refute_0_176, plain,
    (tanh(skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1 | tanh(skoXC1) * (skoXC1 * 8) <= 0),
    inference(arithmetic, [], [refute_0_175])).

cnf(refute_0_177, plain,
    (0 < tanh(skoXC1) * (skoXC1 * 8) | tanh(skoXC1) * (skoXC1 * 8) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_178, plain,
    (0 < tanh(skoXC1) * (skoXC1 * 8) | skoXC1 * 8 <= 0 |
     tanh(skoXC1) <= 0 / (skoXC1 * 8)),
    inference(split, [], [refute_0_177])).

cnf(refute_0_179, plain, (0 < tanh(skoXC1) * (skoXC1 * 8) | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_178])).

cnf(refute_0_180, plain,
    (tanh(skoXC1) * (skoXC1 * -8) < 0 |
     8 + skoXC1 * (-8 + skoXC1 * -8) < 0 |
     0 < tanh(skoXC1) * (skoXC1 * -8) | 3/4 < skoXC1),
    inference(split, [], [refute_0_174])).

cnf(refute_0_181, plain,
    (8 < skoXC1 * (8 + skoXC1 * 8) | 0 < tanh(skoXC1) * (skoXC1 * -8) |
     0 < tanh(skoXC1) * (skoXC1 * 8) | 3/4 < skoXC1),
    inference(arithmetic, [], [refute_0_180])).

cnf(refute_0_182, plain, (8 < skoXC1 * (8 + skoXC1 * 8) | skoXC1 <= 3/4),
    inference(decision, [], [refute_0_51, refute_0_50])).

cnf(refute_0_183, plain,
    (8 < skoXC1 * (8 + skoXC1 * 8) | 0 < tanh(skoXC1) * (skoXC1 * -8) |
     0 < tanh(skoXC1) * (skoXC1 * 8)),
    inference(resolve, [], [refute_0_182, refute_0_181])).

cnf(refute_0_184, plain,
    (0 < tanh(skoXC1) * (skoXC1 * -8) | tanh(skoXC1) * (skoXC1 * -8) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_185, plain,
    (tanh(skoXC1) < 0 / (skoXC1 * -8) | 0 <= skoXC1 * -8 |
     tanh(skoXC1) * (skoXC1 * -8) <= 0),
    inference(split, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (skoXC1 < 0 | tanh(skoXC1) * (skoXC1 * -8) <= 0 | skoXC1 <= 0),
    inference(arithmetic, [], [refute_0_185])).

cnf(refute_0_187, plain, (tanh(skoXC1) * (skoXC1 * -8) <= 0 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_50, refute_0_186])).

cnf(refute_0_188, plain,
    (0 < 0 | skoXC1 * -8 < 0 | 0 < skoXC1 * -8 |
     tanh(skoXC1) * (skoXC1 * -8) <= 0),
    inference(split, [], [refute_0_184])).

cnf(refute_0_189, plain,
    (skoXC1 < 0 | 0 < skoXC1 | tanh(skoXC1) * (skoXC1 * -8) <= 0),
    inference(arithmetic, [], [refute_0_188])).

cnf(refute_0_190, plain, (0 < skoXC1 | tanh(skoXC1) * (skoXC1 * -8) <= 0),
    inference(resolve, [], [refute_0_50, refute_0_189])).

cnf(refute_0_191, plain, (tanh(skoXC1) * (skoXC1 * -8) <= 0),
    inference(resolve, [], [refute_0_187, refute_0_190])).

cnf(refute_0_192, plain,
    (8 < skoXC1 * (8 + skoXC1 * 8) | 0 < tanh(skoXC1) * (skoXC1 * 8)),
    inference(resolve, [], [refute_0_191, refute_0_183])).

cnf(refute_0_193, plain,
    (0 < 0 | skoXC1 * 8 < 0 | 0 < skoXC1 * 8 |
     8 < skoXC1 * (8 + skoXC1 * 8)), inference(split, [], [refute_0_192])).

cnf(refute_0_194, plain,
    (skoXC1 < 0 | 8 < skoXC1 * (8 + skoXC1 * 8) | 0 < skoXC1),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain, (8 < skoXC1 * (8 + skoXC1 * 8) | 0 < skoXC1),
    inference(resolve, [], [refute_0_50, refute_0_194])).

cnf(refute_0_196, plain, (0 < skoXC1 | skoXC1 * (8 + skoXC1 * 8) <= 8),
    inference(decision, [], [refute_0_50, refute_0_51])).

cnf(refute_0_197, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_196, refute_0_195])).

cnf(refute_0_198, plain, (0 < tanh(skoXC1) * (skoXC1 * 8)),
    inference(resolve, [], [refute_0_179, refute_0_197])).

cnf(refute_0_199, plain,
    (tanh(skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_176, refute_0_198])).

cnf(refute_0_200, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_156, refute_0_199])).

cnf(refute_0_201, plain,
    (skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_50, refute_0_200])).

cnf(refute_0_202, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (tanh(skoXC1) * (skoXC1 * -8)) | 3/4 < skoXC1 |
     0 <= tanh(skoXC1) * (skoXC1 * -8)),
    inference(resolve, [], [refute_0_155, refute_0_201])).

cnf(refute_0_203, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1 |
     tanh(skoXC1) * (skoXC1 * 8) <= 0),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_203, refute_0_198])).

cnf(refute_0_205, plain,
    ((8 + skoXC1 * (-8 + skoXC1 * -8)) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1 | 105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0),
    inference(resolve, [], [refute_0_154, refute_0_204])).

cnf(refute_0_206, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1 | skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= -105),
    inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    ((24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) *
     (3 + skoXC1 * skoXC1) < skoXC1 * (skoXC1 * (skoXC1 * -72)) |
     3 + skoXC1 * skoXC1 <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * -72)) / (3 + skoXC1 * skoXC1) <=
     24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_208, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * -24)) |
     0 <= skoXC1 * (skoXC1 * -24) |
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) <= skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_209, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) |
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) <= tanh(skoXC1)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_210, plain,
    ((8 + skoXC1 * (-8 + skoXC1 * -8)) * (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) | 3 + skoXC1 * skoXC1 <= 0 |
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) / (3 + skoXC1 * skoXC1) <=
     8 + skoXC1 * (-8 + skoXC1 * -8)),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_211, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) * (tanh(skoXC1) * (skoXC1 * -8)) |
     0 <= tanh(skoXC1) * (skoXC1 * -8) |
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) <=
     skoXC1 * 3 / (3 + skoXC1 * skoXC1)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_212, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) <=
     tanh(skoXC1)), inference(subst, [], [refute_0_5])).

cnf(refute_0_213, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_212, refute_0_199])).

cnf(refute_0_214, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (8 + skoXC1 * (-8 + skoXC1 * -8)) / (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_50, refute_0_213])).

cnf(refute_0_215, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) * (tanh(skoXC1) * (skoXC1 * -8)) |
     3/4 < skoXC1 | 0 <= tanh(skoXC1) * (skoXC1 * -8)),
    inference(resolve, [], [refute_0_211, refute_0_214])).

cnf(refute_0_216, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) / (3 + skoXC1 * skoXC1) |
     3/4 < skoXC1 | tanh(skoXC1) * (skoXC1 * 8) <= 0),
    inference(arithmetic, [], [refute_0_215])).

cnf(refute_0_217, plain,
    (8 + skoXC1 * (-8 + skoXC1 * -8) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) / (3 + skoXC1 * skoXC1) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_216, refute_0_198])).

cnf(refute_0_218, plain,
    ((8 + skoXC1 * (-8 + skoXC1 * -8)) * (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1 |
     3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_210, refute_0_217])).

cnf(refute_0_219, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1 |
     skoXC1 * skoXC1 <= -3), inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain, (-3 < skoXC1 * skoXC1 | 3/4 < skoXC1),
    inference(decision, [], [refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_221, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     tanh(skoXC1) * (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_219, refute_0_220])).

cnf(refute_0_222, plain,
    (tanh(skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1 |
     0 <= skoXC1 * (skoXC1 * -24)), inference(split, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (tanh(skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1 |
     skoXC1 * (skoXC1 * 24) <= 0),
    inference(arithmetic, [], [refute_0_222])).

cnf(refute_0_224, plain, (0 < skoXC1 * (skoXC1 * 24) | 3/4 < skoXC1),
    inference(decision, [], [refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_225, plain,
    (tanh(skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_223, refute_0_224])).

cnf(refute_0_226, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_209, refute_0_225])).

cnf(refute_0_227, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_50, refute_0_226])).

cnf(refute_0_228, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) * (skoXC1 * (skoXC1 * -24)) |
     3/4 < skoXC1 | 0 <= skoXC1 * (skoXC1 * -24)),
    inference(resolve, [], [refute_0_208, refute_0_227])).

cnf(refute_0_229, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (skoXC1 * (skoXC1 * -72)) / (3 + skoXC1 * skoXC1) |
     3/4 < skoXC1 | skoXC1 * (skoXC1 * 24) <= 0),
    inference(arithmetic, [], [refute_0_228])).

cnf(refute_0_230, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (skoXC1 * (skoXC1 * -72)) / (3 + skoXC1 * skoXC1) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_229, refute_0_224])).

cnf(refute_0_231, plain,
    ((24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) *
     (3 + skoXC1 * skoXC1) < skoXC1 * (skoXC1 * (skoXC1 * -72)) |
     3/4 < skoXC1 | 3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_207, refute_0_230])).

cnf(refute_0_232, plain,
    (72 <
     skoXC1 *
     (72 +
      skoXC1 *
      (24 + skoXC1 * (-24 + skoXC1 * (40 + skoXC1 * (8 + skoXC1 * 8))))) |
     3/4 < skoXC1 | skoXC1 * skoXC1 <= -3),
    inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    (72 <
     skoXC1 *
     (72 +
      skoXC1 *
      (24 + skoXC1 * (-24 + skoXC1 * (40 + skoXC1 * (8 + skoXC1 * 8))))) |
     -3 < skoXC1 * skoXC1 | 3/4 < skoXC1),
    inference(decision, [], [refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_234, plain,
    (72 <
     skoXC1 *
     (72 +
      skoXC1 *
      (24 + skoXC1 * (-24 + skoXC1 * (40 + skoXC1 * (8 + skoXC1 * 8))))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_232, refute_0_233])).

cnf(refute_0_235, plain,
    (72 <
     skoXC1 *
     (72 +
      skoXC1 *
      (24 + skoXC1 * (-24 + skoXC1 * (40 + skoXC1 * (8 + skoXC1 * 8))))) |
     skoXC1 <= 3/4),
    inference(decision, [], [refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_236, plain,
    (72 <
     skoXC1 *
     (72 +
      skoXC1 *
      (24 + skoXC1 * (-24 + skoXC1 * (40 + skoXC1 * (8 + skoXC1 * 8)))))),
    inference(resolve, [], [refute_0_235, refute_0_234])).

cnf(refute_0_237, plain,
    ((24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2520 + skoXC1 * (skoXC1 * -240)))) |
     105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0 |
     skoXC1 * (skoXC1 * (skoXC1 * (-2520 + skoXC1 * (skoXC1 * -240)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <=
     24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_238, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (skoXC1 * (skoXC1 * -24)) | 0 <= skoXC1 * (skoXC1 * -24) |
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) <=
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_239, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) |
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) <= tanh(skoXC1)),
    inference(subst, [], [refute_0_152])).

cnf(refute_0_240, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_239, refute_0_225])).

cnf(refute_0_241, plain,
    (skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) /
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_50, refute_0_240])).

cnf(refute_0_242, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (skoXC1 * (skoXC1 * -24)) | 3/4 < skoXC1 |
     0 <= skoXC1 * (skoXC1 * -24)),
    inference(resolve, [], [refute_0_238, refute_0_241])).

cnf(refute_0_243, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2520 + skoXC1 * (skoXC1 * -240)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1 |
     skoXC1 * (skoXC1 * 24) <= 0),
    inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2520 + skoXC1 * (skoXC1 * -240)))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_243, refute_0_224])).

cnf(refute_0_245, plain,
    ((24 + skoXC1 * (-24 + skoXC1 * (-16 + skoXC1 * (-8 + skoXC1 * -8)))) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     skoXC1 * (skoXC1 * (skoXC1 * (-2520 + skoXC1 * (skoXC1 * -240)))) |
     3/4 < skoXC1 | 105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0),
    inference(resolve, [], [refute_0_237, refute_0_244])).

cnf(refute_0_246, plain,
    (2520 <
     skoXC1 *
     (2520 +
      skoXC1 *
      (600 +
       skoXC1 *
       (-600 +
        skoXC1 *
        (1536 +
         skoXC1 * (144 + skoXC1 * (376 + skoXC1 * (8 + skoXC1 * 8))))))) |
     3/4 < skoXC1 | skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= -105),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (-105 < skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) |
     2520 <
     skoXC1 *
     (2520 +
      skoXC1 *
      (600 +
       skoXC1 *
       (-600 +
        skoXC1 *
        (1536 +
         skoXC1 * (144 + skoXC1 * (376 + skoXC1 * (8 + skoXC1 * 8))))))) |
     3/4 < skoXC1),
    inference(decision, [],
              [refute_0_236, refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_248, plain,
    (2520 <
     skoXC1 *
     (2520 +
      skoXC1 *
      (600 +
       skoXC1 *
       (-600 +
        skoXC1 *
        (1536 +
         skoXC1 * (144 + skoXC1 * (376 + skoXC1 * (8 + skoXC1 * 8))))))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_246, refute_0_247])).

cnf(refute_0_249, plain,
    (2520 <
     skoXC1 *
     (2520 +
      skoXC1 *
      (600 +
       skoXC1 *
       (-600 +
        skoXC1 *
        (1536 +
         skoXC1 * (144 + skoXC1 * (376 + skoXC1 * (8 + skoXC1 * 8))))))) |
     skoXC1 <= 3/4),
    inference(decision, [],
              [refute_0_236, refute_0_197, refute_0_50, refute_0_51])).

cnf(refute_0_250, plain,
    (2520 <
     skoXC1 *
     (2520 +
      skoXC1 *
      (600 +
       skoXC1 *
       (-600 +
        skoXC1 *
        (1536 +
         skoXC1 * (144 + skoXC1 * (376 + skoXC1 * (8 + skoXC1 * 8)))))))),
    inference(resolve, [], [refute_0_249, refute_0_248])).

cnf(refute_0_251, plain,
    (-105 < skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) | 3/4 < skoXC1),
    inference(decision, [],
              [refute_0_51, refute_0_50, refute_0_197, refute_0_236,
               refute_0_250])).

cnf(refute_0_252, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     tanh(skoXC1) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_206, refute_0_251])).

cnf(refute_0_253, plain,
    (tanh(skoXC1) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1 |
     0 <= skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))),
    inference(split, [], [refute_0_252])).

cnf(refute_0_254, plain,
    (tanh(skoXC1) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1 |
     skoXC1 * (skoXC1 * (840 + skoXC1 * (skoXC1 * 80))) <= 0),
    inference(arithmetic, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (0 < skoXC1 * (skoXC1 * (840 + skoXC1 * (skoXC1 * 80))) |
     3/4 < skoXC1),
    inference(decision, [],
              [refute_0_51, refute_0_50, refute_0_197, refute_0_236,
               refute_0_250])).

cnf(refute_0_256, plain,
    (tanh(skoXC1) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_254, refute_0_255])).

cnf(refute_0_257, plain,
    (skoXC1 < 0 |
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_153, refute_0_256])).

cnf(refute_0_258, plain,
    (skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     (840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) /
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_50, refute_0_257])).

cnf(refute_0_259, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     skoXC1 * (105 + skoXC1 * (skoXC1 * 10)) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) *
     (skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))) |
     3/4 < skoXC1 |
     0 <= skoXC1 * (skoXC1 * (-840 + skoXC1 * (skoXC1 * -80)))),
    inference(resolve, [], [refute_0_148, refute_0_258])).

cnf(refute_0_260, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-88200 +
        skoXC1 * (skoXC1 * (-16800 + skoXC1 * (skoXC1 * -800)))))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1 |
     skoXC1 * (skoXC1 * (840 + skoXC1 * (skoXC1 * 80))) <= 0),
    inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (840 +
     skoXC1 *
     (-840 +
      skoXC1 *
      (-480 +
       skoXC1 * (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-88200 +
        skoXC1 * (skoXC1 * (-16800 + skoXC1 * (skoXC1 * -800)))))) /
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) | 3/4 < skoXC1),
    inference(resolve, [], [refute_0_260, refute_0_255])).

cnf(refute_0_262, plain,
    ((840 +
      skoXC1 *
      (-840 +
       skoXC1 *
       (-480 +
        skoXC1 *
        (-360 + skoXC1 * (-352 + skoXC1 * (-8 + skoXC1 * -8)))))) *
     (105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-88200 +
        skoXC1 * (skoXC1 * (-16800 + skoXC1 * (skoXC1 * -800)))))) |
     3/4 < skoXC1 | 105 + skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= 0),
    inference(resolve, [], [refute_0_147, refute_0_261])).

cnf(refute_0_263, plain,
    (88200 <
     skoXC1 *
     (88200 +
      skoXC1 *
      (12600 +
       skoXC1 *
       (-12600 +
        skoXC1 *
        (57720 +
         skoXC1 *
         (1080 +
          skoXC1 *
          (17160 +
           skoXC1 *
           (-80 + skoXC1 * (712 + skoXC1 * (8 + skoXC1 * 8))))))))) |
     3/4 < skoXC1 | skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) <= -105),
    inference(arithmetic, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (-105 < skoXC1 * (skoXC1 * (45 + skoXC1 * skoXC1)) |
     88200 <
     skoXC1 *
     (88200 +
      skoXC1 *
      (12600 +
       skoXC1 *
       (-12600 +
        skoXC1 *
        (57720 +
         skoXC1 *
         (1080 +
          skoXC1 *
          (17160 +
           skoXC1 *
           (-80 + skoXC1 * (712 + skoXC1 * (8 + skoXC1 * 8))))))))) |
     3/4 < skoXC1),
    inference(decision, [],
              [refute_0_51, refute_0_50, refute_0_197, refute_0_236,
               refute_0_250])).

cnf(refute_0_265, plain,
    (88200 <
     skoXC1 *
     (88200 +
      skoXC1 *
      (12600 +
       skoXC1 *
       (-12600 +
        skoXC1 *
        (57720 +
         skoXC1 *
         (1080 +
          skoXC1 *
          (17160 +
           skoXC1 *
           (-80 + skoXC1 * (712 + skoXC1 * (8 + skoXC1 * 8))))))))) |
     3/4 < skoXC1), inference(resolve, [], [refute_0_263, refute_0_264])).

cnf(refute_0_266, plain,
    (88200 <
     skoXC1 *
     (88200 +
      skoXC1 *
      (12600 +
       skoXC1 *
       (-12600 +
        skoXC1 *
        (57720 +
         skoXC1 *
         (1080 +
          skoXC1 *
          (17160 +
           skoXC1 *
           (-80 + skoXC1 * (712 + skoXC1 * (8 + skoXC1 * 8))))))))) |
     skoXC1 <= 3/4),
    inference(decision, [],
              [refute_0_51, refute_0_50, refute_0_197, refute_0_236,
               refute_0_250])).

cnf(refute_0_267, plain,
    (88200 <
     skoXC1 *
     (88200 +
      skoXC1 *
      (12600 +
       skoXC1 *
       (-12600 +
        skoXC1 *
        (57720 +
         skoXC1 *
         (1080 +
          skoXC1 *
          (17160 +
           skoXC1 *
           (-80 + skoXC1 * (712 + skoXC1 * (8 + skoXC1 * 8)))))))))),
    inference(resolve, [], [refute_0_266, refute_0_265])).

cnf(refute_0_268, plain, (2 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_267, refute_0_51, refute_0_50, refute_0_197,
               refute_0_236, refute_0_250])).

cnf(refute_0_269, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) |
     2 < skoXC1), inference(resolve, [], [refute_0_268, refute_0_146])).

cnf(refute_0_270, plain, (skoXC1 <= 2),
    inference(decision, [],
              [refute_0_267, refute_0_51, refute_0_50, refute_0_197,
               refute_0_236, refute_0_250])).

cnf(refute_0_271, plain,
    (2 / (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) <
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64))))))),
    inference(resolve, [], [refute_0_270, refute_0_269])).

cnf(refute_0_272, plain,
    (2 <
     skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) *
     (1 + tanh(skoXC1) * (tanh(skoXC1) * -1)) |
     1 + tanh(skoXC1) * (tanh(skoXC1) * -1) <= 0),
    inference(resolve, [], [refute_0_11, refute_0_271])).

cnf(refute_0_273, plain,
    (2 <
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) +
      tanh(skoXC1) *
      (tanh(skoXC1) *
       (skoXC1 *
        (28800 +
         skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) |
     tanh(skoXC1) * (tanh(skoXC1) * -1) <= -1),
    inference(arithmetic, [], [refute_0_272])).

cnf(refute_0_274, plain,
    (2 <
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) +
      tanh(skoXC1) *
      (tanh(skoXC1) *
       (skoXC1 *
        (28800 +
         skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))))) /
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64))))))),
    inference(resolve, [], [refute_0_273, refute_0_138])).

cnf(refute_0_275, plain,
    (skoXC1 *
     (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))) +
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))))) <
     2 *
     (-14400 +
      skoXC1 *
      (skoXC1 *
       (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))) |
     0 <=
     -14400 +
     skoXC1 *
     (skoXC1 *
      (2880 + skoXC1 * (skoXC1 * (-384 + skoXC1 * (skoXC1 * 64)))))),
    inference(resolve, [], [refute_0_10, refute_0_274])).

cnf(refute_0_276, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     skoXC1 *
     (skoXC1 *
      (-2880 + skoXC1 * (skoXC1 * (384 + skoXC1 * (skoXC1 * -64))))) <=
     -14400), inference(arithmetic, [], [refute_0_275])).

cnf(refute_0_277, plain,
    (-14400 <
     skoXC1 *
     (skoXC1 *
      (-2880 + skoXC1 * (skoXC1 * (384 + skoXC1 * (skoXC1 * -64)))))),
    inference(decision, [],
              [refute_0_250, refute_0_236, refute_0_197, refute_0_50,
               refute_0_51, refute_0_267])).

cnf(refute_0_278, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     tanh(skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768))))))),
    inference(resolve, [], [refute_0_276, refute_0_277])).

cnf(refute_0_279, plain,
    (tanh(skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))),
    inference(split, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (tanh(skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) <=
     0), inference(arithmetic, [], [refute_0_279])).

cnf(refute_0_281, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) |
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_282, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) |
     skoXC1 *
     (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))) <= 0 |
     tanh(skoXC1) <=
     0 /
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))))),
    inference(split, [], [refute_0_281])).

cnf(refute_0_283, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) |
     skoXC1 *
     (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))) <= 0 |
     skoXC1 <= 0), inference(arithmetic, [], [refute_0_282])).

cnf(refute_0_284, plain,
    (0 <
     skoXC1 *
     (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))) |
     skoXC1 <= 0),
    inference(decision, [],
              [refute_0_250, refute_0_236, refute_0_197, refute_0_50,
               refute_0_51, refute_0_267])).

cnf(refute_0_285, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) |
     skoXC1 <= 0), inference(resolve, [], [refute_0_283, refute_0_284])).

cnf(refute_0_286, plain,
    (0 <
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768)))))),
    inference(resolve, [], [refute_0_285, refute_0_197])).

cnf(refute_0_287, plain,
    (tanh(skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768))))))),
    inference(resolve, [], [refute_0_280, refute_0_286])).

cnf(refute_0_288, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768))))))),
    inference(resolve, [], [refute_0_9, refute_0_287])).

cnf(refute_0_289, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) /
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768))))))),
    inference(resolve, [], [refute_0_50, refute_0_288])).

cnf(refute_0_290, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (tanh(skoXC1) *
      (skoXC1 *
       (-28800 +
        skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))) |
     0 <=
     tanh(skoXC1) *
     (skoXC1 *
      (-28800 + skoXC1 * (skoXC1 * (-13440 + skoXC1 * (skoXC1 * -768)))))),
    inference(resolve, [], [refute_0_8, refute_0_289])).

cnf(refute_0_291, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) /
     (3 + skoXC1 * skoXC1) |
     tanh(skoXC1) *
     (skoXC1 *
      (28800 + skoXC1 * (skoXC1 * (13440 + skoXC1 * (skoXC1 * 768))))) <=
     0), inference(arithmetic, [], [refute_0_290])).

cnf(refute_0_292, plain,
    (28800 +
     skoXC1 *
     (-28800 +
      skoXC1 *
      (-5760 +
       skoXC1 *
       (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) /
     (3 + skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_291, refute_0_286])).

cnf(refute_0_293, plain,
    ((28800 +
      skoXC1 *
      (-28800 +
       skoXC1 *
       (-5760 +
        skoXC1 *
        (-13440 + skoXC1 * (768 + skoXC1 * (-768 + skoXC1 * -128)))))) *
     (3 + skoXC1 * skoXC1) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_7, refute_0_292])).

cnf(refute_0_294, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     skoXC1 * skoXC1 <= -3), inference(arithmetic, [], [refute_0_293])).

cnf(refute_0_295, plain, (-3 < skoXC1 * skoXC1),
    inference(decision, [],
              [refute_0_250, refute_0_236, refute_0_197, refute_0_50,
               refute_0_51, refute_0_267])).

cnf(refute_0_296, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     tanh(skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304))))))),
    inference(resolve, [], [refute_0_294, refute_0_295])).

cnf(refute_0_297, plain,
    (tanh(skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-86400 +
       skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))),
    inference(split, [], [refute_0_296])).

cnf(refute_0_298, plain,
    (tanh(skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     skoXC1 *
     (skoXC1 *
      (86400 + skoXC1 * (skoXC1 * (40320 + skoXC1 * (skoXC1 * 2304))))) <=
     0), inference(arithmetic, [], [refute_0_297])).

cnf(refute_0_299, plain,
    (0 <
     skoXC1 *
     (skoXC1 *
      (86400 + skoXC1 * (skoXC1 * (40320 + skoXC1 * (skoXC1 * 2304)))))),
    inference(decision, [],
              [refute_0_250, refute_0_236, refute_0_197, refute_0_50,
               refute_0_51, refute_0_267])).

cnf(refute_0_300, plain,
    (tanh(skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304))))))),
    inference(resolve, [], [refute_0_298, refute_0_299])).

cnf(refute_0_301, plain,
    (skoXC1 < 0 |
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304))))))),
    inference(resolve, [], [refute_0_6, refute_0_300])).

cnf(refute_0_302, plain,
    (skoXC1 * 3 / (3 + skoXC1 * skoXC1) <
     (86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) /
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304))))))),
    inference(resolve, [], [refute_0_50, refute_0_301])).

cnf(refute_0_303, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     skoXC1 * 3 / (3 + skoXC1 * skoXC1) *
     (skoXC1 *
      (skoXC1 *
       (-86400 +
        skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))) |
     0 <=
     skoXC1 *
     (skoXC1 *
      (-86400 +
       skoXC1 * (skoXC1 * (-40320 + skoXC1 * (skoXC1 * -2304)))))),
    inference(resolve, [], [refute_0_1, refute_0_302])).

cnf(refute_0_304, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-259200 +
        skoXC1 * (skoXC1 * (-120960 + skoXC1 * (skoXC1 * -6912)))))) /
     (3 + skoXC1 * skoXC1) |
     skoXC1 *
     (skoXC1 *
      (86400 + skoXC1 * (skoXC1 * (40320 + skoXC1 * (skoXC1 * 2304))))) <=
     0), inference(arithmetic, [], [refute_0_303])).

cnf(refute_0_305, plain,
    (86400 +
     skoXC1 *
     (-86400 +
      skoXC1 *
      (11520 +
       skoXC1 *
       (-69120 +
        skoXC1 *
        (-3456 +
         skoXC1 *
         (-15744 + skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128))))))) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-259200 +
        skoXC1 * (skoXC1 * (-120960 + skoXC1 * (skoXC1 * -6912)))))) /
     (3 + skoXC1 * skoXC1)),
    inference(resolve, [], [refute_0_304, refute_0_299])).

cnf(refute_0_306, plain,
    ((86400 +
      skoXC1 *
      (-86400 +
       skoXC1 *
       (11520 +
        skoXC1 *
        (-69120 +
         skoXC1 *
         (-3456 +
          skoXC1 *
          (-15744 +
           skoXC1 * (384 + skoXC1 * (-768 + skoXC1 * -128)))))))) *
     (3 + skoXC1 * skoXC1) <
     skoXC1 *
     (skoXC1 *
      (skoXC1 *
       (-259200 +
        skoXC1 * (skoXC1 * (-120960 + skoXC1 * (skoXC1 * -6912)))))) |
     3 + skoXC1 * skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_305])).

cnf(refute_0_307, plain,
    (259200 <
     skoXC1 *
     (259200 +
      skoXC1 *
      (-120960 +
       skoXC1 *
       (34560 +
        skoXC1 *
        (-1152 +
         skoXC1 *
         (-4608 +
          skoXC1 *
          (2304 +
           skoXC1 *
           (11136 + skoXC1 * (skoXC1 * (768 + skoXC1 * 128))))))))) |
     skoXC1 * skoXC1 <= -3), inference(arithmetic, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (skoXC1 *
     (259200 +
      skoXC1 *
      (-120960 +
       skoXC1 *
       (34560 +
        skoXC1 *
        (-1152 +
         skoXC1 *
         (-4608 +
          skoXC1 *
          (2304 +
           skoXC1 *
           (11136 + skoXC1 * (skoXC1 * (768 + skoXC1 * 128))))))))) <=
     259200 | skoXC1 * skoXC1 <= -3),
    inference(decision, [],
              [refute_0_250, refute_0_236, refute_0_197, refute_0_50,
               refute_0_51, refute_0_267])).

cnf(refute_0_309, plain, (skoXC1 * skoXC1 <= -3),
    inference(resolve, [], [refute_0_308, refute_0_307])).

cnf(refute_0_310, plain, ($false),
    inference(resolve, [], [refute_0_309, refute_0_295])).
% SZS output end CNFRefutation for tanh-8-weak2.tptp
