% SZS output start CNFRefutation for trigpoly-351-1a.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(tan_lower_bound_pos0, axiom,
    (X < 0 | 157/100 < X | ~ lgen(R, Y, X) | lgen(R, Y, tan(X)))).

cnf(tan_lower_bound_pos2, axiom,
    (X < 0 | 314159/200000 < X |
     ~ lgen(R, Y,
            -(21 * (X ^ 4 - 60 * X ^ 2 + 495)) * X /
            (X ^ 6 - 210 * X ^ 4 + 4725 * X ^ 2 - 10395)) |
     lgen(R, Y, tan(X)))).

cnf(sin_lower_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040) |
     lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_5_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - 1/6 * X ^ 3 + 1/120 * X ^ 5 - 1/5040 * X ^ 7 +
            1/362880 * X ^ 9 - 1/39916800 * X ^ 11 +
            1/6227020800 * X ^ 13 - 1/1307674368000 * X ^ 15 +
            1/355687428096000 * X ^ 17 - 1/121645100408832000 * X ^ 19 +
            1/51090942171709440000 * X ^ 21 -
            1/25852016738884976640000 * X ^ 23) | lgen(R, Y, sin(X)))).

fof(trig_poly_351, conjecture,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        -1/2 * tan(X / 4) <= 1/250000000 + sin(X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        -1/2 * tan(X / 4) <= 1/250000000 + sin(X))),
    inference(strip, [], [trig_poly_351])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 62831/10000) =>
          -1/2 * tan(X / 4) <= 1/250000000 + sin(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 62831/10000 &
        1/250000000 + sin(X) < -1/2 * tan(X / 4))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 62831/10000 &
     1/250000000 + sin(skoXC1) < -1/2 * tan(skoXC1 / 4) & 0 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/250000000 + sin(skoXC1) < -1/2 * tan(skoXC1 / 4)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 62831/10000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     0 <=
     -20790 +
     skoXC1 *
     (skoXC1 *
      (4725/8 +
       skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))) |
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000))))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000111, sin(X_000110)) | X_000111 <= sin(X_000110)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (X_000110 < 0 |
     ~ lgen(0, X_000111,
            X_000110 - 1/6 * X_000110 ^ 3 + 1/120 * X_000110 ^ 5 -
            1/5040 * X_000110 ^ 7 + 1/362880 * X_000110 ^ 9 -
            1/39916800 * X_000110 ^ 11 + 1/6227020800 * X_000110 ^ 13 -
            1/1307674368000 * X_000110 ^ 15 +
            1/355687428096000 * X_000110 ^ 17 -
            1/121645100408832000 * X_000110 ^ 19 +
            1/51090942171709440000 * X_000110 ^ 21 -
            1/25852016738884976640000 * X_000110 ^ 23) |
     lgen(0, X_000111, sin(X_000110))),
    inference(subst, [], [sin_lower_bound_5_pos])).

cnf(refute_0_3, plain,
    (X_000110 < 0 |
     ~ lgen(0, X_000111,
            X_000110 - 1/6 * X_000110 ^ 3 + 1/120 * X_000110 ^ 5 -
            1/5040 * X_000110 ^ 7 + 1/362880 * X_000110 ^ 9 -
            1/39916800 * X_000110 ^ 11 + 1/6227020800 * X_000110 ^ 13 -
            1/1307674368000 * X_000110 ^ 15 +
            1/355687428096000 * X_000110 ^ 17 -
            1/121645100408832000 * X_000110 ^ 19 +
            1/51090942171709440000 * X_000110 ^ 21 -
            1/25852016738884976640000 * X_000110 ^ 23) |
     X_000111 <= sin(X_000110)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000110 *
     (1 +
      X_000110 *
      (X_000110 *
       (-1/6 +
        X_000110 *
        (X_000110 *
         (1/120 +
          X_000110 *
          (X_000110 *
           (-1/5040 +
            X_000110 *
            (X_000110 *
             (1/362880 +
              X_000110 *
              (X_000110 *
               (-1/39916800 +
                X_000110 *
                (X_000110 *
                 (1/6227020800 +
                  X_000110 *
                  (X_000110 *
                   (-1/1307674368000 +
                    X_000110 *
                    (X_000110 *
                     (1/355687428096000 +
                      X_000110 *
                      (X_000110 *
                       (-1/121645100408832000 +
                        X_000110 *
                        (X_000110 *
                         (1/51090942171709440000 +
                          X_000110 *
                          (X_000110 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) <
     X_000111 | X_000110 < 0 | X_000111 <= sin(X_000110)),
    inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) <=
     sin(skoXC1)), inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    ((-1/125000000 + sin(skoXC1) * -2) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))) <
     skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/16 +
       skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096))))) |
     -1/125000000 + sin(skoXC1) * -2 <=
     skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096))))))),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_7, plain,
    (1/250000000 + sin(skoXC1) < -1/2 * tan(skoXC1 / 4)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_8, plain,
    (tan(skoXC1 * 1/4) < -1/125000000 + sin(skoXC1) * -2),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, X_000146, tan(X_000145)) | X_000146 <= tan(X_000145)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (314159/200000 < X_000145 | X_000145 < 0 |
     ~ lgen(0, X_000146,
            -(21 * (X_000145 ^ 4 - 60 * X_000145 ^ 2 + 495)) * X_000145 /
            (X_000145 ^ 6 - 210 * X_000145 ^ 4 + 4725 * X_000145 ^ 2 -
             10395)) | lgen(0, X_000146, tan(X_000145))),
    inference(subst, [], [tan_lower_bound_pos2])).

cnf(refute_0_11, plain,
    (314159/200000 < X_000145 | X_000145 < 0 |
     ~ lgen(0, X_000146,
            -(21 * (X_000145 ^ 4 - 60 * X_000145 ^ 2 + 495)) * X_000145 /
            (X_000145 ^ 6 - 210 * X_000145 ^ 4 + 4725 * X_000145 ^ 2 -
             10395)) | X_000146 <= tan(X_000145)),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (314159/200000 < X_000145 |
     X_000145 *
     (-10395 +
      X_000145 * (X_000145 * (1260 + X_000145 * (X_000145 * -21)))) /
     (-10395 +
      X_000145 *
      (X_000145 *
       (4725 + X_000145 * (X_000145 * (-210 + X_000145 * X_000145))))) <
     X_000146 | X_000145 < 0 | X_000146 <= tan(X_000145)),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (skoXC1 * 1/4 < 0 | 314159/200000 < skoXC1 * 1/4 |
     skoXC1 * 1/4 *
     (-10395 +
      skoXC1 * 1/4 *
      (skoXC1 * 1/4 * (1260 + skoXC1 * 1/4 * (skoXC1 * 1/4 * -21)))) /
     (-10395 +
      skoXC1 * 1/4 *
      (skoXC1 * 1/4 *
       (4725 +
        skoXC1 * 1/4 *
        (skoXC1 * 1/4 * (-210 + skoXC1 * 1/4 * (skoXC1 * 1/4)))))) <
     -1/125000000 + sin(skoXC1) * -2 |
     -1/125000000 + sin(skoXC1) * -2 <= tan(skoXC1 * 1/4)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 * 1/4 < 0 | 314159/200000 < skoXC1 * 1/4 |
     skoXC1 * 1/4 *
     (-10395 +
      skoXC1 * 1/4 *
      (skoXC1 * 1/4 * (1260 + skoXC1 * 1/4 * (skoXC1 * 1/4 * -21)))) /
     (-10395 +
      skoXC1 * 1/4 *
      (skoXC1 * 1/4 *
       (4725 +
        skoXC1 * 1/4 *
        (skoXC1 * 1/4 * (-210 + skoXC1 * 1/4 * (skoXC1 * 1/4)))))) <
     -1/125000000 + sin(skoXC1) * -2),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))) <
     -1/125000000 + sin(skoXC1) * -2 | 314159/50000 < skoXC1),
    inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (~ lgen(0, X_000083, sin(X_000082)) | X_000083 <= sin(X_000082)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_17, plain,
    (X_000082 < 0 |
     ~ lgen(0, X_000083,
            X_000082 - X_000082 ^ 3 / 6 + X_000082 ^ 5 / 120 -
            X_000082 ^ 7 / 5040 + X_000082 ^ 9 / 362880 -
            X_000082 ^ 11 / 39916800 + X_000082 ^ 13 / 6227020800 -
            X_000082 ^ 15 / 1307674368000) |
     lgen(0, X_000083, sin(X_000082))),
    inference(subst, [], [sin_lower_bound_3_pos])).

cnf(refute_0_18, plain,
    (X_000082 < 0 |
     ~ lgen(0, X_000083,
            X_000082 - X_000082 ^ 3 / 6 + X_000082 ^ 5 / 120 -
            X_000082 ^ 7 / 5040 + X_000082 ^ 9 / 362880 -
            X_000082 ^ 11 / 39916800 + X_000082 ^ 13 / 6227020800 -
            X_000082 ^ 15 / 1307674368000) | X_000083 <= sin(X_000082)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (X_000082 *
     (1 +
      X_000082 *
      (X_000082 *
       (-1/6 +
        X_000082 *
        (X_000082 *
         (1/120 +
          X_000082 *
          (X_000082 *
           (-1/5040 +
            X_000082 *
            (X_000082 *
             (1/362880 +
              X_000082 *
              (X_000082 *
               (-1/39916800 +
                X_000082 *
                (X_000082 *
                 (1/6227020800 +
                  X_000082 * (X_000082 * -1/1307674368000)))))))))))))) <
     X_000083 | X_000082 < 0 | X_000083 <= sin(X_000082)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 * (skoXC1 * -1/1307674368000)))))))))))))) <
     -1/250000000 + skoXC1 * -1/8 |
     -1/250000000 + skoXC1 * -1/8 <= sin(skoXC1)),
    inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (~ lgen(0, X_000118, tan(X_000117)) | X_000118 <= tan(X_000117)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_22, plain,
    (157/100 < X_000117 | X_000117 < 0 | ~ lgen(0, X_000118, X_000117) |
     lgen(0, X_000118, tan(X_000117))),
    inference(subst, [], [tan_lower_bound_pos0])).

cnf(refute_0_23, plain,
    (157/100 < X_000117 | X_000117 < 0 | ~ lgen(0, X_000118, X_000117) |
     X_000118 <= tan(X_000117)),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (157/100 < X_000117 | X_000117 < X_000118 | X_000117 < 0 |
     X_000118 <= tan(X_000117)), inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 * 1/4 < 0 | 157/100 < skoXC1 * 1/4 |
     skoXC1 * 1/4 < -1/125000000 + sin(skoXC1) * -2 |
     -1/125000000 + sin(skoXC1) * -2 <= tan(skoXC1 * 1/4)),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 * 1/4 < 0 | 157/100 < skoXC1 * 1/4 |
     skoXC1 * 1/4 < -1/125000000 + sin(skoXC1) * -2),
    inference(resolve, [], [refute_0_25, refute_0_8])).

cnf(refute_0_27, plain,
    (skoXC1 < 0 | sin(skoXC1) < -1/250000000 + skoXC1 * -1/8 |
     157/25 < skoXC1), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_29, plain, (skoXC1 < 62831/10000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_30, plain, (157/25 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_31, plain,
    (sin(skoXC1) < -1/250000000 + skoXC1 * -1/8 | 157/25 < skoXC1),
    inference(resolve, [], [refute_0_30, refute_0_27])).

cnf(refute_0_32, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 * (skoXC1 * -1/1307674368000)))))))))))))) <
     -1/250000000 + skoXC1 * -1/8 | 157/25 < skoXC1),
    inference(resolve, [], [refute_0_20, refute_0_31])).

cnf(refute_0_33, plain,
    (skoXC1 < 0 |
     1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (-1/120 +
          skoXC1 *
          (skoXC1 *
           (1/5040 +
            skoXC1 *
            (skoXC1 *
             (-1/362880 +
              skoXC1 *
              (skoXC1 *
               (1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (-1/6227020800 +
                  skoXC1 * (skoXC1 * 1/1307674368000)))))))))))))) |
     157/25 < skoXC1), inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (-1/120 +
          skoXC1 *
          (skoXC1 *
           (1/5040 +
            skoXC1 *
            (skoXC1 *
             (-1/362880 +
              skoXC1 *
              (skoXC1 *
               (1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (-1/6227020800 +
                  skoXC1 * (skoXC1 * 1/1307674368000)))))))))))))) |
     157/25 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_35, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 *
        (skoXC1 *
         (-1/120 +
          skoXC1 *
          (skoXC1 *
           (1/5040 +
            skoXC1 *
            (skoXC1 *
             (-1/362880 +
              skoXC1 *
              (skoXC1 *
               (1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (-1/6227020800 +
                  skoXC1 * (skoXC1 * 1/1307674368000)))))))))))))) |
     157/25 < skoXC1), inference(resolve, [], [refute_0_34, refute_0_33])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000055, sin(X_000054)) | X_000055 <= sin(X_000054)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_37, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            X_000054 - X_000054 ^ 3 / 6 + X_000054 ^ 5 / 120 -
            X_000054 ^ 7 / 5040) | lgen(0, X_000055, sin(X_000054))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_38, plain,
    (X_000054 < 0 |
     ~ lgen(0, X_000055,
            X_000054 - X_000054 ^ 3 / 6 + X_000054 ^ 5 / 120 -
            X_000054 ^ 7 / 5040) | X_000055 <= sin(X_000054)),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (X_000054 *
     (1 +
      X_000054 *
      (X_000054 *
       (-1/6 +
        X_000054 *
        (X_000054 * (1/120 + X_000054 * (X_000054 * -1/5040)))))) <
     X_000055 | X_000054 < 0 | X_000055 <= sin(X_000054)),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) <
     -1/250000000 + skoXC1 * -1/8 |
     -1/250000000 + skoXC1 * -1/8 <= sin(skoXC1)),
    inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) <
     -1/250000000 + skoXC1 * -1/8 | 157/25 < skoXC1),
    inference(resolve, [], [refute_0_40, refute_0_31])).

cnf(refute_0_42, plain,
    (skoXC1 < 0 |
     1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 * (skoXC1 * (-1/120 + skoXC1 * (skoXC1 * 1/5040)))))) |
     157/25 < skoXC1), inference(arithmetic, [], [refute_0_41])).

cnf(refute_0_43, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 * (skoXC1 * (-1/120 + skoXC1 * (skoXC1 * 1/5040)))))) |
     157/25 < skoXC1 | 0 <= skoXC1),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_44, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 * (skoXC1 * (-1/120 + skoXC1 * (skoXC1 * 1/5040)))))) |
     157/25 < skoXC1), inference(resolve, [], [refute_0_43, refute_0_42])).

cnf(refute_0_45, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 * (skoXC1 * (-1/120 + skoXC1 * (skoXC1 * 1/5040)))))) |
     skoXC1 <= 157/25),
    inference(decision, [], [refute_0_28, refute_0_29])).

cnf(refute_0_46, plain,
    (1/250000000 <
     skoXC1 *
     (-9/8 +
      skoXC1 *
      (skoXC1 *
       (1/6 +
        skoXC1 * (skoXC1 * (-1/120 + skoXC1 * (skoXC1 * 1/5040))))))),
    inference(resolve, [], [refute_0_45, refute_0_44])).

cnf(refute_0_47, plain, (314159/50000 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_48, plain,
    (skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))) <
     -1/125000000 + sin(skoXC1) * -2 | 314159/50000 < skoXC1),
    inference(resolve, [], [refute_0_47, refute_0_15])).

cnf(refute_0_49, plain, (skoXC1 <= 314159/50000),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_50, plain,
    (skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) /
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))) <
     -1/125000000 + sin(skoXC1) * -2),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    ((-1/125000000 + sin(skoXC1) * -2) *
     (-10395 +
      skoXC1 *
      (skoXC1 *
       (4725/16 +
        skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))) <
     skoXC1 *
     (-10395/4 +
      skoXC1 * (skoXC1 * (315/16 + skoXC1 * (skoXC1 * -21/1024)))) |
     0 <=
     -10395 +
     skoXC1 *
     (skoXC1 *
      (4725/16 +
       skoXC1 * (skoXC1 * (-105/128 + skoXC1 * (skoXC1 * 1/4096)))))),
    inference(resolve, [], [refute_0_6, refute_0_50])).

cnf(refute_0_52, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     sin(skoXC1) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/16 +
       skoXC1 * (skoXC1 * (105/128 + skoXC1 * (skoXC1 * -1/4096))))) <=
     -10395), inference(arithmetic, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (-10395 <
     skoXC1 *
     (skoXC1 *
      (-4725/16 +
       skoXC1 * (skoXC1 * (105/128 + skoXC1 * (skoXC1 * -1/4096)))))),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_54, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     sin(skoXC1) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain,
    (sin(skoXC1) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     0 <=
     -20790 +
     skoXC1 *
     (skoXC1 *
      (4725/8 +
       skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))),
    inference(split, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (sin(skoXC1) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) <=
     -20790), inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (-20790 <
     skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048)))))),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_58, plain,
    (sin(skoXC1) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_56, refute_0_57])).

cnf(refute_0_59, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_5, refute_0_58])).

cnf(refute_0_60, plain, (0 <= skoXC1),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_61, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_60, refute_0_59])).

cnf(refute_0_62, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 *
        (skoXC1 *
         (1/120 +
          skoXC1 *
          (skoXC1 *
           (-1/5040 +
            skoXC1 *
            (skoXC1 *
             (1/362880 +
              skoXC1 *
              (skoXC1 *
               (-1/39916800 +
                skoXC1 *
                (skoXC1 *
                 (1/6227020800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/1307674368000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/355687428096000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/121645100408832000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/51090942171709440000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           25852016738884976640000)))))))))))))))))))))) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     0 <=
     -20790 +
     skoXC1 *
     (skoXC1 *
      (4725/8 +
       skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))),
    inference(resolve, [], [refute_0_0, refute_0_61])).

cnf(refute_0_63, plain,
    (2079/25000000 <
     skoXC1 *
     (-93555/4 +
      skoXC1 *
      (189/80000000 +
       skoXC1 *
       (65205/16 +
        skoXC1 *
        (-21/3200000000 +
         skoXC1 *
         (-279909/1024 +
          skoXC1 *
          (1/512000000000 +
           skoXC1 *
           (19089/2048 +
            skoXC1 *
            (skoXC1 *
             (-771/4096 +
              skoXC1 *
              (skoXC1 *
               (203/81920 +
                skoXC1 *
                (skoXC1 *
                 (-100753/4428103680 +
                  skoXC1 *
                  (skoXC1 *
                   (11629/75910348800 +
                    skoXC1 *
                    (skoXC1 *
                     (-14197/18066663014400 +
                      skoXC1 *
                      (skoXC1 *
                       (79/24964843438080 +
                        skoXC1 *
                        (skoXC1 *
                         (-197/19222929447321600 +
                          skoXC1 *
                          (skoXC1 *
                           (25999/954995134942937088000 +
                            skoXC1 *
                            (skoXC1 *
                             (-59/1000471093749743616000 +
                              skoXC1 *
                              (skoXC1 *
                               (1933/26472465140618216079360000 +
                                skoXC1 *
                                (skoXC1 *
                                 -1/
                                 52944930281236432158720000)))))))))))))))))))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) <=
     -20790), inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     0 <=
     -20790 +
     skoXC1 *
     (skoXC1 *
      (4725/8 +
       skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))) |
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) <=
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_65, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) <=
     sin(skoXC1)), inference(subst, [], [refute_0_39])).

cnf(refute_0_66, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_65, refute_0_58])).

cnf(refute_0_67, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) <
     (2079/25000000 +
      skoXC1 *
      (10395/4 +
       skoXC1 *
       (-189/80000000 +
        skoXC1 *
        (-315/16 +
         skoXC1 *
         (21/3200000000 +
          skoXC1 * (21/1024 + skoXC1 * -1/512000000000)))))) /
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048))))))),
    inference(resolve, [], [refute_0_60, refute_0_66])).

cnf(refute_0_68, plain,
    (2079/25000000 +
     skoXC1 *
     (10395/4 +
      skoXC1 *
      (-189/80000000 +
       skoXC1 *
       (-315/16 +
        skoXC1 *
        (21/3200000000 +
         skoXC1 * (21/1024 + skoXC1 * -1/512000000000))))) <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/6 +
        skoXC1 * (skoXC1 * (1/120 + skoXC1 * (skoXC1 * -1/5040)))))) *
     (-20790 +
      skoXC1 *
      (skoXC1 *
       (4725/8 +
        skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))) |
     0 <=
     -20790 +
     skoXC1 *
     (skoXC1 *
      (4725/8 +
       skoXC1 * (skoXC1 * (-105/64 + skoXC1 * (skoXC1 * 1/2048)))))),
    inference(resolve, [], [refute_0_64, refute_0_67])).

cnf(refute_0_69, plain,
    (2079/25000000 <
     skoXC1 *
     (-93555/4 +
      skoXC1 *
      (189/80000000 +
       skoXC1 *
       (65205/16 +
        skoXC1 *
        (-21/3200000000 +
         skoXC1 *
         (-279909/1024 +
          skoXC1 *
          (1/512000000000 +
           skoXC1 *
           (19089/2048 +
            skoXC1 *
            (skoXC1 *
             (-1609/12288 +
              skoXC1 *
              (skoXC1 *
               (27/81920 + skoXC1 * (skoXC1 * -1/10321920)))))))))))) |
     skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) <=
     -20790), inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (-20790 <
     skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) |
     2079/25000000 <
     skoXC1 *
     (-93555/4 +
      skoXC1 *
      (189/80000000 +
       skoXC1 *
       (65205/16 +
        skoXC1 *
        (-21/3200000000 +
         skoXC1 *
         (-279909/1024 +
          skoXC1 *
          (1/512000000000 +
           skoXC1 *
           (19089/2048 +
            skoXC1 *
            (skoXC1 *
             (-1609/12288 +
              skoXC1 *
              (skoXC1 *
               (27/81920 + skoXC1 * (skoXC1 * -1/10321920))))))))))))),
    inference(decision, [],
              [refute_0_35, refute_0_46, refute_0_28, refute_0_29])).

cnf(refute_0_71, plain,
    (2079/25000000 <
     skoXC1 *
     (-93555/4 +
      skoXC1 *
      (189/80000000 +
       skoXC1 *
       (65205/16 +
        skoXC1 *
        (-21/3200000000 +
         skoXC1 *
         (-279909/1024 +
          skoXC1 *
          (1/512000000000 +
           skoXC1 *
           (19089/2048 +
            skoXC1 *
            (skoXC1 *
             (-1609/12288 +
              skoXC1 *
              (skoXC1 *
               (27/81920 + skoXC1 * (skoXC1 * -1/10321920))))))))))))),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) <=
     -20790 |
     skoXC1 *
     (-93555/4 +
      skoXC1 *
      (189/80000000 +
       skoXC1 *
       (65205/16 +
        skoXC1 *
        (-21/3200000000 +
         skoXC1 *
         (-279909/1024 +
          skoXC1 *
          (1/512000000000 +
           skoXC1 *
           (19089/2048 +
            skoXC1 *
            (skoXC1 *
             (-771/4096 +
              skoXC1 *
              (skoXC1 *
               (203/81920 +
                skoXC1 *
                (skoXC1 *
                 (-100753/4428103680 +
                  skoXC1 *
                  (skoXC1 *
                   (11629/75910348800 +
                    skoXC1 *
                    (skoXC1 *
                     (-14197/18066663014400 +
                      skoXC1 *
                      (skoXC1 *
                       (79/24964843438080 +
                        skoXC1 *
                        (skoXC1 *
                         (-197/19222929447321600 +
                          skoXC1 *
                          (skoXC1 *
                           (25999/954995134942937088000 +
                            skoXC1 *
                            (skoXC1 *
                             (-59/1000471093749743616000 +
                              skoXC1 *
                              (skoXC1 *
                               (1933/26472465140618216079360000 +
                                skoXC1 *
                                (skoXC1 *
                                 -1/
                                 52944930281236432158720000)))))))))))))))))))))))))))) <=
     2079/25000000),
    inference(decision, [],
              [refute_0_71, refute_0_35, refute_0_46, refute_0_28,
               refute_0_29])).

cnf(refute_0_73, plain,
    (skoXC1 *
     (skoXC1 *
      (-4725/8 +
       skoXC1 * (skoXC1 * (105/64 + skoXC1 * (skoXC1 * -1/2048))))) <=
     -20790), inference(resolve, [], [refute_0_72, refute_0_63])).

cnf(refute_0_74, plain, ($false),
    inference(resolve, [], [refute_0_73, refute_0_57])).
% SZS output end CNFRefutation for trigpoly-351-1a.tptp
