% SZS output start CNFRefutation for trigpoly-352-1-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120, Y) |
     lgen(R, sin(X), Y))).

cnf(sin_upper_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800, Y) |
     lgen(R, sin(X), Y))).

cnf(sin_upper_bound_5_pos, axiom,
    (X < 0 |
     ~ lgen(R,
            X - 1/6 * X ^ 3 + 1/120 * X ^ 5 - 1/5040 * X ^ 7 +
            1/362880 * X ^ 9 - 1/39916800 * X ^ 11 +
            1/6227020800 * X ^ 13 - 1/1307674368000 * X ^ 15 +
            1/355687428096000 * X ^ 17 - 1/121645100408832000 * X ^ 19 +
            1/51090942171709440000 * X ^ 21, Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040) |
     lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_1, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320,
            Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_3, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200 +
            X ^ 16 / 20922789888000, Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_5, axiom,
    (~ lgen(R,
            1 - 1/2 * X ^ 2 + 1/24 * X ^ 4 - 1/720 * X ^ 6 +
            1/40320 * X ^ 8 - 1/3628800 * X ^ 10 + 1/479001600 * X ^ 12 -
            1/87178291200 * X ^ 14 + 1/20922789888000 * X ^ 16 -
            1/6402373705728000 * X ^ 18 + 1/2432902008176640000 * X ^ 20 -
            1/1124000727777607680000 * X ^ 22 +
            1/620448401733239439360000 * X ^ 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_1, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_3, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_5, axiom,
    (~ lgen(R, Y,
            1 - 1/2 * X ^ 2 + 1/24 * X ^ 4 - 1/720 * X ^ 6 +
            1/40320 * X ^ 8 - 1/3628800 * X ^ 10 + 1/479001600 * X ^ 12 -
            1/87178291200 * X ^ 14 + 1/20922789888000 * X ^ 16 -
            1/6402373705728000 * X ^ 18 + 1/2432902008176640000 * X ^ 20 -
            1/1124000727777607680000 * X ^ 22) | lgen(R, Y, cos(X)))).

fof(trig_poly_352, conjecture,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        abs(1/2 + cos(X)) <= 50000001/50000000 / (2 * sin(X / 2))))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        abs(1/2 + cos(X)) <= 50000001/50000000 / (2 * sin(X / 2)))),
    inference(strip, [], [trig_poly_352])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 62831/10000) =>
          abs(1/2 + cos(X)) <= 50000001/50000000 / (2 * sin(X / 2)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 62831/10000 &
        50000001/50000000 / (2 * sin(X / 2)) < abs(1/2 + cos(X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 62831/10000 &
     50000001/50000000 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1)) &
     0 < skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (50000001/50000000 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 < 62831/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/48 +
        skoXC1 *
        (skoXC1 *
         (-1/3840 +
          skoXC1 *
          (skoXC1 *
           (1/645120 +
            skoXC1 *
            (skoXC1 *
             (-1/185794560 +
              skoXC1 *
              (skoXC1 *
               (1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (-1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         -1/
                         107145471557284795514880000)))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) <=
     (50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, cos(X_000046), X_000047) | cos(X_000046) <= X_000047),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0,
            1 - 1/2 * X_000046 ^ 2 + 1/24 * X_000046 ^ 4 -
            1/720 * X_000046 ^ 6 + 1/40320 * X_000046 ^ 8 -
            1/3628800 * X_000046 ^ 10 + 1/479001600 * X_000046 ^ 12 -
            1/87178291200 * X_000046 ^ 14 +
            1/20922789888000 * X_000046 ^ 16 -
            1/6402373705728000 * X_000046 ^ 18 +
            1/2432902008176640000 * X_000046 ^ 20 -
            1/1124000727777607680000 * X_000046 ^ 22 +
            1/620448401733239439360000 * X_000046 ^ 24, X_000047) |
     lgen(0, cos(X_000046), X_000047)),
    inference(subst, [], [cos_upper_bound_5])).

cnf(refute_0_3, plain,
    (~ lgen(0,
            1 - 1/2 * X_000046 ^ 2 + 1/24 * X_000046 ^ 4 -
            1/720 * X_000046 ^ 6 + 1/40320 * X_000046 ^ 8 -
            1/3628800 * X_000046 ^ 10 + 1/479001600 * X_000046 ^ 12 -
            1/87178291200 * X_000046 ^ 14 +
            1/20922789888000 * X_000046 ^ 16 -
            1/6402373705728000 * X_000046 ^ 18 +
            1/2432902008176640000 * X_000046 ^ 20 -
            1/1124000727777607680000 * X_000046 ^ 22 +
            1/620448401733239439360000 * X_000046 ^ 24, X_000047) |
     cos(X_000046) <= X_000047),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000047 <
     1 +
     X_000046 *
     (X_000046 *
      (-1/2 +
       X_000046 *
       (X_000046 *
        (1/24 +
         X_000046 *
         (X_000046 *
          (-1/720 +
           X_000046 *
           (X_000046 *
            (1/40320 +
             X_000046 *
             (X_000046 *
              (-1/3628800 +
               X_000046 *
               (X_000046 *
                (1/479001600 +
                 X_000046 *
                 (X_000046 *
                  (-1/87178291200 +
                   X_000046 *
                   (X_000046 *
                    (1/20922789888000 +
                     X_000046 *
                     (X_000046 *
                      (-1/6402373705728000 +
                       X_000046 *
                       (X_000046 *
                        (1/2432902008176640000 +
                         X_000046 *
                         (X_000046 *
                          (-1/1124000727777607680000 +
                           X_000046 *
                           (X_000046 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(X_000046) <= X_000047), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(skoXC1) <=
     (50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000)))))))))))))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         107145471557284795514880000)))))))))))))))))))) *
     (1 + cos(skoXC1) * 2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         107145471557284795514880000)))))))))))))))))))) <=
     50000001/50000000 / (1 + cos(skoXC1) * 2) | 1 + cos(skoXC1) * 2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    (50000001/50000000 < (1/2 + cos(skoXC1)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0 |
     1/2 + cos(skoXC1) <= 50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_8, plain,
    (50000001/50000000 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < abs(1/2 + cos(skoXC1))),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (1/2 + cos(skoXC1) < 0 | abs(1/2 + cos(skoXC1)) = 1/2 + cos(skoXC1)),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_11, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) != 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) <=
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    introduced(tautology, [equality])).

cnf(refute_0_12, plain,
    (1/2 + cos(skoXC1) < 0 |
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) <=
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (1/2 + cos(skoXC1) < 0 |
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < 1/2 + cos(skoXC1)),
    inference(resolve, [], [refute_0_12, refute_0_9])).

cnf(refute_0_14, plain,
    (cos(skoXC1) < -1/2 |
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < 1/2 + cos(skoXC1)),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 *
           (1/322560 +
            skoXC1 *
            (skoXC1 *
             (-1/92897280 +
              skoXC1 *
              (skoXC1 *
               (1/40874803200 +
                skoXC1 * (skoXC1 * -1/25505877196800)))))))))))) |
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_16, plain,
    (~ lgen(0, X_000040, cos(X_000039)) | X_000040 <= cos(X_000039)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_17, plain,
    (~ lgen(0, X_000040,
            1 - 1/2 * X_000039 ^ 2 + 1/24 * X_000039 ^ 4 -
            1/720 * X_000039 ^ 6 + 1/40320 * X_000039 ^ 8 -
            1/3628800 * X_000039 ^ 10 + 1/479001600 * X_000039 ^ 12 -
            1/87178291200 * X_000039 ^ 14 +
            1/20922789888000 * X_000039 ^ 16 -
            1/6402373705728000 * X_000039 ^ 18 +
            1/2432902008176640000 * X_000039 ^ 20 -
            1/1124000727777607680000 * X_000039 ^ 22) |
     lgen(0, X_000040, cos(X_000039))),
    inference(subst, [], [cos_lower_bound_5])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000040,
            1 - 1/2 * X_000039 ^ 2 + 1/24 * X_000039 ^ 4 -
            1/720 * X_000039 ^ 6 + 1/40320 * X_000039 ^ 8 -
            1/3628800 * X_000039 ^ 10 + 1/479001600 * X_000039 ^ 12 -
            1/87178291200 * X_000039 ^ 14 +
            1/20922789888000 * X_000039 ^ 16 -
            1/6402373705728000 * X_000039 ^ 18 +
            1/2432902008176640000 * X_000039 ^ 20 -
            1/1124000727777607680000 * X_000039 ^ 22) |
     X_000040 <= cos(X_000039)),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (1 +
     X_000039 *
     (X_000039 *
      (-1/2 +
       X_000039 *
       (X_000039 *
        (1/24 +
         X_000039 *
         (X_000039 *
          (-1/720 +
           X_000039 *
           (X_000039 *
            (1/40320 +
             X_000039 *
             (X_000039 *
              (-1/3628800 +
               X_000039 *
               (X_000039 *
                (1/479001600 +
                 X_000039 *
                 (X_000039 *
                  (-1/87178291200 +
                   X_000039 *
                   (X_000039 *
                    (1/20922789888000 +
                     X_000039 *
                     (X_000039 *
                      (-1/6402373705728000 +
                       X_000039 *
                       (X_000039 *
                        (1/2432902008176640000 +
                         X_000039 *
                         (X_000039 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     X_000040 | X_000040 <= cos(X_000039)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) *
     (-1 + cos(skoXC1) * -2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) <=
     50000001/50000000 / (-1 + cos(skoXC1) * -2) |
     -1 + cos(skoXC1) * -2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_22, plain,
    (0 <= 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) = -(1/2 + cos(skoXC1))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_23, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < -(1/2 + cos(skoXC1)) |
     abs(1/2 + cos(skoXC1)) != -(1/2 + cos(skoXC1)) |
     abs(1/2 + cos(skoXC1)) <=
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    introduced(tautology, [equality])).

cnf(refute_0_24, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < -(1/2 + cos(skoXC1)) |
     0 <= 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) <=
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < -(1/2 + cos(skoXC1)) |
     0 <= 1/2 + cos(skoXC1)),
    inference(resolve, [], [refute_0_24, refute_0_9])).

cnf(refute_0_26, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) <
     -1/2 + cos(skoXC1) * -1 | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (50000001/50000000 <
     (-1/2 + cos(skoXC1) * -1) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0 |
     -1/2 + cos(skoXC1) * -1 <=
     50000001/50000000 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_28, plain,
    (50000001/50000000 <
     (-1/2 + cos(skoXC1) * -1) * (sin(skoXC1 * 1/2) * 2) |
     -1/2 <= cos(skoXC1) | sin(skoXC1 * 1/2) * 2 <= 0),
    inference(resolve, [], [refute_0_27, refute_0_26])).

cnf(refute_0_29, plain,
    (50000001/50000000 < sin(skoXC1 * 1/2) * (-1 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1) | sin(skoXC1 * 1/2) <= 0),
    inference(arithmetic, [], [refute_0_28])).

cnf(refute_0_30, plain,
    (X_000103 < sin(X_000102) | ~ lgen(1, X_000103, sin(X_000102))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_31, plain,
    (X_000102 < 0 |
     ~ lgen(1, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040) | lgen(1, X_000103, sin(X_000102))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_32, plain,
    (X_000102 < 0 | X_000103 < sin(X_000102) |
     ~ lgen(1, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040)),
    inference(resolve, [], [refute_0_31, refute_0_30])).

cnf(refute_0_33, plain,
    (X_000102 < 0 | X_000103 < sin(X_000102) |
     X_000102 *
     (1 +
      X_000102 *
      (X_000102 *
       (-1/6 +
        X_000102 *
        (X_000102 * (1/120 + X_000102 * (X_000102 * -1/5040)))))) <=
     X_000103), inference(arithmetic, [], [refute_0_32])).

cnf(refute_0_34, plain,
    (skoXC1 * 1/2 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <= 0),
    inference(subst, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0), inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (~ lgen(0, X_000026, cos(X_000025)) | X_000026 <= cos(X_000025)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_37, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 - X_000025 ^ 6 / 720 +
            X_000025 ^ 8 / 40320 - X_000025 ^ 10 / 3628800 +
            X_000025 ^ 12 / 479001600 - X_000025 ^ 14 / 87178291200) |
     lgen(0, X_000026, cos(X_000025))),
    inference(subst, [], [cos_lower_bound_3])).

cnf(refute_0_38, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 - X_000025 ^ 6 / 720 +
            X_000025 ^ 8 / 40320 - X_000025 ^ 10 / 3628800 +
            X_000025 ^ 12 / 479001600 - X_000025 ^ 14 / 87178291200) |
     X_000026 <= cos(X_000025)),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (1 +
     X_000025 *
     (X_000025 *
      (-1/2 +
       X_000025 *
       (X_000025 *
        (1/24 +
         X_000025 *
         (X_000025 *
          (-1/720 +
           X_000025 *
           (X_000025 *
            (1/40320 +
             X_000025 *
             (X_000025 *
              (-1/3628800 +
               X_000025 *
               (X_000025 *
                (1/479001600 +
                 X_000025 * (X_000025 * -1/87178291200))))))))))))) <
     X_000026 | X_000026 <= cos(X_000025)),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) < -1/2 |
     -1/2 <= cos(skoXC1)), inference(subst, [], [refute_0_39])).

cnf(refute_0_41, plain,
    (3/2 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (~ lgen(0, X_000012, cos(X_000011)) | X_000012 <= cos(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_43, plain,
    (~ lgen(0, X_000012,
            1 - X_000011 ^ 2 / 2 + X_000011 ^ 4 / 24 -
            X_000011 ^ 6 / 720) | lgen(0, X_000012, cos(X_000011))),
    inference(subst, [], [cos_lower_bound_1])).

cnf(refute_0_44, plain,
    (~ lgen(0, X_000012,
            1 - X_000011 ^ 2 / 2 + X_000011 ^ 4 / 24 -
            X_000011 ^ 6 / 720) | X_000012 <= cos(X_000011)),
    inference(resolve, [], [refute_0_43, refute_0_42])).

cnf(refute_0_45, plain,
    (1 +
     X_000011 *
     (X_000011 *
      (-1/2 +
       X_000011 * (X_000011 * (1/24 + X_000011 * (X_000011 * -1/720))))) <
     X_000012 | X_000012 <= cos(X_000011)),
    inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     -1/2 | -1/2 <= cos(skoXC1)), inference(subst, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (3/2 <
     skoXC1 *
     (skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/720))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain, (skoXC1 < 62831/10000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_49, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_50, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_41, refute_0_47, refute_0_48, refute_0_49])).

cnf(refute_0_51, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_35, refute_0_50])).

cnf(refute_0_52, plain, (-1/2 <= cos(skoXC1) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_41, refute_0_47, refute_0_48, refute_0_49])).

cnf(refute_0_53, plain, (0 < sin(skoXC1 * 1/2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_52, refute_0_51])).

cnf(refute_0_54, plain,
    (50000001/50000000 < sin(skoXC1 * 1/2) * (-1 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_29, refute_0_53])).

cnf(refute_0_55, plain,
    (50000001/50000000 / (-1 + cos(skoXC1) * -2) < sin(skoXC1 * 1/2) |
     -1/2 <= cos(skoXC1) | -1 + cos(skoXC1) * -2 <= 0),
    inference(split, [], [refute_0_54])).

cnf(refute_0_56, plain,
    (50000001/50000000 / (-1 + cos(skoXC1) * -2) < sin(skoXC1 * 1/2) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (~ lgen(0, sin(X_000111), X_000112) | sin(X_000111) <= X_000112),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_58, plain,
    (X_000111 < 0 |
     ~ lgen(0,
            X_000111 - X_000111 ^ 3 / 6 + X_000111 ^ 5 / 120 -
            X_000111 ^ 7 / 5040 + X_000111 ^ 9 / 362880 -
            X_000111 ^ 11 / 39916800 + X_000111 ^ 13 / 6227020800,
            X_000112) | lgen(0, sin(X_000111), X_000112)),
    inference(subst, [], [sin_upper_bound_3_pos])).

cnf(refute_0_59, plain,
    (X_000111 < 0 |
     ~ lgen(0,
            X_000111 - X_000111 ^ 3 / 6 + X_000111 ^ 5 / 120 -
            X_000111 ^ 7 / 5040 + X_000111 ^ 9 / 362880 -
            X_000111 ^ 11 / 39916800 + X_000111 ^ 13 / 6227020800,
            X_000112) | sin(X_000111) <= X_000112),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (X_000111 < 0 |
     X_000112 <
     X_000111 *
     (1 +
      X_000111 *
      (X_000111 *
       (-1/6 +
        X_000111 *
        (X_000111 *
         (1/120 +
          X_000111 *
          (X_000111 *
           (-1/5040 +
            X_000111 *
            (X_000111 *
             (1/362880 +
              X_000111 *
              (X_000111 *
               (-1/39916800 +
                X_000111 * (X_000111 * 1/6227020800)))))))))))) |
     sin(X_000111) <= X_000112), inference(arithmetic, [], [refute_0_59])).

cnf(refute_0_61, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/6227020800)))))))))))) |
     sin(skoXC1 * 1/2) <= 50000001/50000000 / (-1 + cos(skoXC1) * -2)),
    inference(subst, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/6227020800)))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_61, refute_0_56])).

cnf(refute_0_63, plain,
    (skoXC1 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_52, refute_0_63])).

cnf(refute_0_65, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) *
     (-1 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1) |
     -1 + cos(skoXC1) * -2 <= 0),
    inference(resolve, [], [refute_0_21, refute_0_64])).

cnf(refute_0_66, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_65])).

cnf(refute_0_67, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 *
           (1/322560 +
            skoXC1 *
            (skoXC1 *
             (-1/92897280 +
              skoXC1 *
              (skoXC1 *
               (1/40874803200 +
                skoXC1 * (skoXC1 * -1/25505877196800))))))))))))),
    inference(split, [], [refute_0_66])).

cnf(refute_0_68, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920)))) |
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_70, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_45])).

cnf(refute_0_71, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (-1 + cos(skoXC1) * -2) |
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <=
     50000001/50000000 / (-1 + cos(skoXC1) * -2) |
     -1 + cos(skoXC1) * -2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_72, plain,
    (~ lgen(0, sin(X_000081), X_000082) | sin(X_000081) <= X_000082),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_73, plain,
    (X_000081 < 0 |
     ~ lgen(0, X_000081 - X_000081 ^ 3 / 6 + X_000081 ^ 5 / 120,
            X_000082) | lgen(0, sin(X_000081), X_000082)),
    inference(subst, [], [sin_upper_bound_1_pos])).

cnf(refute_0_74, plain,
    (X_000081 < 0 |
     ~ lgen(0, X_000081 - X_000081 ^ 3 / 6 + X_000081 ^ 5 / 120,
            X_000082) | sin(X_000081) <= X_000082),
    inference(resolve, [], [refute_0_73, refute_0_72])).

cnf(refute_0_75, plain,
    (X_000081 < 0 |
     X_000082 <
     X_000081 *
     (1 + X_000081 * (X_000081 * (-1/6 + X_000081 * (X_000081 * 1/120)))) |
     sin(X_000081) <= X_000082), inference(arithmetic, [], [refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     sin(skoXC1 * 1/2) <= 50000001/50000000 / (-1 + cos(skoXC1) * -2)),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_76, refute_0_56])).

cnf(refute_0_78, plain,
    (skoXC1 < 0 |
     50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (50000001/50000000 / (-1 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_52, refute_0_78])).

cnf(refute_0_80, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (-1 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1) |
     -1 + cos(skoXC1) * -2 <= 0),
    inference(resolve, [], [refute_0_71, refute_0_79])).

cnf(refute_0_81, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))),
    inference(split, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_49, refute_0_48, refute_0_47, refute_0_41])).

cnf(refute_0_85, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_70, refute_0_85])).

cnf(refute_0_87, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))),
    inference(resolve, [], [refute_0_69, refute_0_86])).

cnf(refute_0_88, plain,
    (50000001/50000000 <
     skoXC1 *
     (-3/2 +
      skoXC1 *
      (skoXC1 *
       (9/16 +
        skoXC1 *
        (skoXC1 *
         (-81/1280 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (50000001/50000000 <
     skoXC1 *
     (-3/2 +
      skoXC1 *
      (skoXC1 *
       (9/16 +
        skoXC1 *
        (skoXC1 *
         (-81/1280 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_49, refute_0_48, refute_0_47, refute_0_41])).

cnf(refute_0_90, plain,
    (50000001/50000000 <
     skoXC1 *
     (-3/2 +
      skoXC1 *
      (skoXC1 *
       (9/16 +
        skoXC1 *
        (skoXC1 *
         (-81/1280 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_88, refute_0_89])).

cnf(refute_0_91, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_90, refute_0_49, refute_0_48, refute_0_47,
               refute_0_41])).

cnf(refute_0_92, plain,
    (cos(skoXC1) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_68, refute_0_91])).

cnf(refute_0_93, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (50000001/50000000 +
      skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/48 +
         skoXC1 *
         (skoXC1 *
          (1/3840 +
           skoXC1 *
           (skoXC1 *
            (-1/645120 +
             skoXC1 *
             (skoXC1 *
              (1/185794560 +
               skoXC1 *
               (skoXC1 *
                (-1/81749606400 +
                 skoXC1 * (skoXC1 * 1/51011754393600))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_20, refute_0_92])).

cnf(refute_0_94, plain,
    (50000001/50000000 +
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 *
            (1/322560 +
             skoXC1 *
             (skoXC1 *
              (-1/92897280 +
               skoXC1 *
               (skoXC1 *
                (1/40874803200 +
                 skoXC1 * (skoXC1 * -1/25505877196800))))))))))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 *
           (1/322560 +
            skoXC1 *
            (skoXC1 *
             (-1/92897280 +
              skoXC1 *
              (skoXC1 *
               (1/40874803200 +
                skoXC1 * (skoXC1 * -1/25505877196800))))))))))))),
    inference(resolve, [], [refute_0_15, refute_0_93])).

cnf(refute_0_95, plain,
    (50000001/50000000 <
     skoXC1 *
     (-3/2 +
      skoXC1 *
      (skoXC1 *
       (9/16 +
        skoXC1 *
        (skoXC1 *
         (-81/1280 +
          skoXC1 *
          (skoXC1 *
           (243/71680 +
            skoXC1 *
            (skoXC1 *
             (-243/2293760 +
              skoXC1 *
              (skoXC1 *
               (2187/1009254400 +
                skoXC1 *
                (skoXC1 *
                 (-6561/209924915200 +
                  skoXC1 *
                  (skoXC1 *
                   (597871/1785411403776000 +
                    skoXC1 *
                    (skoXC1 *
                     (-59347/21424936845312000 +
                      skoXC1 *
                      (skoXC1 *
                       (39821/2185343558221824000 +
                        skoXC1 *
                        (skoXC1 *
                         (-3402739/34878083189220311040000 +
                          skoXC1 *
                          (skoXC1 *
                           (110723/255772610054282280960000 +
                            skoXC1 *
                            (skoXC1 *
                             (-74231/46039069809770810572800000 +
                              skoXC1 *
                              (skoXC1 *
                               (4789/950571382542914971238400000 +
                                skoXC1 *
                                (skoXC1 *
                                 (-9323/723955164944684042095165440000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (2797/
                                    108593274741702606314274816000000 +
                                    skoXC1 *
                                    (skoXC1 *
                                     (-181/
                                      4778104088634914677828091904000000 +
                                      skoXC1 *
                                      (skoXC1 *
                                       1/
                                       28668624531809488066968551424000000)))))))))))))))))))))))))))))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_94])).

cnf(refute_0_96, plain,
    (-1/2 <= cos(skoXC1) |
     skoXC1 *
     (-3/2 +
      skoXC1 *
      (skoXC1 *
       (9/16 +
        skoXC1 *
        (skoXC1 *
         (-81/1280 +
          skoXC1 *
          (skoXC1 *
           (243/71680 +
            skoXC1 *
            (skoXC1 *
             (-243/2293760 +
              skoXC1 *
              (skoXC1 *
               (2187/1009254400 +
                skoXC1 *
                (skoXC1 *
                 (-6561/209924915200 +
                  skoXC1 *
                  (skoXC1 *
                   (597871/1785411403776000 +
                    skoXC1 *
                    (skoXC1 *
                     (-59347/21424936845312000 +
                      skoXC1 *
                      (skoXC1 *
                       (39821/2185343558221824000 +
                        skoXC1 *
                        (skoXC1 *
                         (-3402739/34878083189220311040000 +
                          skoXC1 *
                          (skoXC1 *
                           (110723/255772610054282280960000 +
                            skoXC1 *
                            (skoXC1 *
                             (-74231/46039069809770810572800000 +
                              skoXC1 *
                              (skoXC1 *
                               (4789/950571382542914971238400000 +
                                skoXC1 *
                                (skoXC1 *
                                 (-9323/723955164944684042095165440000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (2797/
                                    108593274741702606314274816000000 +
                                    skoXC1 *
                                    (skoXC1 *
                                     (-181/
                                      4778104088634914677828091904000000 +
                                      skoXC1 *
                                      (skoXC1 *
                                       1/
                                       28668624531809488066968551424000000)))))))))))))))))))))))))))))))))) <=
     50000001/50000000 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(decision, [],
              [refute_0_90, refute_0_49, refute_0_48, refute_0_47,
               refute_0_41])).

cnf(refute_0_97, plain,
    (-1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(resolve, [], [refute_0_96, refute_0_95])).

cnf(refute_0_98, plain, (-1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_97, refute_0_91])).

cnf(refute_0_99, plain,
    (50000001/50000000 / (sin(skoXC1 * 1/2) * 2) < 1/2 + cos(skoXC1)),
    inference(resolve, [], [refute_0_98, refute_0_14])).

cnf(refute_0_100, plain,
    (50000001/50000000 < (1/2 + cos(skoXC1)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0),
    inference(resolve, [], [refute_0_7, refute_0_99])).

cnf(refute_0_101, plain,
    (50000001/50000000 < sin(skoXC1 * 1/2) * (1 + cos(skoXC1) * 2) |
     sin(skoXC1 * 1/2) <= 0), inference(arithmetic, [], [refute_0_100])).

cnf(refute_0_102, plain,
    (X_000133 < sin(X_000132) | ~ lgen(1, X_000133, sin(X_000132))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_103, plain,
    (X_000132 < 0 |
     ~ lgen(1, X_000133,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880 -
            X_000132 ^ 11 / 39916800 + X_000132 ^ 13 / 6227020800 -
            X_000132 ^ 15 / 1307674368000) |
     lgen(1, X_000133, sin(X_000132))),
    inference(subst, [], [sin_lower_bound_3_pos])).

cnf(refute_0_104, plain,
    (X_000132 < 0 | X_000133 < sin(X_000132) |
     ~ lgen(1, X_000133,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880 -
            X_000132 ^ 11 / 39916800 + X_000132 ^ 13 / 6227020800 -
            X_000132 ^ 15 / 1307674368000)),
    inference(resolve, [], [refute_0_103, refute_0_102])).

cnf(refute_0_105, plain,
    (X_000132 < 0 | X_000133 < sin(X_000132) |
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 *
           (-1/5040 +
            X_000132 *
            (X_000132 *
             (1/362880 +
              X_000132 *
              (X_000132 *
               (-1/39916800 +
                X_000132 *
                (X_000132 *
                 (1/6227020800 +
                  X_000132 * (X_000132 * -1/1307674368000)))))))))))))) <=
     X_000133), inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (skoXC1 * 1/2 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 *
                (skoXC1 * 1/2 *
                 (1/6227020800 +
                  skoXC1 * 1/2 *
                  (skoXC1 * 1/2 * -1/1307674368000)))))))))))))) <= 0),
    inference(subst, [], [refute_0_105])).

cnf(refute_0_107, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 * (skoXC1 * -1/42849873690624000)))))))))))))) <=
     0), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (cos(X_000018) < X_000019 | ~ lgen(1, cos(X_000018), X_000019)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_109, plain,
    (~ lgen(1,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) |
     lgen(1, cos(X_000018), X_000019)),
    inference(subst, [], [cos_upper_bound_1])).

cnf(refute_0_110, plain,
    (cos(X_000018) < X_000019 |
     ~ lgen(1,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019)),
    inference(resolve, [], [refute_0_109, refute_0_108])).

cnf(refute_0_111, plain,
    (cos(X_000018) < X_000019 |
     X_000019 <=
     1 +
     X_000018 *
     (X_000018 *
      (-1/2 +
       X_000018 *
       (X_000018 *
        (1/24 +
         X_000018 *
         (X_000018 * (-1/720 + X_000018 * (X_000018 * 1/40320)))))))),
    inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (cos(skoXC1) < -1/2 |
     -1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(subst, [], [refute_0_111])).

cnf(refute_0_113, plain,
    (-1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_98, refute_0_112])).

cnf(refute_0_114, plain,
    (skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/720 + skoXC1 * (skoXC1 * -1/40320))))))) <=
     3/2), inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0),
    inference(decision, [], [refute_0_114, refute_0_48, refute_0_49])).

cnf(refute_0_116, plain,
    (0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0), inference(resolve, [], [refute_0_115, refute_0_35])).

cnf(refute_0_117, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 * (skoXC1 * -1/42849873690624000)))))))))))))) |
     0 < sin(skoXC1 * 1/2)),
    inference(decision, [],
              [refute_0_116, refute_0_114, refute_0_48, refute_0_49])).

cnf(refute_0_118, plain, (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_107, refute_0_117])).

cnf(refute_0_119, plain, (0 < sin(skoXC1 * 1/2) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_116, refute_0_114, refute_0_48, refute_0_49])).

cnf(refute_0_120, plain, (0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_119, refute_0_118])).

cnf(refute_0_121, plain,
    (50000001/50000000 < sin(skoXC1 * 1/2) * (1 + cos(skoXC1) * 2)),
    inference(resolve, [], [refute_0_101, refute_0_120])).

cnf(refute_0_122, plain,
    (50000001/50000000 / (1 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2) |
     1 + cos(skoXC1) * 2 <= 0), inference(split, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (50000001/50000000 / (1 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2) |
     cos(skoXC1) <= -1/2), inference(arithmetic, [], [refute_0_122])).

cnf(refute_0_124, plain,
    (50000001/50000000 < 0 | 1 + cos(skoXC1) * 2 < 0 |
     0 < 1 + cos(skoXC1) * 2), inference(split, [], [refute_0_121])).

cnf(refute_0_125, plain, (cos(skoXC1) < -1/2 | -1/2 < cos(skoXC1)),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain, (-1/2 < cos(skoXC1)),
    inference(resolve, [], [refute_0_98, refute_0_125])).

cnf(refute_0_127, plain,
    (50000001/50000000 / (1 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_123, refute_0_126])).

cnf(refute_0_128, plain,
    (~ lgen(0, sin(X_000139), X_000140) | sin(X_000139) <= X_000140),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_129, plain,
    (X_000139 < 0 |
     ~ lgen(0,
            X_000139 - 1/6 * X_000139 ^ 3 + 1/120 * X_000139 ^ 5 -
            1/5040 * X_000139 ^ 7 + 1/362880 * X_000139 ^ 9 -
            1/39916800 * X_000139 ^ 11 + 1/6227020800 * X_000139 ^ 13 -
            1/1307674368000 * X_000139 ^ 15 +
            1/355687428096000 * X_000139 ^ 17 -
            1/121645100408832000 * X_000139 ^ 19 +
            1/51090942171709440000 * X_000139 ^ 21, X_000140) |
     lgen(0, sin(X_000139), X_000140)),
    inference(subst, [], [sin_upper_bound_5_pos])).

cnf(refute_0_130, plain,
    (X_000139 < 0 |
     ~ lgen(0,
            X_000139 - 1/6 * X_000139 ^ 3 + 1/120 * X_000139 ^ 5 -
            1/5040 * X_000139 ^ 7 + 1/362880 * X_000139 ^ 9 -
            1/39916800 * X_000139 ^ 11 + 1/6227020800 * X_000139 ^ 13 -
            1/1307674368000 * X_000139 ^ 15 +
            1/355687428096000 * X_000139 ^ 17 -
            1/121645100408832000 * X_000139 ^ 19 +
            1/51090942171709440000 * X_000139 ^ 21, X_000140) |
     sin(X_000139) <= X_000140),
    inference(resolve, [], [refute_0_129, refute_0_128])).

cnf(refute_0_131, plain,
    (X_000139 < 0 |
     X_000140 <
     X_000139 *
     (1 +
      X_000139 *
      (X_000139 *
       (-1/6 +
        X_000139 *
        (X_000139 *
         (1/120 +
          X_000139 *
          (X_000139 *
           (-1/5040 +
            X_000139 *
            (X_000139 *
             (1/362880 +
              X_000139 *
              (X_000139 *
               (-1/39916800 +
                X_000139 *
                (X_000139 *
                 (1/6227020800 +
                  X_000139 *
                  (X_000139 *
                   (-1/1307674368000 +
                    X_000139 *
                    (X_000139 *
                     (1/355687428096000 +
                      X_000139 *
                      (X_000139 *
                       (-1/121645100408832000 +
                        X_000139 *
                        (X_000139 *
                         1/51090942171709440000)))))))))))))))))))) |
     sin(X_000139) <= X_000140),
    inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 *
                (skoXC1 * 1/2 *
                 (1/6227020800 +
                  skoXC1 * 1/2 *
                  (skoXC1 * 1/2 *
                   (-1/1307674368000 +
                    skoXC1 * 1/2 *
                    (skoXC1 * 1/2 *
                     (1/355687428096000 +
                      skoXC1 * 1/2 *
                      (skoXC1 * 1/2 *
                       (-1/121645100408832000 +
                        skoXC1 * 1/2 *
                        (skoXC1 * 1/2 *
                         1/51090942171709440000)))))))))))))))))))) |
     sin(skoXC1 * 1/2) <= 50000001/50000000 / (1 + cos(skoXC1) * 2)),
    inference(subst, [], [refute_0_131])).

cnf(refute_0_133, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 *
                (skoXC1 * 1/2 *
                 (1/6227020800 +
                  skoXC1 * 1/2 *
                  (skoXC1 * 1/2 *
                   (-1/1307674368000 +
                    skoXC1 * 1/2 *
                    (skoXC1 * 1/2 *
                     (1/355687428096000 +
                      skoXC1 * 1/2 *
                      (skoXC1 * 1/2 *
                       (-1/121645100408832000 +
                        skoXC1 * 1/2 *
                        (skoXC1 * 1/2 *
                         1/51090942171709440000))))))))))))))))))))),
    inference(resolve, [], [refute_0_132, refute_0_127])).

cnf(refute_0_134, plain,
    (skoXC1 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         107145471557284795514880000))))))))))))))))))))),
    inference(arithmetic, [], [refute_0_133])).

cnf(refute_0_135, plain,
    (~ lgen(0, cos(X_000018), X_000019) | cos(X_000018) <= X_000019),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_136, plain,
    (~ lgen(0,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) |
     lgen(0, cos(X_000018), X_000019)),
    inference(subst, [], [cos_upper_bound_1])).

cnf(refute_0_137, plain,
    (~ lgen(0,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) | cos(X_000018) <= X_000019),
    inference(resolve, [], [refute_0_136, refute_0_135])).

cnf(refute_0_138, plain,
    (X_000019 <
     1 +
     X_000018 *
     (X_000018 *
      (-1/2 +
       X_000018 *
       (X_000018 *
        (1/24 +
         X_000018 *
         (X_000018 * (-1/720 + X_000018 * (X_000018 * 1/40320))))))) |
     cos(X_000018) <= X_000019),
    inference(arithmetic, [], [refute_0_137])).

cnf(refute_0_139, plain,
    (-1/2 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <= -1/2), inference(subst, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (-1/2 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_139, refute_0_126])).

cnf(refute_0_141, plain,
    (-3/2 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(arithmetic, [], [refute_0_140])).

cnf(refute_0_142, plain,
    (~ lgen(0, cos(X_000032), X_000033) | cos(X_000032) <= X_000033),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_143, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033) |
     lgen(0, cos(X_000032), X_000033)),
    inference(subst, [], [cos_upper_bound_3])).

cnf(refute_0_144, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033) |
     cos(X_000032) <= X_000033),
    inference(resolve, [], [refute_0_143, refute_0_142])).

cnf(refute_0_145, plain,
    (X_000033 <
     1 +
     X_000032 *
     (X_000032 *
      (-1/2 +
       X_000032 *
       (X_000032 *
        (1/24 +
         X_000032 *
         (X_000032 *
          (-1/720 +
           X_000032 *
           (X_000032 *
            (1/40320 +
             X_000032 *
             (X_000032 *
              (-1/3628800 +
               X_000032 *
               (X_000032 *
                (1/479001600 +
                 X_000032 *
                 (X_000032 *
                  (-1/87178291200 +
                   X_000032 * (X_000032 * 1/20922789888000))))))))))))))) |
     cos(X_000032) <= X_000033),
    inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain,
    (-1/2 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     cos(skoXC1) <= -1/2), inference(subst, [], [refute_0_145])).

cnf(refute_0_147, plain,
    (-1/2 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(resolve, [], [refute_0_146, refute_0_126])).

cnf(refute_0_148, plain,
    (-3/2 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(arithmetic, [], [refute_0_147])).

cnf(refute_0_149, plain, (0 <= skoXC1),
    inference(decision, [],
              [refute_0_49, refute_0_48, refute_0_114, refute_0_141,
               refute_0_148])).

cnf(refute_0_150, plain,
    (50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         107145471557284795514880000))))))))))))))))))))),
    inference(resolve, [], [refute_0_149, refute_0_134])).

cnf(refute_0_151, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         107145471557284795514880000)))))))))))))))))))) *
     (1 + cos(skoXC1) * 2) | 1 + cos(skoXC1) * 2 <= 0),
    inference(resolve, [], [refute_0_6, refute_0_150])).

cnf(refute_0_152, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/48 +
        skoXC1 *
        (skoXC1 *
         (-1/3840 +
          skoXC1 *
          (skoXC1 *
           (1/645120 +
            skoXC1 *
            (skoXC1 *
             (-1/185794560 +
              skoXC1 *
              (skoXC1 *
               (1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (-1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         -1/
                         107145471557284795514880000)))))))))))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) |
     cos(skoXC1) <= -1/2), inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/48 +
        skoXC1 *
        (skoXC1 *
         (-1/3840 +
          skoXC1 *
          (skoXC1 *
           (1/645120 +
            skoXC1 *
            (skoXC1 *
             (-1/185794560 +
              skoXC1 *
              (skoXC1 *
               (1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (-1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         -1/
                         107145471557284795514880000)))))))))))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000)))))))))))))))))))))),
    inference(resolve, [], [refute_0_152, refute_0_126])).

cnf(refute_0_154, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) <
     cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0), inference(split, [], [refute_0_153])).

cnf(refute_0_155, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) <=
     (50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_156, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <=
     (50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(subst, [], [refute_0_138])).

cnf(refute_0_157, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (1 + cos(skoXC1) * 2) |
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <=
     50000001/50000000 / (1 + cos(skoXC1) * 2) | 1 + cos(skoXC1) * 2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_158, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     sin(skoXC1 * 1/2) <= 50000001/50000000 / (1 + cos(skoXC1) * 2)),
    inference(subst, [], [refute_0_75])).

cnf(refute_0_159, plain,
    (skoXC1 * 1/2 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120))))),
    inference(resolve, [], [refute_0_158, refute_0_127])).

cnf(refute_0_160, plain,
    (skoXC1 < 0 |
     50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))),
    inference(arithmetic, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (50000001/50000000 / (1 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))),
    inference(resolve, [], [refute_0_149, refute_0_160])).

cnf(refute_0_162, plain,
    (50000001/50000000 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (1 + cos(skoXC1) * 2) | 1 + cos(skoXC1) * 2 <= 0),
    inference(resolve, [], [refute_0_157, refute_0_161])).

cnf(refute_0_163, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840)))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     cos(skoXC1) <= -1/2), inference(arithmetic, [], [refute_0_162])).

cnf(refute_0_164, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840)))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(resolve, [], [refute_0_163, refute_0_126])).

cnf(refute_0_165, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(split, [], [refute_0_164])).

cnf(refute_0_166, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))),
    inference(decision, [],
              [refute_0_148, refute_0_141, refute_0_114, refute_0_48,
               refute_0_49])).

cnf(refute_0_167, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     cos(skoXC1)), inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_156, refute_0_167])).

cnf(refute_0_169, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 + skoXC1 * (skoXC1 * (1/48 + skoXC1 * (skoXC1 * -1/3840)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(resolve, [], [refute_0_155, refute_0_168])).

cnf(refute_0_170, plain,
    (50000001/50000000 <
     skoXC1 *
     (3/2 +
      skoXC1 *
      (skoXC1 *
       (-9/16 +
        skoXC1 *
        (skoXC1 *
         (81/1280 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400)))))))))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     50000001/50000000 <
     skoXC1 *
     (3/2 +
      skoXC1 *
      (skoXC1 *
       (-9/16 +
        skoXC1 *
        (skoXC1 *
         (81/1280 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400))))))))))))),
    inference(decision, [],
              [refute_0_148, refute_0_141, refute_0_114, refute_0_48,
               refute_0_49])).

cnf(refute_0_172, plain,
    (50000001/50000000 <
     skoXC1 *
     (3/2 +
      skoXC1 *
      (skoXC1 *
       (-9/16 +
        skoXC1 *
        (skoXC1 *
         (81/1280 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400))))))))))))),
    inference(resolve, [], [refute_0_170, refute_0_171])).

cnf(refute_0_173, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/53572735778642397757440000))))))))))))))))))))),
    inference(decision, [],
              [refute_0_49, refute_0_48, refute_0_114, refute_0_141,
               refute_0_148, refute_0_172])).

cnf(refute_0_174, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) <
     cos(skoXC1)), inference(resolve, [], [refute_0_154, refute_0_173])).

cnf(refute_0_175, plain,
    ((50000001/50000000 +
      skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/48 +
         skoXC1 *
         (skoXC1 *
          (-1/3840 +
           skoXC1 *
           (skoXC1 *
            (1/645120 +
             skoXC1 *
             (skoXC1 *
              (-1/185794560 +
               skoXC1 *
               (skoXC1 *
                (1/81749606400 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/51011754393600 +
                   skoXC1 *
                   (skoXC1 *
                    (1/42849873690624000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/46620662575398912000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/63777066403145711616000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/
                          107145471557284795514880000))))))))))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000)))))))))))))))))))))))),
    inference(resolve, [], [refute_0_5, refute_0_174])).

cnf(refute_0_176, plain,
    (50000001/50000000 +
     skoXC1 *
     (-1/2 +
      skoXC1 *
      (skoXC1 *
       (1/48 +
        skoXC1 *
        (skoXC1 *
         (-1/3840 +
          skoXC1 *
          (skoXC1 *
           (1/645120 +
            skoXC1 *
            (skoXC1 *
             (-1/185794560 +
              skoXC1 *
              (skoXC1 *
               (1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (-1/51011754393600 +
                  skoXC1 *
                  (skoXC1 *
                   (1/42849873690624000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/46620662575398912000 +
                      skoXC1 *
                      (skoXC1 *
                       (1/63777066403145711616000 +
                        skoXC1 *
                        (skoXC1 *
                         -1/
                         107145471557284795514880000)))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 *
                 (skoXC1 *
                  (1/25505877196800 +
                   skoXC1 *
                   (skoXC1 *
                    (-1/21424936845312000 +
                     skoXC1 *
                     (skoXC1 *
                      (1/23310331287699456000 +
                       skoXC1 *
                       (skoXC1 *
                        (-1/31888533201572855808000 +
                         skoXC1 *
                         (skoXC1 *
                          1/
                          53572735778642397757440000))))))))))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0), inference(resolve, [], [refute_0_0, refute_0_175])).

cnf(refute_0_177, plain,
    (50000001/50000000 <
     skoXC1 *
     (3/2 +
      skoXC1 *
      (skoXC1 *
       (-9/16 +
        skoXC1 *
        (skoXC1 *
         (81/1280 +
          skoXC1 *
          (skoXC1 *
           (-243/71680 +
            skoXC1 *
            (skoXC1 *
             (243/2293760 +
              skoXC1 *
              (skoXC1 *
               (-2187/1009254400 +
                skoXC1 *
                (skoXC1 *
                 (6561/209924915200 +
                  skoXC1 *
                  (skoXC1 *
                   (-19683/58778976256000 +
                    skoXC1 *
                    (skoXC1 *
                     (177147/63951526166528000 +
                      skoXC1 *
                      (skoXC1 *
                       (-177147/9720631977312256000 +
                        skoXC1 *
                        (skoXC1 *
                         (531441/5443553907294863360000 +
                          skoXC1 *
                          (skoXC1 *
                           (-170549237/392866729043377583554560000 +
                            skoXC1 *
                            (skoXC1 *
                             (98065811/60239565119984562811699200000 +
                              skoXC1 *
                              (skoXC1 *
                               (-1984888813/
                                380593572428062467844315545600000 +
                                skoXC1 *
                                (skoXC1 *
                                 (305874259/
                                  21313240055971498199281670553600000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (-21449773/
                                    639397201679144945978450116608000000 +
                                    skoXC1 *
                                    (skoXC1 *
                                     (5483501/
                                      84400430621647132869155415392256000000 +
                                      skoXC1 *
                                      (skoXC1 *
                                       (-2558887/
                                        25109128109940022028573736079196160000000 +
                                        skoXC1 *
                                        (skoXC1 *
                                         (1853263/
                                          14462857791325452688458471981616988160000000 +
                                          skoXC1 *
                                          (skoXC1 *
                                           (-213019/
                                            1648765788211101606484265805904336650240000000 +
                                            skoXC1 *
                                            (skoXC1 *
                                             (1151/
                                              10991771921407344043228438706028911001600000000 +
                                              skoXC1 *
                                              (skoXC1 *
                                               (-31/
                                                461654420699108449815594425653214262067200000000 +
                                                skoXC1 *
                                                (skoXC1 *
                                                 1/
                                                 33239118290335808386722798647031426868838400000000)))))))))))))))))))))))))))))))))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0), inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0 |
     skoXC1 *
     (3/2 +
      skoXC1 *
      (skoXC1 *
       (-9/16 +
        skoXC1 *
        (skoXC1 *
         (81/1280 +
          skoXC1 *
          (skoXC1 *
           (-243/71680 +
            skoXC1 *
            (skoXC1 *
             (243/2293760 +
              skoXC1 *
              (skoXC1 *
               (-2187/1009254400 +
                skoXC1 *
                (skoXC1 *
                 (6561/209924915200 +
                  skoXC1 *
                  (skoXC1 *
                   (-19683/58778976256000 +
                    skoXC1 *
                    (skoXC1 *
                     (177147/63951526166528000 +
                      skoXC1 *
                      (skoXC1 *
                       (-177147/9720631977312256000 +
                        skoXC1 *
                        (skoXC1 *
                         (531441/5443553907294863360000 +
                          skoXC1 *
                          (skoXC1 *
                           (-170549237/392866729043377583554560000 +
                            skoXC1 *
                            (skoXC1 *
                             (98065811/60239565119984562811699200000 +
                              skoXC1 *
                              (skoXC1 *
                               (-1984888813/
                                380593572428062467844315545600000 +
                                skoXC1 *
                                (skoXC1 *
                                 (305874259/
                                  21313240055971498199281670553600000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   (-21449773/
                                    639397201679144945978450116608000000 +
                                    skoXC1 *
                                    (skoXC1 *
                                     (5483501/
                                      84400430621647132869155415392256000000 +
                                      skoXC1 *
                                      (skoXC1 *
                                       (-2558887/
                                        25109128109940022028573736079196160000000 +
                                        skoXC1 *
                                        (skoXC1 *
                                         (1853263/
                                          14462857791325452688458471981616988160000000 +
                                          skoXC1 *
                                          (skoXC1 *
                                           (-213019/
                                            1648765788211101606484265805904336650240000000 +
                                            skoXC1 *
                                            (skoXC1 *
                                             (1151/
                                              10991771921407344043228438706028911001600000000 +
                                              skoXC1 *
                                              (skoXC1 *
                                               (-31/
                                                461654420699108449815594425653214262067200000000 +
                                                skoXC1 *
                                                (skoXC1 *
                                                 1/
                                                 33239118290335808386722798647031426868838400000000)))))))))))))))))))))))))))))))))))))))))))) <=
     50000001/50000000),
    inference(decision, [],
              [refute_0_49, refute_0_48, refute_0_114, refute_0_141,
               refute_0_148, refute_0_172])).

cnf(refute_0_179, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 *
                (skoXC1 *
                 (1/25505877196800 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/21424936845312000 +
                    skoXC1 *
                    (skoXC1 *
                     (1/23310331287699456000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/31888533201572855808000 +
                        skoXC1 *
                        (skoXC1 *
                         1/
                         53572735778642397757440000)))))))))))))))))))) <=
     0), inference(resolve, [], [refute_0_178, refute_0_177])).

cnf(refute_0_180, plain, ($false),
    inference(resolve, [], [refute_0_179, refute_0_173])).
% SZS output end CNFRefutation for trigpoly-352-1-weak.tptp
