% SZS output start CNFRefutation for trigpoly-352-1.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120, Y) |
     lgen(R, sin(X), Y))).

cnf(sin_upper_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800, Y) |
     lgen(R, sin(X), Y))).

cnf(sin_lower_bound_1_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040) |
     lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_3_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800 + X ^ 13 / 6227020800 -
            X ^ 15 / 1307674368000) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_1, axiom,
    (~ lgen(R, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320,
            Y) | lgen(R, cos(X), Y))).

cnf(cos_upper_bound_3, axiom,
    (~ lgen(R,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200 +
            X ^ 16 / 20922789888000, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_1, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720) |
     lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_3, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800 + X ^ 12 / 479001600 - X ^ 14 / 87178291200) |
     lgen(R, Y, cos(X)))).

fof(trig_poly_352, conjecture,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        abs(1/2 + cos(X)) <= 10 ^ -2 + 1 / (2 * sin(X / 2))))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 62831/10000) =>
        abs(1/2 + cos(X)) <= 10 ^ -2 + 1 / (2 * sin(X / 2)))),
    inference(strip, [], [trig_poly_352])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 62831/10000) =>
          abs(1/2 + cos(X)) <= 10 ^ -2 + 1 / (2 * sin(X / 2)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 62831/10000 &
        10 ^ -2 + 1 / (2 * sin(X / 2)) < abs(1/2 + cos(X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 62831/10000 &
     10 ^ -2 + 1 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1)) &
     0 < skoXC1), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (10 ^ -2 + 1 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 < 62831/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 *
         (-49/192000 +
          skoXC1 *
          (skoXC1 *
           (7/4608000 +
            skoXC1 *
            (skoXC1 *
             (-7/1327104000 +
              skoXC1 *
              (skoXC1 *
               (7/583925760000 +
                skoXC1 * (skoXC1 * -7/364369674240000)))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800)))))))))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, cos(X_000032), X_000033) | cos(X_000032) <= X_000033),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033) |
     lgen(0, cos(X_000032), X_000033)),
    inference(subst, [], [cos_upper_bound_3])).

cnf(refute_0_3, plain,
    (~ lgen(0,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033) |
     cos(X_000032) <= X_000033),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000033 <
     1 +
     X_000032 *
     (X_000032 *
      (-1/2 +
       X_000032 *
       (X_000032 *
        (1/24 +
         X_000032 *
         (X_000032 *
          (-1/720 +
           X_000032 *
           (X_000032 *
            (1/40320 +
             X_000032 *
             (X_000032 *
              (-1/3628800 +
               X_000032 *
               (X_000032 *
                (1/479001600 +
                 X_000032 *
                 (X_000032 *
                  (-1/87178291200 +
                   X_000032 * (X_000032 * 1/20922789888000))))))))))))))) |
     cos(X_000032) <= X_000033), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     cos(skoXC1) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800)))))))))))))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) *
     (49/50 + cos(skoXC1) * 2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) <=
     1 / (49/50 + cos(skoXC1) * 2) | 49/50 + cos(skoXC1) * 2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    (~ lgen(0, sin(X_000111), X_000112) | sin(X_000111) <= X_000112),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_8, plain,
    (X_000111 < 0 |
     ~ lgen(0,
            X_000111 - X_000111 ^ 3 / 6 + X_000111 ^ 5 / 120 -
            X_000111 ^ 7 / 5040 + X_000111 ^ 9 / 362880 -
            X_000111 ^ 11 / 39916800 + X_000111 ^ 13 / 6227020800,
            X_000112) | lgen(0, sin(X_000111), X_000112)),
    inference(subst, [], [sin_upper_bound_3_pos])).

cnf(refute_0_9, plain,
    (X_000111 < 0 |
     ~ lgen(0,
            X_000111 - X_000111 ^ 3 / 6 + X_000111 ^ 5 / 120 -
            X_000111 ^ 7 / 5040 + X_000111 ^ 9 / 362880 -
            X_000111 ^ 11 / 39916800 + X_000111 ^ 13 / 6227020800,
            X_000112) | sin(X_000111) <= X_000112),
    inference(resolve, [], [refute_0_8, refute_0_7])).

cnf(refute_0_10, plain,
    (X_000111 < 0 |
     X_000112 <
     X_000111 *
     (1 +
      X_000111 *
      (X_000111 *
       (-1/6 +
        X_000111 *
        (X_000111 *
         (1/120 +
          X_000111 *
          (X_000111 *
           (-1/5040 +
            X_000111 *
            (X_000111 *
             (1/362880 +
              X_000111 *
              (X_000111 *
               (-1/39916800 +
                X_000111 * (X_000111 * 1/6227020800)))))))))))) |
     sin(X_000111) <= X_000112), inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/6227020800)))))))))))) |
     sin(skoXC1 * 1/2) <= 1 / (49/50 + cos(skoXC1) * 2)),
    inference(subst, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (1 < (49/100 + cos(skoXC1)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0 |
     49/100 + cos(skoXC1) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_13, plain,
    (10 ^ -2 + 1 / (2 * sin(skoXC1 / 2)) < abs(1/2 + cos(skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_14, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + abs(1/2 + cos(skoXC1))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (1/2 + cos(skoXC1) < 0 | abs(1/2 + cos(skoXC1)) = 1/2 + cos(skoXC1)),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_16, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + (1/2 + cos(skoXC1)) |
     abs(1/2 + cos(skoXC1)) != 1/2 + cos(skoXC1) |
     -1/100 + abs(1/2 + cos(skoXC1)) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    introduced(tautology, [equality])).

cnf(refute_0_17, plain,
    (1/2 + cos(skoXC1) < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + (1/2 + cos(skoXC1)) |
     -1/100 + abs(1/2 + cos(skoXC1)) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (1/2 + cos(skoXC1) < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + (1/2 + cos(skoXC1))),
    inference(resolve, [], [refute_0_17, refute_0_14])).

cnf(refute_0_19, plain,
    (cos(skoXC1) < -1/2 |
     1 / (sin(skoXC1 * 1/2) * 2) < 49/100 + cos(skoXC1)),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560)))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200)))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_21, plain,
    (~ lgen(0, X_000026, cos(X_000025)) | X_000026 <= cos(X_000025)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_22, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 - X_000025 ^ 6 / 720 +
            X_000025 ^ 8 / 40320 - X_000025 ^ 10 / 3628800 +
            X_000025 ^ 12 / 479001600 - X_000025 ^ 14 / 87178291200) |
     lgen(0, X_000026, cos(X_000025))),
    inference(subst, [], [cos_lower_bound_3])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000026,
            1 - X_000025 ^ 2 / 2 + X_000025 ^ 4 / 24 - X_000025 ^ 6 / 720 +
            X_000025 ^ 8 / 40320 - X_000025 ^ 10 / 3628800 +
            X_000025 ^ 12 / 479001600 - X_000025 ^ 14 / 87178291200) |
     X_000026 <= cos(X_000025)),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (1 +
     X_000025 *
     (X_000025 *
      (-1/2 +
       X_000025 *
       (X_000025 *
        (1/24 +
         X_000025 *
         (X_000025 *
          (-1/720 +
           X_000025 *
           (X_000025 *
            (1/40320 +
             X_000025 *
             (X_000025 *
              (-1/3628800 +
               X_000025 *
               (X_000025 *
                (1/479001600 +
                 X_000025 * (X_000025 * -1/87178291200))))))))))))) <
     X_000026 | X_000026 <= cos(X_000025)),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) *
     (-51/50 + cos(skoXC1) * -2) | 0 <= -51/50 + cos(skoXC1) * -2 |
     1 / (-51/50 + cos(skoXC1) * -2) <=
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_27, plain,
    (~ lgen(0, X_000103, sin(X_000102)) | X_000103 <= sin(X_000102)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_28, plain,
    (X_000102 < 0 |
     ~ lgen(0, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040) | lgen(0, X_000103, sin(X_000102))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_29, plain,
    (X_000102 < 0 |
     ~ lgen(0, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040) | X_000103 <= sin(X_000102)),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (X_000102 *
     (1 +
      X_000102 *
      (X_000102 *
       (-1/6 +
        X_000102 *
        (X_000102 * (1/120 + X_000102 * (X_000102 * -1/5040)))))) <
     X_000103 | X_000102 < 0 | X_000103 <= sin(X_000102)),
    inference(arithmetic, [], [refute_0_29])).

cnf(refute_0_31, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <
     1 / (-51/50 + cos(skoXC1) * -2) |
     1 / (-51/50 + cos(skoXC1) * -2) <= sin(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (0 <= 1/2 + cos(skoXC1) |
     abs(1/2 + cos(skoXC1)) = -(1/2 + cos(skoXC1))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_33, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + -(1/2 + cos(skoXC1)) |
     abs(1/2 + cos(skoXC1)) != -(1/2 + cos(skoXC1)) |
     -1/100 + abs(1/2 + cos(skoXC1)) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    introduced(tautology, [equality])).

cnf(refute_0_34, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + -(1/2 + cos(skoXC1)) |
     0 <= 1/2 + cos(skoXC1) |
     -1/100 + abs(1/2 + cos(skoXC1)) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -1/100 + -(1/2 + cos(skoXC1)) |
     0 <= 1/2 + cos(skoXC1)),
    inference(resolve, [], [refute_0_34, refute_0_14])).

cnf(refute_0_36, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < -51/100 + cos(skoXC1) * -1 |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_35])).

cnf(refute_0_37, plain,
    (1 < (-51/100 + cos(skoXC1) * -1) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0 |
     -51/100 + cos(skoXC1) * -1 <= 1 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_38, plain,
    (1 < (-51/100 + cos(skoXC1) * -1) * (sin(skoXC1 * 1/2) * 2) |
     -1/2 <= cos(skoXC1) | sin(skoXC1 * 1/2) * 2 <= 0),
    inference(resolve, [], [refute_0_37, refute_0_36])).

cnf(refute_0_39, plain,
    (1 < sin(skoXC1 * 1/2) * (-51/50 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1) | sin(skoXC1 * 1/2) <= 0),
    inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (X_000103 < sin(X_000102) | ~ lgen(1, X_000103, sin(X_000102))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_41, plain,
    (X_000102 < 0 |
     ~ lgen(1, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040) | lgen(1, X_000103, sin(X_000102))),
    inference(subst, [], [sin_lower_bound_1_pos])).

cnf(refute_0_42, plain,
    (X_000102 < 0 | X_000103 < sin(X_000102) |
     ~ lgen(1, X_000103,
            X_000102 - X_000102 ^ 3 / 6 + X_000102 ^ 5 / 120 -
            X_000102 ^ 7 / 5040)),
    inference(resolve, [], [refute_0_41, refute_0_40])).

cnf(refute_0_43, plain,
    (X_000102 < 0 | X_000103 < sin(X_000102) |
     X_000102 *
     (1 +
      X_000102 *
      (X_000102 *
       (-1/6 +
        X_000102 *
        (X_000102 * (1/120 + X_000102 * (X_000102 * -1/5040)))))) <=
     X_000103), inference(arithmetic, [], [refute_0_42])).

cnf(refute_0_44, plain,
    (skoXC1 * 1/2 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <= 0),
    inference(subst, [], [refute_0_43])).

cnf(refute_0_45, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0), inference(arithmetic, [], [refute_0_44])).

cnf(refute_0_46, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) < -1/2 |
     -1/2 <= cos(skoXC1)), inference(subst, [], [refute_0_24])).

cnf(refute_0_47, plain,
    (3/2 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (~ lgen(0, X_000012, cos(X_000011)) | X_000012 <= cos(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_49, plain,
    (~ lgen(0, X_000012,
            1 - X_000011 ^ 2 / 2 + X_000011 ^ 4 / 24 -
            X_000011 ^ 6 / 720) | lgen(0, X_000012, cos(X_000011))),
    inference(subst, [], [cos_lower_bound_1])).

cnf(refute_0_50, plain,
    (~ lgen(0, X_000012,
            1 - X_000011 ^ 2 / 2 + X_000011 ^ 4 / 24 -
            X_000011 ^ 6 / 720) | X_000012 <= cos(X_000011)),
    inference(resolve, [], [refute_0_49, refute_0_48])).

cnf(refute_0_51, plain,
    (1 +
     X_000011 *
     (X_000011 *
      (-1/2 +
       X_000011 * (X_000011 * (1/24 + X_000011 * (X_000011 * -1/720))))) <
     X_000012 | X_000012 <= cos(X_000011)),
    inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     -1/2 | -1/2 <= cos(skoXC1)), inference(subst, [], [refute_0_51])).

cnf(refute_0_53, plain,
    (3/2 <
     skoXC1 *
     (skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/720))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_52])).

cnf(refute_0_54, plain, (skoXC1 < 62831/10000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_55, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_56, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_47, refute_0_53, refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_45, refute_0_56])).

cnf(refute_0_58, plain, (-1/2 <= cos(skoXC1) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_47, refute_0_53, refute_0_54, refute_0_55])).

cnf(refute_0_59, plain, (0 < sin(skoXC1 * 1/2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_58, refute_0_57])).

cnf(refute_0_60, plain,
    (1 < sin(skoXC1 * 1/2) * (-51/50 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_39, refute_0_59])).

cnf(refute_0_61, plain,
    (sin(skoXC1 * 1/2) < 1 / (-51/50 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1) | 0 <= -51/50 + cos(skoXC1) * -2),
    inference(split, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (sin(skoXC1 * 1/2) < 1 / (-51/50 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1) | cos(skoXC1) <= -51/100),
    inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (1 < 0 | -51/50 + cos(skoXC1) * -2 < 0 |
     0 < -51/50 + cos(skoXC1) * -2 | -1/2 <= cos(skoXC1)),
    inference(split, [], [refute_0_60])).

cnf(refute_0_64, plain,
    (cos(skoXC1) < -51/100 | -51/100 < cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_63])).

cnf(refute_0_65, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920)))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200)))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_66, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_24])).

cnf(refute_0_67, plain,
    (1 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (-51/50 + cos(skoXC1) * -2) |
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <=
     1 / (-51/50 + cos(skoXC1) * -2) | -51/50 + cos(skoXC1) * -2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_68, plain,
    (~ lgen(0, sin(X_000081), X_000082) | sin(X_000081) <= X_000082),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_69, plain,
    (X_000081 < 0 |
     ~ lgen(0, X_000081 - X_000081 ^ 3 / 6 + X_000081 ^ 5 / 120,
            X_000082) | lgen(0, sin(X_000081), X_000082)),
    inference(subst, [], [sin_upper_bound_1_pos])).

cnf(refute_0_70, plain,
    (X_000081 < 0 |
     ~ lgen(0, X_000081 - X_000081 ^ 3 / 6 + X_000081 ^ 5 / 120,
            X_000082) | sin(X_000081) <= X_000082),
    inference(resolve, [], [refute_0_69, refute_0_68])).

cnf(refute_0_71, plain,
    (X_000081 < 0 |
     X_000082 <
     X_000081 *
     (1 + X_000081 * (X_000081 * (-1/6 + X_000081 * (X_000081 * 1/120)))) |
     sin(X_000081) <= X_000082), inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (-51/50 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     sin(skoXC1 * 1/2) <= 1 / (-51/50 + cos(skoXC1) * -2)),
    inference(subst, [], [refute_0_71])).

cnf(refute_0_73, plain,
    (1 / (-51/50 + cos(skoXC1) * -2) < sin(skoXC1 * 1/2) |
     -1/2 <= cos(skoXC1) | -51/50 + cos(skoXC1) * -2 <= 0),
    inference(split, [], [refute_0_60])).

cnf(refute_0_74, plain,
    (1 / (-51/50 + cos(skoXC1) * -2) < sin(skoXC1 * 1/2) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (-51/50 + cos(skoXC1) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_72, refute_0_74])).

cnf(refute_0_76, plain,
    (skoXC1 < 0 |
     1 / (-51/50 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_75])).

cnf(refute_0_77, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     -51/100 | -51/100 <= cos(skoXC1)),
    inference(subst, [], [refute_0_51])).

cnf(refute_0_78, plain,
    (151/100 <
     skoXC1 *
     (skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/720))))) |
     -51/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) < -51/100 |
     -51/100 <= cos(skoXC1)), inference(subst, [], [refute_0_24])).

cnf(refute_0_80, plain,
    (151/100 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     -51/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_79])).

cnf(refute_0_81, plain,
    (-51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_47, refute_0_53, refute_0_54, refute_0_55,
               refute_0_78, refute_0_80])).

cnf(refute_0_82, plain,
    (1 / (-51/50 + cos(skoXC1) * -2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_81, refute_0_76])).

cnf(refute_0_83, plain,
    (1 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (-51/50 + cos(skoXC1) * -2) | -51/100 <= cos(skoXC1) |
     -1/2 <= cos(skoXC1) | -51/50 + cos(skoXC1) * -2 <= 0),
    inference(resolve, [], [refute_0_67, refute_0_82])).

cnf(refute_0_84, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000)))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))),
    inference(split, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_85])).

cnf(refute_0_87, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_80, refute_0_78, refute_0_55, refute_0_54,
               refute_0_53, refute_0_47])).

cnf(refute_0_88, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_86, refute_0_87])).

cnf(refute_0_89, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_66, refute_0_88])).

cnf(refute_0_90, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))),
    inference(resolve, [], [refute_0_65, refute_0_89])).

cnf(refute_0_91, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-101/967680 +
              skoXC1 *
              (skoXC1 *
               (59/29030400 +
                skoXC1 *
                (skoXC1 *
                 (-29/1094860800 +
                  skoXC1 *
                  (skoXC1 *
                   (5063/20922789888000 +
                    skoXC1 *
                    (skoXC1 *
                     (-131/83691159552000 +
                      skoXC1 *
                      (skoXC1 * 1/167382319104000)))))))))))))))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_90])).

cnf(refute_0_92, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920)))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_93, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_51])).

cnf(refute_0_94, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000))))) /
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_93, refute_0_88])).

cnf(refute_0_95, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 * (skoXC1 * (-17/800 + skoXC1 * (skoXC1 * 17/64000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/1920))))),
    inference(resolve, [], [refute_0_92, refute_0_94])).

cnf(refute_0_96, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_95])).

cnf(refute_0_97, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_80, refute_0_78, refute_0_55, refute_0_54,
               refute_0_53, refute_0_47])).

cnf(refute_0_98, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-11/138240 + skoXC1 * (skoXC1 * 1/1382400)))))))))) |
     -51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_96, refute_0_97])).

cnf(refute_0_99, plain,
    (-51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (13/3840 +
            skoXC1 *
            (skoXC1 *
             (-101/967680 +
              skoXC1 *
              (skoXC1 *
               (59/29030400 +
                skoXC1 *
                (skoXC1 *
                 (-29/1094860800 +
                  skoXC1 *
                  (skoXC1 *
                   (5063/20922789888000 +
                    skoXC1 *
                    (skoXC1 *
                     (-131/83691159552000 +
                      skoXC1 *
                      (skoXC1 * 1/167382319104000)))))))))))))))))) <= 1 |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(decision, [],
              [refute_0_98, refute_0_80, refute_0_78, refute_0_55,
               refute_0_54, refute_0_53, refute_0_47])).

cnf(refute_0_100, plain,
    (-51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(resolve, [], [refute_0_99, refute_0_91])).

cnf(refute_0_101, plain, (-51/100 <= cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_100, refute_0_87])).

cnf(refute_0_102, plain, (-51/100 < cos(skoXC1) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_101, refute_0_64])).

cnf(refute_0_103, plain,
    (sin(skoXC1 * 1/2) < 1 / (-51/50 + cos(skoXC1) * -2) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_62, refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <
     1 / (-51/50 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_31, refute_0_103])).

cnf(refute_0_105, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <
     1 / (-51/50 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_104])).

cnf(refute_0_106, plain,
    (skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <
     1 / (-51/50 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_58, refute_0_105])).

cnf(refute_0_107, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) *
     (-51/50 + cos(skoXC1) * -2) | -1/2 <= cos(skoXC1) |
     0 <= -51/50 + cos(skoXC1) * -2),
    inference(resolve, [], [refute_0_26, refute_0_106])).

cnf(refute_0_108, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1) | cos(skoXC1) <= -51/100),
    inference(arithmetic, [], [refute_0_107])).

cnf(refute_0_109, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_108, refute_0_102])).

cnf(refute_0_110, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))),
    inference(split, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(arithmetic, [], [refute_0_110])).

cnf(refute_0_112, plain,
    (-51/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     cos(skoXC1) <= -51/100), inference(subst, [], [refute_0_4])).

cnf(refute_0_113, plain,
    (-51/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_112, refute_0_102])).

cnf(refute_0_114, plain,
    (-151/100 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_113])).

cnf(refute_0_115, plain,
    (~ lgen(0, cos(X_000018), X_000019) | cos(X_000018) <= X_000019),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_116, plain,
    (~ lgen(0,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) |
     lgen(0, cos(X_000018), X_000019)),
    inference(subst, [], [cos_upper_bound_1])).

cnf(refute_0_117, plain,
    (~ lgen(0,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) | cos(X_000018) <= X_000019),
    inference(resolve, [], [refute_0_116, refute_0_115])).

cnf(refute_0_118, plain,
    (X_000019 <
     1 +
     X_000018 *
     (X_000018 *
      (-1/2 +
       X_000018 *
       (X_000018 *
        (1/24 +
         X_000018 *
         (X_000018 * (-1/720 + X_000018 * (X_000018 * 1/40320))))))) |
     cos(X_000018) <= X_000019),
    inference(arithmetic, [], [refute_0_117])).

cnf(refute_0_119, plain,
    (-51/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <= -51/100), inference(subst, [], [refute_0_118])).

cnf(refute_0_120, plain,
    (-51/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_119, refute_0_102])).

cnf(refute_0_121, plain,
    (-151/100 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     -1/2 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_120])).

cnf(refute_0_122, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_114, refute_0_121, refute_0_47, refute_0_53,
               refute_0_54, refute_0_55])).

cnf(refute_0_123, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_111, refute_0_122])).

cnf(refute_0_124, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_25, refute_0_123])).

cnf(refute_0_125, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 * (skoXC1 * -1/87178291200)))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))),
    inference(resolve, [], [refute_0_20, refute_0_124])).

cnf(refute_0_126, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (109351/32256000 +
            skoXC1 *
            (skoXC1 *
             (-41/387072 +
              skoXC1 *
              (skoXC1 *
               (251/116121600 +
                skoXC1 *
                (skoXC1 *
                 (-59/1916006400 +
                  skoXC1 *
                  (skoXC1 *
                   (26687/83691159552000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/413289676800 +
                      skoXC1 *
                      (skoXC1 *
                       (1/80343513169920 +
                        skoXC1 *
                        (skoXC1 *
                         -1/28120229609472000)))))))))))))))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(arithmetic, [], [refute_0_125])).

cnf(refute_0_127, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560)))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_128, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_51])).

cnf(refute_0_129, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     (1 +
      skoXC1 *
      (51/100 +
       skoXC1 *
       (skoXC1 *
        (-17/800 +
         skoXC1 *
         (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_128, refute_0_123])).

cnf(refute_0_130, plain,
    (1 +
     skoXC1 *
     (51/100 +
      skoXC1 *
      (skoXC1 *
       (-17/800 +
        skoXC1 *
        (skoXC1 * (17/64000 + skoXC1 * (skoXC1 * -17/10752000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720)))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))) |
     -1/2 <= cos(skoXC1) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 * (skoXC1 * (-1/1920 + skoXC1 * (skoXC1 * 1/322560))))))),
    inference(resolve, [], [refute_0_127, refute_0_129])).

cnf(refute_0_131, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (109351/32256000 +
            skoXC1 *
            (skoXC1 *
             (-157/1935360 +
              skoXC1 *
              (skoXC1 *
               (11/12902400 + skoXC1 * (skoXC1 * -1/232243200)))))))))))) |
     -1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(arithmetic, [], [refute_0_130])).

cnf(refute_0_132, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (109351/32256000 +
            skoXC1 *
            (skoXC1 *
             (-157/1935360 +
              skoXC1 *
              (skoXC1 *
               (11/12902400 + skoXC1 * (skoXC1 * -1/232243200)))))))))))) |
     0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) |
     -1/2 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_114, refute_0_121, refute_0_47, refute_0_53,
               refute_0_54, refute_0_55])).

cnf(refute_0_133, plain,
    (1 <
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (109351/32256000 +
            skoXC1 *
            (skoXC1 *
             (-157/1935360 +
              skoXC1 *
              (skoXC1 *
               (11/12902400 + skoXC1 * (skoXC1 * -1/232243200)))))))))))) |
     -1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_131, refute_0_132])).

cnf(refute_0_134, plain,
    (-1/2 <= cos(skoXC1) |
     skoXC1 *
     (-151/100 +
      skoXC1 *
      (skoXC1 *
       (1351/2400 +
        skoXC1 *
        (skoXC1 *
         (-12151/192000 +
          skoXC1 *
          (skoXC1 *
           (109351/32256000 +
            skoXC1 *
            (skoXC1 *
             (-41/387072 +
              skoXC1 *
              (skoXC1 *
               (251/116121600 +
                skoXC1 *
                (skoXC1 *
                 (-59/1916006400 +
                  skoXC1 *
                  (skoXC1 *
                   (26687/83691159552000 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/413289676800 +
                      skoXC1 *
                      (skoXC1 *
                       (1/80343513169920 +
                        skoXC1 *
                        (skoXC1 *
                         -1/28120229609472000)))))))))))))))))))) <= 1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0),
    inference(decision, [],
              [refute_0_133, refute_0_114, refute_0_121, refute_0_47,
               refute_0_53, refute_0_54, refute_0_55])).

cnf(refute_0_135, plain,
    (-1/2 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(resolve, [], [refute_0_134, refute_0_126])).

cnf(refute_0_136, plain, (-1/2 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_135, refute_0_122])).

cnf(refute_0_137, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < 49/100 + cos(skoXC1)),
    inference(resolve, [], [refute_0_136, refute_0_19])).

cnf(refute_0_138, plain,
    (1 < (49/100 + cos(skoXC1)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0),
    inference(resolve, [], [refute_0_12, refute_0_137])).

cnf(refute_0_139, plain,
    (1 < sin(skoXC1 * 1/2) * (49/50 + cos(skoXC1) * 2) |
     sin(skoXC1 * 1/2) <= 0), inference(arithmetic, [], [refute_0_138])).

cnf(refute_0_140, plain,
    (X_000133 < sin(X_000132) | ~ lgen(1, X_000133, sin(X_000132))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_141, plain,
    (X_000132 < 0 |
     ~ lgen(1, X_000133,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880 -
            X_000132 ^ 11 / 39916800 + X_000132 ^ 13 / 6227020800 -
            X_000132 ^ 15 / 1307674368000) |
     lgen(1, X_000133, sin(X_000132))),
    inference(subst, [], [sin_lower_bound_3_pos])).

cnf(refute_0_142, plain,
    (X_000132 < 0 | X_000133 < sin(X_000132) |
     ~ lgen(1, X_000133,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880 -
            X_000132 ^ 11 / 39916800 + X_000132 ^ 13 / 6227020800 -
            X_000132 ^ 15 / 1307674368000)),
    inference(resolve, [], [refute_0_141, refute_0_140])).

cnf(refute_0_143, plain,
    (X_000132 < 0 | X_000133 < sin(X_000132) |
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 *
           (-1/5040 +
            X_000132 *
            (X_000132 *
             (1/362880 +
              X_000132 *
              (X_000132 *
               (-1/39916800 +
                X_000132 *
                (X_000132 *
                 (1/6227020800 +
                  X_000132 * (X_000132 * -1/1307674368000)))))))))))))) <=
     X_000133), inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (skoXC1 * 1/2 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 *
                (skoXC1 * 1/2 *
                 (1/6227020800 +
                  skoXC1 * 1/2 *
                  (skoXC1 * 1/2 * -1/1307674368000)))))))))))))) <= 0),
    inference(subst, [], [refute_0_143])).

cnf(refute_0_145, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 * (skoXC1 * -1/42849873690624000)))))))))))))) <=
     0), inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain,
    (cos(X_000018) < X_000019 | ~ lgen(1, cos(X_000018), X_000019)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_147, plain,
    (~ lgen(1,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019) |
     lgen(1, cos(X_000018), X_000019)),
    inference(subst, [], [cos_upper_bound_1])).

cnf(refute_0_148, plain,
    (cos(X_000018) < X_000019 |
     ~ lgen(1,
            1 - X_000018 ^ 2 / 2 + X_000018 ^ 4 / 24 - X_000018 ^ 6 / 720 +
            X_000018 ^ 8 / 40320, X_000019)),
    inference(resolve, [], [refute_0_147, refute_0_146])).

cnf(refute_0_149, plain,
    (cos(X_000018) < X_000019 |
     X_000019 <=
     1 +
     X_000018 *
     (X_000018 *
      (-1/2 +
       X_000018 *
       (X_000018 *
        (1/24 +
         X_000018 *
         (X_000018 * (-1/720 + X_000018 * (X_000018 * 1/40320)))))))),
    inference(arithmetic, [], [refute_0_148])).

cnf(refute_0_150, plain,
    (cos(skoXC1) < -1/2 |
     -1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(subst, [], [refute_0_149])).

cnf(refute_0_151, plain,
    (-1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_136, refute_0_150])).

cnf(refute_0_152, plain,
    (skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/720 + skoXC1 * (skoXC1 * -1/40320))))))) <=
     3/2), inference(arithmetic, [], [refute_0_151])).

cnf(refute_0_153, plain,
    (0 <= skoXC1 |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0),
    inference(decision, [], [refute_0_152, refute_0_55, refute_0_54])).

cnf(refute_0_154, plain,
    (0 < sin(skoXC1 * 1/2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <=
     0), inference(resolve, [], [refute_0_153, refute_0_45])).

cnf(refute_0_155, plain,
    (skoXC1 < 0 |
     0 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 *
                (skoXC1 *
                 (1/51011754393600 +
                  skoXC1 * (skoXC1 * -1/42849873690624000)))))))))))))) |
     0 < sin(skoXC1 * 1/2)),
    inference(decision, [],
              [refute_0_154, refute_0_152, refute_0_55, refute_0_54])).

cnf(refute_0_156, plain, (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_145, refute_0_155])).

cnf(refute_0_157, plain, (0 < sin(skoXC1 * 1/2) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_154, refute_0_152, refute_0_55, refute_0_54])).

cnf(refute_0_158, plain, (0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_157, refute_0_156])).

cnf(refute_0_159, plain,
    (1 < sin(skoXC1 * 1/2) * (49/50 + cos(skoXC1) * 2)),
    inference(resolve, [], [refute_0_139, refute_0_158])).

cnf(refute_0_160, plain,
    (1 / (49/50 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2) |
     49/50 + cos(skoXC1) * 2 <= 0), inference(split, [], [refute_0_159])).

cnf(refute_0_161, plain,
    (1 / (49/50 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2) |
     cos(skoXC1) <= -49/100), inference(arithmetic, [], [refute_0_160])).

cnf(refute_0_162, plain,
    (1 < 0 | 49/50 + cos(skoXC1) * 2 < 0 | 0 < 49/50 + cos(skoXC1) * 2),
    inference(split, [], [refute_0_159])).

cnf(refute_0_163, plain, (cos(skoXC1) < -49/100 | -49/100 < cos(skoXC1)),
    inference(arithmetic, [], [refute_0_162])).

cnf(refute_0_164, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_165, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))))),
    inference(subst, [], [refute_0_118])).

cnf(refute_0_166, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) *
     (49/50 + cos(skoXC1) * 2) | 0 <= 49/50 + cos(skoXC1) * 2 |
     1 / (49/50 + cos(skoXC1) * 2) <=
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_167, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <
     1 / (49/50 + cos(skoXC1) * 2) |
     1 / (49/50 + cos(skoXC1) * 2) <= sin(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_30])).

cnf(refute_0_168, plain,
    (sin(skoXC1 * 1/2) < 1 / (49/50 + cos(skoXC1) * 2) |
     0 <= 49/50 + cos(skoXC1) * 2), inference(split, [], [refute_0_159])).

cnf(refute_0_169, plain,
    (sin(skoXC1 * 1/2) < 1 / (49/50 + cos(skoXC1) * 2) |
     -49/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_168])).

cnf(refute_0_170, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/5040)))))) <
     1 / (49/50 + cos(skoXC1) * 2) | -49/100 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_167, refute_0_169])).

cnf(refute_0_171, plain,
    (skoXC1 < 0 |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <
     1 / (49/50 + cos(skoXC1) * 2) | -49/100 <= cos(skoXC1)),
    inference(arithmetic, [], [refute_0_170])).

cnf(refute_0_172, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 + skoXC1 * (skoXC1 * (1/24 + skoXC1 * (skoXC1 * -1/720))))) <
     -49/100 | -49/100 <= cos(skoXC1)),
    inference(subst, [], [refute_0_51])).

cnf(refute_0_173, plain,
    (149/100 <
     skoXC1 *
     (skoXC1 *
      (1/2 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/720))))) |
     -49/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_172])).

cnf(refute_0_174, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 * (skoXC1 * -1/87178291200))))))))))))) < -49/100 |
     -49/100 <= cos(skoXC1)), inference(subst, [], [refute_0_24])).

cnf(refute_0_175, plain,
    (149/100 <
     skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 * (skoXC1 * 1/87178291200))))))))))))) |
     -49/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_174])).

cnf(refute_0_176, plain, (-49/100 <= cos(skoXC1) | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_152, refute_0_55, refute_0_54, refute_0_173,
               refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) <
     1 / (49/50 + cos(skoXC1) * 2) | -49/100 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_176, refute_0_171])).

cnf(refute_0_178, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 * (skoXC1 * (1/3840 + skoXC1 * (skoXC1 * -1/645120)))))) *
     (49/50 + cos(skoXC1) * 2) | -49/100 <= cos(skoXC1) |
     0 <= 49/50 + cos(skoXC1) * 2),
    inference(resolve, [], [refute_0_166, refute_0_177])).

cnf(refute_0_179, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000)))))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) |
     -49/100 <= cos(skoXC1)), inference(arithmetic, [], [refute_0_178])).

cnf(refute_0_180, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) <
     cos(skoXC1) | -49/100 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(split, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (cos(X_000032) < X_000033 | ~ lgen(1, cos(X_000032), X_000033)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_182, plain,
    (~ lgen(1,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033) |
     lgen(1, cos(X_000032), X_000033)),
    inference(subst, [], [cos_upper_bound_3])).

cnf(refute_0_183, plain,
    (cos(X_000032) < X_000033 |
     ~ lgen(1,
            1 - X_000032 ^ 2 / 2 + X_000032 ^ 4 / 24 - X_000032 ^ 6 / 720 +
            X_000032 ^ 8 / 40320 - X_000032 ^ 10 / 3628800 +
            X_000032 ^ 12 / 479001600 - X_000032 ^ 14 / 87178291200 +
            X_000032 ^ 16 / 20922789888000, X_000033)),
    inference(resolve, [], [refute_0_182, refute_0_181])).

cnf(refute_0_184, plain,
    (cos(X_000032) < X_000033 |
     X_000033 <=
     1 +
     X_000032 *
     (X_000032 *
      (-1/2 +
       X_000032 *
       (X_000032 *
        (1/24 +
         X_000032 *
         (X_000032 *
          (-1/720 +
           X_000032 *
           (X_000032 *
            (1/40320 +
             X_000032 *
             (X_000032 *
              (-1/3628800 +
               X_000032 *
               (X_000032 *
                (1/479001600 +
                 X_000032 *
                 (X_000032 *
                  (-1/87178291200 +
                   X_000032 * (X_000032 * 1/20922789888000)))))))))))))))),
    inference(arithmetic, [], [refute_0_183])).

cnf(refute_0_185, plain,
    (cos(skoXC1) < -1/2 |
     -1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(subst, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (-1/2 <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(resolve, [], [refute_0_136, refute_0_185])).

cnf(refute_0_187, plain,
    (skoXC1 *
     (skoXC1 *
      (1/2 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/720 +
           skoXC1 *
           (skoXC1 *
            (-1/40320 +
             skoXC1 *
             (skoXC1 *
              (1/3628800 +
               skoXC1 *
               (skoXC1 *
                (-1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (1/87178291200 +
                   skoXC1 * (skoXC1 * -1/20922789888000))))))))))))))) <=
     3/2), inference(arithmetic, [], [refute_0_186])).

cnf(refute_0_188, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) |
     -49/100 <= cos(skoXC1)),
    inference(decision, [],
              [refute_0_175, refute_0_173, refute_0_54, refute_0_55,
               refute_0_152, refute_0_187])).

cnf(refute_0_189, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) <
     cos(skoXC1) | -49/100 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_180, refute_0_188])).

cnf(refute_0_190, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     -49/100 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_165, refute_0_189])).

cnf(refute_0_191, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 * (-49/192000 + skoXC1 * (skoXC1 * 7/4608000)))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560))))))) |
     -49/100 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(resolve, [], [refute_0_164, refute_0_190])).

cnf(refute_0_192, plain,
    (1 <
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-109349/32256000 +
            skoXC1 *
            (skoXC1 *
             (41/387072 +
              skoXC1 *
              (skoXC1 *
               (-73/38707200 +
                skoXC1 *
                (skoXC1 *
                 (1/58060800 +
                  skoXC1 * (skoXC1 * -1/13005619200)))))))))))))) |
     -49/100 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(arithmetic, [], [refute_0_191])).

cnf(refute_0_193, plain,
    (-49/100 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0 |
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-109349/32256000 +
            skoXC1 *
            (skoXC1 *
             (41/387072 +
              skoXC1 *
              (skoXC1 *
               (-73/38707200 +
                skoXC1 *
                (skoXC1 *
                 (1/58060800 +
                  skoXC1 * (skoXC1 * -1/13005619200)))))))))))))) <= 1),
    inference(decision, [],
              [refute_0_175, refute_0_173, refute_0_54, refute_0_55,
               refute_0_152, refute_0_187])).

cnf(refute_0_194, plain,
    (-49/100 <= cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 * (skoXC1 * (1/1920 + skoXC1 * (skoXC1 * -1/322560)))))) <=
     0), inference(resolve, [], [refute_0_193, refute_0_192])).

cnf(refute_0_195, plain, (-49/100 <= cos(skoXC1)),
    inference(resolve, [], [refute_0_194, refute_0_188])).

cnf(refute_0_196, plain, (-49/100 < cos(skoXC1)),
    inference(resolve, [], [refute_0_195, refute_0_163])).

cnf(refute_0_197, plain,
    (1 / (49/50 + cos(skoXC1) * 2) < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_161, refute_0_196])).

cnf(refute_0_198, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 +
            skoXC1 * 1/2 *
            (skoXC1 * 1/2 *
             (1/362880 +
              skoXC1 * 1/2 *
              (skoXC1 * 1/2 *
               (-1/39916800 +
                skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/6227020800))))))))))))),
    inference(resolve, [], [refute_0_11, refute_0_197])).

cnf(refute_0_199, plain,
    (skoXC1 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600))))))))))))),
    inference(arithmetic, [], [refute_0_198])).

cnf(refute_0_200, plain,
    (-49/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <= -49/100), inference(subst, [], [refute_0_118])).

cnf(refute_0_201, plain,
    (-49/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_200, refute_0_196])).

cnf(refute_0_202, plain,
    (-149/100 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(arithmetic, [], [refute_0_201])).

cnf(refute_0_203, plain,
    (-49/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000))))))))))))))) |
     cos(skoXC1) <= -49/100), inference(subst, [], [refute_0_4])).

cnf(refute_0_204, plain,
    (-49/100 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(resolve, [], [refute_0_203, refute_0_196])).

cnf(refute_0_205, plain,
    (-149/100 <
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain, (0 <= skoXC1),
    inference(decision, [],
              [refute_0_54, refute_0_55, refute_0_152, refute_0_187,
               refute_0_202, refute_0_205])).

cnf(refute_0_207, plain,
    (1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600))))))))))))),
    inference(resolve, [], [refute_0_206, refute_0_199])).

cnf(refute_0_208, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 *
           (-1/645120 +
            skoXC1 *
            (skoXC1 *
             (1/185794560 +
              skoXC1 *
              (skoXC1 *
               (-1/81749606400 +
                skoXC1 * (skoXC1 * 1/51011754393600)))))))))))) *
     (49/50 + cos(skoXC1) * 2) | 49/50 + cos(skoXC1) * 2 <= 0),
    inference(resolve, [], [refute_0_6, refute_0_207])).

cnf(refute_0_209, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 *
         (-49/192000 +
          skoXC1 *
          (skoXC1 *
           (7/4608000 +
            skoXC1 *
            (skoXC1 *
             (-7/1327104000 +
              skoXC1 *
              (skoXC1 *
               (7/583925760000 +
                skoXC1 * (skoXC1 * -7/364369674240000)))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) |
     cos(skoXC1) <= -49/100), inference(arithmetic, [], [refute_0_208])).

cnf(refute_0_210, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 *
         (-49/192000 +
          skoXC1 *
          (skoXC1 *
           (7/4608000 +
            skoXC1 *
            (skoXC1 *
             (-7/1327104000 +
              skoXC1 *
              (skoXC1 *
               (7/583925760000 +
                skoXC1 * (skoXC1 * -7/364369674240000)))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800)))))))))))))),
    inference(resolve, [], [refute_0_209, refute_0_196])).

cnf(refute_0_211, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) <
     cos(skoXC1) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(split, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_213, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320))))))) |
     cos(skoXC1) <=
     (1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(subst, [], [refute_0_118])).

cnf(refute_0_214, plain,
    (1 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (49/50 + cos(skoXC1) * 2) |
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) <=
     1 / (49/50 + cos(skoXC1) * 2) | 49/50 + cos(skoXC1) * 2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_215, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120)))) |
     sin(skoXC1 * 1/2) <= 1 / (49/50 + cos(skoXC1) * 2)),
    inference(subst, [], [refute_0_71])).

cnf(refute_0_216, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 * (-1/6 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/120))))),
    inference(resolve, [], [refute_0_215, refute_0_197])).

cnf(refute_0_217, plain,
    (skoXC1 < 0 |
     1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))),
    inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (1 / (49/50 + cos(skoXC1) * 2) <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840))))),
    inference(resolve, [], [refute_0_206, refute_0_217])).

cnf(refute_0_219, plain,
    (1 <
     skoXC1 *
     (1/2 + skoXC1 * (skoXC1 * (-1/48 + skoXC1 * (skoXC1 * 1/3840)))) *
     (49/50 + cos(skoXC1) * 2) | 49/50 + cos(skoXC1) * 2 <= 0),
    inference(resolve, [], [refute_0_214, refute_0_218])).

cnf(refute_0_220, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000)))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     cos(skoXC1) <= -49/100), inference(arithmetic, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000)))) <
     cos(skoXC1) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))))),
    inference(resolve, [], [refute_0_220, refute_0_196])).

cnf(refute_0_222, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     cos(skoXC1) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(split, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))),
    inference(decision, [],
              [refute_0_205, refute_0_202, refute_0_187, refute_0_152,
               refute_0_55, refute_0_54])).

cnf(refute_0_224, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     cos(skoXC1)), inference(resolve, [], [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000))))) /
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))),
    inference(resolve, [], [refute_0_213, refute_0_224])).

cnf(refute_0_226, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 * (skoXC1 * (49/2400 + skoXC1 * (skoXC1 * -49/192000)))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 * (skoXC1 * (-1/720 + skoXC1 * (skoXC1 * 1/40320)))))))) *
     (skoXC1 *
      (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(resolve, [], [refute_0_212, refute_0_225])).

cnf(refute_0_227, plain,
    (1 <
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400)))))))))))) |
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) <= 0),
    inference(arithmetic, [], [refute_0_226])).

cnf(refute_0_228, plain,
    (0 <
     skoXC1 *
     (1 + skoXC1 * (skoXC1 * (-1/24 + skoXC1 * (skoXC1 * 1/1920)))) |
     1 <
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400))))))))))))),
    inference(decision, [],
              [refute_0_205, refute_0_202, refute_0_187, refute_0_152,
               refute_0_55, refute_0_54])).

cnf(refute_0_229, plain,
    (1 <
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-13/3840 +
            skoXC1 *
            (skoXC1 *
             (101/967680 +
              skoXC1 *
              (skoXC1 *
               (-17/9676800 + skoXC1 * (skoXC1 * 1/77414400))))))))))))),
    inference(resolve, [], [refute_0_227, refute_0_228])).

cnf(refute_0_230, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800))))))))))))),
    inference(decision, [],
              [refute_0_229, refute_0_205, refute_0_202, refute_0_187,
               refute_0_152, refute_0_55, refute_0_54])).

cnf(refute_0_231, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) <
     cos(skoXC1)), inference(resolve, [], [refute_0_211, refute_0_230])).

cnf(refute_0_232, plain,
    ((1 +
      skoXC1 *
      (-49/100 +
       skoXC1 *
       (skoXC1 *
        (49/2400 +
         skoXC1 *
         (skoXC1 *
          (-49/192000 +
           skoXC1 *
           (skoXC1 *
            (7/4608000 +
             skoXC1 *
             (skoXC1 *
              (-7/1327104000 +
               skoXC1 *
               (skoXC1 *
                (7/583925760000 +
                 skoXC1 * (skoXC1 * -7/364369674240000))))))))))))) /
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))),
    inference(resolve, [], [refute_0_5, refute_0_231])).

cnf(refute_0_233, plain,
    (1 +
     skoXC1 *
     (-49/100 +
      skoXC1 *
      (skoXC1 *
       (49/2400 +
        skoXC1 *
        (skoXC1 *
         (-49/192000 +
          skoXC1 *
          (skoXC1 *
           (7/4608000 +
            skoXC1 *
            (skoXC1 *
             (-7/1327104000 +
              skoXC1 *
              (skoXC1 *
               (7/583925760000 +
                skoXC1 * (skoXC1 * -7/364369674240000)))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 * (skoXC1 * 1/20922789888000)))))))))))))))) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 *
            (-1/322560 +
             skoXC1 *
             (skoXC1 *
              (1/92897280 +
               skoXC1 *
               (skoXC1 *
                (-1/40874803200 +
                 skoXC1 * (skoXC1 * 1/25505877196800))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_232])).

cnf(refute_0_234, plain,
    (1 <
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-109349/32256000 +
            skoXC1 *
            (skoXC1 *
             (984149/9289728000 +
              skoXC1 *
              (skoXC1 *
               (-8857349/4087480320000 +
                skoXC1 *
                (skoXC1 *
                 (79716149/2550587719680000 +
                  skoXC1 *
                  (skoXC1 *
                   (-597871/1785411403776000 +
                    skoXC1 *
                    (skoXC1 *
                     (59347/21424936845312000 +
                      skoXC1 *
                      (skoXC1 *
                       (-6967/385648863215616000 +
                        skoXC1 *
                        (skoXC1 *
                         (29363/323945045101117440000 +
                          skoXC1 *
                          (skoXC1 *
                           (-2377/7126790992224583680000 +
                            skoXC1 *
                            (skoXC1 *
                             (1/1140286558755933388800 +
                              skoXC1 *
                              (skoXC1 *
                               (-1/617655219326130585600000 +
                                skoXC1 *
                                (skoXC1 *
                                 1/
                                 533654109497776825958400000)))))))))))))))))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(arithmetic, [], [refute_0_233])).

cnf(refute_0_235, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0 |
     skoXC1 *
     (149/100 +
      skoXC1 *
      (skoXC1 *
       (-1349/2400 +
        skoXC1 *
        (skoXC1 *
         (12149/192000 +
          skoXC1 *
          (skoXC1 *
           (-109349/32256000 +
            skoXC1 *
            (skoXC1 *
             (984149/9289728000 +
              skoXC1 *
              (skoXC1 *
               (-8857349/4087480320000 +
                skoXC1 *
                (skoXC1 *
                 (79716149/2550587719680000 +
                  skoXC1 *
                  (skoXC1 *
                   (-597871/1785411403776000 +
                    skoXC1 *
                    (skoXC1 *
                     (59347/21424936845312000 +
                      skoXC1 *
                      (skoXC1 *
                       (-6967/385648863215616000 +
                        skoXC1 *
                        (skoXC1 *
                         (29363/323945045101117440000 +
                          skoXC1 *
                          (skoXC1 *
                           (-2377/7126790992224583680000 +
                            skoXC1 *
                            (skoXC1 *
                             (1/1140286558755933388800 +
                              skoXC1 *
                              (skoXC1 *
                               (-1/617655219326130585600000 +
                                skoXC1 *
                                (skoXC1 *
                                 1/
                                 533654109497776825958400000)))))))))))))))))))))))))))) <=
     1),
    inference(decision, [],
              [refute_0_229, refute_0_205, refute_0_202, refute_0_187,
               refute_0_152, refute_0_55, refute_0_54])).

cnf(refute_0_236, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 *
           (-1/322560 +
            skoXC1 *
            (skoXC1 *
             (1/92897280 +
              skoXC1 *
              (skoXC1 *
               (-1/40874803200 +
                skoXC1 * (skoXC1 * 1/25505877196800)))))))))))) <= 0),
    inference(resolve, [], [refute_0_235, refute_0_234])).

cnf(refute_0_237, plain, ($false),
    inference(resolve, [], [refute_0_236, refute_0_230])).
% SZS output end CNFRefutation for trigpoly-352-1.tptp
