% SZS output start CNFRefutation for trigpoly-352-2-weak.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880,
            Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(cos_upper_bound_7, axiom,
    (~ lgen(R,
            1 - 1/2 * X ^ 2 + 1/24 * X ^ 4 - 1/720 * X ^ 6 +
            1/40320 * X ^ 8 - 1/3628800 * X ^ 10 + 1/479001600 * X ^ 12 -
            1/87178291200 * X ^ 14 + 1/20922789888000 * X ^ 16 -
            1/6402373705728000 * X ^ 18 + 1/2432902008176640000 * X ^ 20 -
            1/1124000727777607680000 * X ^ 22 +
            1/620448401733239439360000 * X ^ 24, Y) | lgen(R, cos(X), Y))).

cnf(cos_lower_bound_7, axiom,
    (~ lgen(R, Y,
            1 - 1/2 * X ^ 2 + 1/24 * X ^ 4 - 1/720 * X ^ 6 +
            1/40320 * X ^ 8 - 1/3628800 * X ^ 10 + 1/479001600 * X ^ 12 -
            1/87178291200 * X ^ 14 + 1/20922789888000 * X ^ 16 -
            1/6402373705728000 * X ^ 18 + 1/2432902008176640000 * X ^ 20 -
            1/1124000727777607680000 * X ^ 22) | lgen(R, Y, cos(X)))).

fof(trig_poly_352, conjecture,
    (! [X] :
       ((0 < X & X < 24/5) =>
        abs(1/2 + cos(X) + cos(2 * X)) <= 1/5 + 1 / abs(2 * sin(X / 2))))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 < X & X < 24/5) =>
        abs(1/2 + cos(X) + cos(2 * X)) <= 1/5 + 1 / abs(2 * sin(X / 2)))),
    inference(strip, [], [trig_poly_352])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 < X & X < 24/5) =>
          abs(1/2 + cos(X) + cos(2 * X)) <=
          1/5 + 1 / abs(2 * sin(X / 2)))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 < X & X < 24/5 &
        1/5 + 1 / abs(2 * sin(X / 2)) < abs(1/2 + cos(X) + cos(2 * X)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoXC1 < 24/5 &
     1/5 + 1 / abs(2 * sin(skoXC1 / 2)) <
     abs(1/2 + cos(skoXC1) + cos(2 * skoXC1)) & 0 < skoXC1),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1/5 + 1 / abs(2 * sin(skoXC1 / 2)) <
     abs(1/2 + cos(skoXC1) + cos(2 * skoXC1))),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 < skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 < 24/5),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (2 +
        skoXC1 *
        (skoXC1 *
         (-2/3 +
          skoXC1 *
          (skoXC1 *
           (4/45 +
            skoXC1 *
            (skoXC1 *
             (-2/315 +
              skoXC1 *
              (skoXC1 *
               (4/14175 +
                skoXC1 *
                (skoXC1 *
                 (-4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (-2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (-4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             -4/
                             147926426347074375)))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     skoXC1 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) <=
     (1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1 | skoXC1 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_1, plain,
    (~ lgen(0, cos(X_000018), X_000019) | cos(X_000018) <= X_000019),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain,
    (~ lgen(0,
            1 - 1/2 * X_000018 ^ 2 + 1/24 * X_000018 ^ 4 -
            1/720 * X_000018 ^ 6 + 1/40320 * X_000018 ^ 8 -
            1/3628800 * X_000018 ^ 10 + 1/479001600 * X_000018 ^ 12 -
            1/87178291200 * X_000018 ^ 14 +
            1/20922789888000 * X_000018 ^ 16 -
            1/6402373705728000 * X_000018 ^ 18 +
            1/2432902008176640000 * X_000018 ^ 20 -
            1/1124000727777607680000 * X_000018 ^ 22 +
            1/620448401733239439360000 * X_000018 ^ 24, X_000019) |
     lgen(0, cos(X_000018), X_000019)),
    inference(subst, [], [cos_upper_bound_7])).

cnf(refute_0_3, plain,
    (~ lgen(0,
            1 - 1/2 * X_000018 ^ 2 + 1/24 * X_000018 ^ 4 -
            1/720 * X_000018 ^ 6 + 1/40320 * X_000018 ^ 8 -
            1/3628800 * X_000018 ^ 10 + 1/479001600 * X_000018 ^ 12 -
            1/87178291200 * X_000018 ^ 14 +
            1/20922789888000 * X_000018 ^ 16 -
            1/6402373705728000 * X_000018 ^ 18 +
            1/2432902008176640000 * X_000018 ^ 20 -
            1/1124000727777607680000 * X_000018 ^ 22 +
            1/620448401733239439360000 * X_000018 ^ 24, X_000019) |
     cos(X_000018) <= X_000019),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (X_000019 <
     1 +
     X_000018 *
     (X_000018 *
      (-1/2 +
       X_000018 *
       (X_000018 *
        (1/24 +
         X_000018 *
         (X_000018 *
          (-1/720 +
           X_000018 *
           (X_000018 *
            (1/40320 +
             X_000018 *
             (X_000018 *
              (-1/3628800 +
               X_000018 *
               (X_000018 *
                (1/479001600 +
                 X_000018 *
                 (X_000018 *
                  (-1/87178291200 +
                   X_000018 *
                   (X_000018 *
                    (1/20922789888000 +
                     X_000018 *
                     (X_000018 *
                      (-1/6402373705728000 +
                       X_000018 *
                       (X_000018 *
                        (1/2432902008176640000 +
                         X_000018 *
                         (X_000018 *
                          (-1/1124000727777607680000 +
                           X_000018 *
                           (X_000018 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(X_000018) <= X_000019), inference(arithmetic, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(skoXC1) <=
     (1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1), inference(subst, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (-8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             4/147926426347074375)))))))))))))))))))))))) *
     skoXC1 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          (-8/2143861251406875 +
                           skoXC1 *
                           (skoXC1 *
                            4/147926426347074375))))))))))))))))))))))) <=
     (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 |
     skoXC1 <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_7, plain,
    ((1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 <
     1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          (-1/1124000727777607680000 +
                           skoXC1 * 2 *
                           (skoXC1 * 2 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(skoXC1 * 2) <=
     (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_8, plain,
    (1 < skoXC1 * 1/2 * (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     skoXC1 * 1/2 <= 1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_9, plain,
    (~ lgen(0, sin(X_000049), X_000050) | sin(X_000049) <= X_000050),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (X_000049 < 0 | ~ lgen(0, X_000049, X_000050) |
     lgen(0, sin(X_000049), X_000050)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_11, plain,
    (X_000049 < 0 | ~ lgen(0, X_000049, X_000050) |
     sin(X_000049) <= X_000050),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (X_000050 < X_000049 | X_000049 < 0 | sin(X_000049) <= X_000050),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) < skoXC1 * 1/2 |
     sin(skoXC1 * 1/2) <=
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2)),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (1 < (3/10 + cos(skoXC1) + cos(skoXC1 * 2)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0 |
     3/10 + cos(skoXC1) + cos(skoXC1 * 2) <= 1 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_15, plain,
    (1/5 + 1 / abs(2 * sin(skoXC1 / 2)) <
     abs(1/2 + cos(skoXC1) + cos(2 * skoXC1))),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_16, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2))),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (1/2 + cos(skoXC1) + cos(skoXC1 * 2) < 0 |
     abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) =
     1/2 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_18, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + (1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) !=
     1/2 + cos(skoXC1) + cos(skoXC1 * 2) |
     -1/5 + abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)), introduced(tautology, [equality])).

cnf(refute_0_19, plain,
    (1/2 + cos(skoXC1) + cos(skoXC1 * 2) < 0 |
     1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + (1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     -1/5 + abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_17, refute_0_18])).

cnf(refute_0_20, plain,
    (1/2 + cos(skoXC1) + cos(skoXC1 * 2) < 0 |
     1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + (1/2 + cos(skoXC1) + cos(skoXC1 * 2))),
    inference(resolve, [], [refute_0_19, refute_0_16])).

cnf(refute_0_21, plain,
    (cos(skoXC1 * 2) < -1/2 + cos(skoXC1) * -1 |
     1 / abs(sin(skoXC1 * 1/2) * 2) <
     3/10 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (3/4 +
          skoXC1 *
          (skoXC1 *
           (-7/60 +
            skoXC1 *
            (skoXC1 *
             (19/1890 +
              skoXC1 *
              (skoXC1 *
               (-31/56700 +
                skoXC1 *
                (skoXC1 *
                 (19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (-139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (-67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (-31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             1/6431583754220625)))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000012, cos(X_000011)) | X_000012 <= cos(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_24, plain,
    (~ lgen(0, X_000012,
            1 - 1/2 * X_000011 ^ 2 + 1/24 * X_000011 ^ 4 -
            1/720 * X_000011 ^ 6 + 1/40320 * X_000011 ^ 8 -
            1/3628800 * X_000011 ^ 10 + 1/479001600 * X_000011 ^ 12 -
            1/87178291200 * X_000011 ^ 14 +
            1/20922789888000 * X_000011 ^ 16 -
            1/6402373705728000 * X_000011 ^ 18 +
            1/2432902008176640000 * X_000011 ^ 20 -
            1/1124000727777607680000 * X_000011 ^ 22) |
     lgen(0, X_000012, cos(X_000011))),
    inference(subst, [], [cos_lower_bound_7])).

cnf(refute_0_25, plain,
    (~ lgen(0, X_000012,
            1 - 1/2 * X_000011 ^ 2 + 1/24 * X_000011 ^ 4 -
            1/720 * X_000011 ^ 6 + 1/40320 * X_000011 ^ 8 -
            1/3628800 * X_000011 ^ 10 + 1/479001600 * X_000011 ^ 12 -
            1/87178291200 * X_000011 ^ 14 +
            1/20922789888000 * X_000011 ^ 16 -
            1/6402373705728000 * X_000011 ^ 18 +
            1/2432902008176640000 * X_000011 ^ 20 -
            1/1124000727777607680000 * X_000011 ^ 22) |
     X_000012 <= cos(X_000011)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (1 +
     X_000011 *
     (X_000011 *
      (-1/2 +
       X_000011 *
       (X_000011 *
        (1/24 +
         X_000011 *
         (X_000011 *
          (-1/720 +
           X_000011 *
           (X_000011 *
            (1/40320 +
             X_000011 *
             (X_000011 *
              (-1/3628800 +
               X_000011 *
               (X_000011 *
                (1/479001600 +
                 X_000011 *
                 (X_000011 *
                  (-1/87178291200 +
                   X_000011 *
                   (X_000011 *
                    (1/20922789888000 +
                     X_000011 *
                     (X_000011 *
                      (-1/6402373705728000 +
                       X_000011 *
                       (X_000011 *
                        (1/2432902008176640000 +
                         X_000011 *
                         (X_000011 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     X_000012 | X_000012 <= cos(X_000011)),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <= cos(skoXC1)),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           -8/2143861251406875)))))))))))))))))))))) *
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          -8/2143861251406875)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_29, plain,
    (1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <= cos(skoXC1 * 2)),
    inference(subst, [], [refute_0_26])).

cnf(refute_0_30, plain,
    (1 <
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) *
     (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     0 <= -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 |
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <=
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_31, plain,
    (~ lgen(0, X_000071, sin(X_000070)) | X_000071 <= sin(X_000070)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_32, plain,
    (X_000070 < 0 | ~ lgen(0, X_000071, X_000070 - X_000070 ^ 3 / 6) |
     lgen(0, X_000071, sin(X_000070))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_33, plain,
    (X_000070 < 0 | ~ lgen(0, X_000071, X_000070 - X_000070 ^ 3 / 6) |
     X_000071 <= sin(X_000070)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (X_000070 * (1 + X_000070 * (X_000070 * -1/6)) < X_000071 |
     X_000070 < 0 | X_000071 <= sin(X_000070)),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <=
     sin(skoXC1 * 1/2)), inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (1 <
     (-7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1) *
     (sin(skoXC1 * 1/2) * 2) | sin(skoXC1 * 1/2) * 2 <= 0 |
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 <=
     1 / (sin(skoXC1 * 1/2) * 2)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_37, plain,
    (0 <= 1/2 + cos(skoXC1) + cos(skoXC1 * 2) |
     abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) =
     -(1/2 + cos(skoXC1) + cos(skoXC1 * 2))),
    inference(subst, [], [abs_negative])).

cnf(refute_0_38, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + -(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) !=
     -(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     -1/5 + abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)), introduced(tautology, [equality])).

cnf(refute_0_39, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + -(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     0 <= 1/2 + cos(skoXC1) + cos(skoXC1 * 2) |
     -1/5 + abs(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_37, refute_0_38])).

cnf(refute_0_40, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -1/5 + -(1/2 + cos(skoXC1) + cos(skoXC1 * 2)) |
     0 <= 1/2 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_39, refute_0_16])).

cnf(refute_0_41, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_40])).

cnf(refute_0_42, plain,
    (sin(skoXC1 * 1/2) * 2 < 0 |
     abs(sin(skoXC1 * 1/2) * 2) = sin(skoXC1 * 1/2) * 2),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_43, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     abs(sin(skoXC1 * 1/2) * 2) != sin(skoXC1 * 1/2) * 2 |
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 <=
     1 / abs(sin(skoXC1 * 1/2) * 2)), introduced(tautology, [equality])).

cnf(refute_0_44, plain,
    (sin(skoXC1 * 1/2) * 2 < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 <=
     1 / abs(sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (sin(skoXC1 * 1/2) * 2 < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_44, refute_0_41])).

cnf(refute_0_46, plain,
    (sin(skoXC1 * 1/2) < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) < 0 |
     skoXC1 * 1/2 < 0 | 0 <= sin(skoXC1 * 1/2)),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_48, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/48)) |
     0 <= sin(skoXC1 * 1/2)), inference(arithmetic, [], [refute_0_47])).

cnf(refute_0_49, plain, (0 < skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_50, plain, (skoXC1 < 24/5),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_51, plain,
    (skoXC1 < 0 | skoXC1 * (-1/2 + skoXC1 * (skoXC1 * 1/48)) <= 0),
    inference(decision, [], [refute_0_49, refute_0_50])).

cnf(refute_0_52, plain, (skoXC1 < 0 | 0 <= sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_51, refute_0_48])).

cnf(refute_0_53, plain, (0 <= skoXC1),
    inference(decision, [], [refute_0_49, refute_0_50])).

cnf(refute_0_54, plain, (0 <= sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_53, refute_0_52])).

cnf(refute_0_55, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) <
     -7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_54, refute_0_46])).

cnf(refute_0_56, plain,
    (1 <
     (-7/10 + cos(skoXC1) * -1 + cos(skoXC1 * 2) * -1) *
     (sin(skoXC1 * 1/2) * 2) | sin(skoXC1 * 1/2) * 2 <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_36, refute_0_55])).

cnf(refute_0_57, plain,
    (1 <
     sin(skoXC1 * 1/2) * (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) | sin(skoXC1 * 1/2) <= 0),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (X_000071 < sin(X_000070) | ~ lgen(1, X_000071, sin(X_000070))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_59, plain,
    (X_000070 < 0 | ~ lgen(1, X_000071, X_000070 - X_000070 ^ 3 / 6) |
     lgen(1, X_000071, sin(X_000070))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_60, plain,
    (X_000070 < 0 | X_000071 < sin(X_000070) |
     ~ lgen(1, X_000071, X_000070 - X_000070 ^ 3 / 6)),
    inference(resolve, [], [refute_0_59, refute_0_58])).

cnf(refute_0_61, plain,
    (X_000070 < 0 | X_000071 < sin(X_000070) |
     X_000070 * (1 + X_000070 * (X_000070 * -1/6)) <= X_000071),
    inference(arithmetic, [], [refute_0_60])).

cnf(refute_0_62, plain,
    (skoXC1 * 1/2 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) <= 0),
    inference(subst, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2) |
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) <= 0),
    inference(arithmetic, [], [refute_0_62])).

cnf(refute_0_64, plain,
    (skoXC1 < 0 | 0 < skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48))),
    inference(decision, [], [refute_0_50, refute_0_49])).

cnf(refute_0_65, plain, (skoXC1 < 0 | 0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain, (0 < sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_53, refute_0_65])).

cnf(refute_0_67, plain,
    (1 <
     sin(skoXC1 * 1/2) * (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_57, refute_0_66])).

cnf(refute_0_68, plain,
    (sin(skoXC1 * 1/2) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     0 <= -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (sin(skoXC1 * 1/2) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     cos(skoXC1 * 2) <= -7/10 + cos(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (1 < 0 | -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 < 0 |
     0 < -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_67])).

cnf(refute_0_71, plain,
    (cos(skoXC1 * 2) < -7/10 + cos(skoXC1) * -1 |
     -7/10 + cos(skoXC1) * -1 < cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_70])).

cnf(refute_0_72, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (14417/19200 +
          skoXC1 *
          (skoXC1 *
           (-379697/3225600 +
            skoXC1 *
            (skoXC1 *
             (1381031/132710400 +
              skoXC1 *
              (skoXC1 *
               (-5119/8601600 +
                skoXC1 *
                (skoXC1 *
                 (20351/851558400 +
                  skoXC1 *
                  (skoXC1 *
                   (-413717/581188608000 +
                    skoXC1 *
                    (skoXC1 *
                     (57049/3487131648000 +
                      skoXC1 *
                      (skoXC1 *
                       (-79733/266765571072000 +
                        skoXC1 *
                        (skoXC1 *
                         (281/63243887616000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32051/585417045717504000 +
                            skoXC1 *
                            (skoXC1 *
                             (5813/10752557982566400000 +
                              skoXC1 *
                              (skoXC1 *
                               (-8821/2370939035155891200000 +
                                skoXC1 *
                                (skoXC1 *
                                 (269/16596573246091238400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1/
                                   24894859869136857600000)))))))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_73, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) <=
     cos(skoXC1)), inference(subst, [], [refute_0_26])).

cnf(refute_0_74, plain,
    (1 +
     skoXC1 *
     (7/10 +
      skoXC1 *
      (skoXC1 *
       (-7/240 +
        skoXC1 *
        (skoXC1 *
         (7/19200 +
          skoXC1 *
          (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           -8/2143861251406875)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) <=
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          -8/2143861251406875)))))))))))))))))))))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_75, plain,
    (1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) <=
     cos(skoXC1 * 2)), inference(subst, [], [refute_0_26])).

cnf(refute_0_76, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 * (-1/645120 + skoXC1 * (skoXC1 * 1/185794560)))))))) *
     (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 * (-1/645120 + skoXC1 * (skoXC1 * 1/185794560)))))))) <=
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_77, plain,
    (1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     sin(skoXC1 * 1/2) | -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 <= 0),
    inference(split, [], [refute_0_67])).

cnf(refute_0_78, plain,
    (1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     sin(skoXC1 * 1/2) | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_77])).

cnf(refute_0_79, plain,
    (~ lgen(0, sin(X_000077), X_000078) | sin(X_000077) <= X_000078),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_80, plain,
    (X_000077 < 0 |
     ~ lgen(0,
            X_000077 - X_000077 ^ 3 / 6 + X_000077 ^ 5 / 120 -
            X_000077 ^ 7 / 5040 + X_000077 ^ 9 / 362880, X_000078) |
     lgen(0, sin(X_000077), X_000078)),
    inference(subst, [], [sin_upper_bound_2_pos])).

cnf(refute_0_81, plain,
    (X_000077 < 0 |
     ~ lgen(0,
            X_000077 - X_000077 ^ 3 / 6 + X_000077 ^ 5 / 120 -
            X_000077 ^ 7 / 5040 + X_000077 ^ 9 / 362880, X_000078) |
     sin(X_000077) <= X_000078),
    inference(resolve, [], [refute_0_80, refute_0_79])).

cnf(refute_0_82, plain,
    (X_000077 < 0 |
     X_000078 <
     X_000077 *
     (1 +
      X_000077 *
      (X_000077 *
       (-1/6 +
        X_000077 *
        (X_000077 *
         (1/120 +
          X_000077 *
          (X_000077 * (-1/5040 + X_000077 * (X_000077 * 1/362880)))))))) |
     sin(X_000077) <= X_000078), inference(arithmetic, [], [refute_0_81])).

cnf(refute_0_83, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/362880)))))))) |
     sin(skoXC1 * 1/2) <=
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2)),
    inference(subst, [], [refute_0_82])).

cnf(refute_0_84, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     skoXC1 * 1/2 *
     (1 +
      skoXC1 * 1/2 *
      (skoXC1 * 1/2 *
       (-1/6 +
        skoXC1 * 1/2 *
        (skoXC1 * 1/2 *
         (1/120 +
          skoXC1 * 1/2 *
          (skoXC1 * 1/2 *
           (-1/5040 + skoXC1 * 1/2 * (skoXC1 * 1/2 * 1/362880)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_83, refute_0_78])).

cnf(refute_0_85, plain,
    (skoXC1 < 0 |
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 * (-1/645120 + skoXC1 * (skoXC1 * 1/185794560)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 * (-1/645120 + skoXC1 * (skoXC1 * 1/185794560)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_53, refute_0_85])).

cnf(refute_0_87, plain,
    (1 <
     skoXC1 *
     (1/2 +
      skoXC1 *
      (skoXC1 *
       (-1/48 +
        skoXC1 *
        (skoXC1 *
         (1/3840 +
          skoXC1 *
          (skoXC1 * (-1/645120 + skoXC1 * (skoXC1 * 1/185794560)))))))) *
     (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 <= 0),
    inference(resolve, [], [refute_0_76, refute_0_86])).

cnf(refute_0_88, plain,
    (1 +
     skoXC1 *
     (7/10 +
      skoXC1 *
      (skoXC1 *
       (-7/240 +
        skoXC1 *
        (skoXC1 *
         (7/19200 +
          skoXC1 *
          (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280))))))))) <
     cos(skoXC1 * 2) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_87])).

cnf(refute_0_89, plain,
    (cos(skoXC1 * 2) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_88])).

cnf(refute_0_90, plain,
    (cos(skoXC1 * 2) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0 | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (0 <
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280))))))))),
    inference(decision, [], [refute_0_50, refute_0_49])).

cnf(refute_0_92, plain,
    (cos(skoXC1 * 2) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain,
    (1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_75, refute_0_92])).

cnf(refute_0_94, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          -8/2143861251406875))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (7/10 +
       skoXC1 *
       (skoXC1 *
        (-7/240 +
         skoXC1 *
         (skoXC1 *
          (7/19200 +
           skoXC1 *
           (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
      cos(skoXC1) *
      (skoXC1 *
       (1 +
        skoXC1 *
        (skoXC1 *
         (-1/24 +
          skoXC1 *
          (skoXC1 *
           (1/1920 +
            skoXC1 *
            (skoXC1 *
             (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_93])).

cnf(refute_0_95, plain,
    (1 +
     skoXC1 *
     (7/10 +
      skoXC1 *
      (skoXC1 *
       (-7/240 +
        skoXC1 *
        (skoXC1 *
         (7/19200 +
          skoXC1 *
          (skoXC1 * (-1/460800 + skoXC1 * (skoXC1 * 1/132710400)))))))) +
     cos(skoXC1) *
     (skoXC1 *
      (1 +
       skoXC1 *
       (skoXC1 *
        (-1/24 +
         skoXC1 *
         (skoXC1 *
          (1/1920 +
           skoXC1 *
           (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           -8/2143861251406875)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_74, refute_0_94])).

cnf(refute_0_96, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (14417/19200 +
          skoXC1 *
          (skoXC1 *
           (-379697/3225600 +
            skoXC1 *
            (skoXC1 *
             (1381031/132710400 +
              skoXC1 *
              (skoXC1 *
               (-5119/8601600 +
                skoXC1 *
                (skoXC1 *
                 (20351/851558400 +
                  skoXC1 *
                  (skoXC1 *
                   (-413717/581188608000 +
                    skoXC1 *
                    (skoXC1 *
                     (57049/3487131648000 +
                      skoXC1 *
                      (skoXC1 *
                       (-79733/266765571072000 +
                        skoXC1 *
                        (skoXC1 *
                         (281/63243887616000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32051/585417045717504000 +
                            skoXC1 *
                            (skoXC1 *
                             (5813/10752557982566400000 +
                              skoXC1 *
                              (skoXC1 *
                               (-8821/2370939035155891200000 +
                                skoXC1 *
                                (skoXC1 *
                                 (269/16596573246091238400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1/
                                   24894859869136857600000)))))))))))))))))))))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0 | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_95])).

cnf(refute_0_97, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (14417/19200 +
          skoXC1 *
          (skoXC1 *
           (-379697/3225600 +
            skoXC1 *
            (skoXC1 *
             (1381031/132710400 +
              skoXC1 *
              (skoXC1 *
               (-5119/8601600 +
                skoXC1 *
                (skoXC1 *
                 (20351/851558400 +
                  skoXC1 *
                  (skoXC1 *
                   (-413717/581188608000 +
                    skoXC1 *
                    (skoXC1 *
                     (57049/3487131648000 +
                      skoXC1 *
                      (skoXC1 *
                       (-79733/266765571072000 +
                        skoXC1 *
                        (skoXC1 *
                         (281/63243887616000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32051/585417045717504000 +
                            skoXC1 *
                            (skoXC1 *
                             (5813/10752557982566400000 +
                              skoXC1 *
                              (skoXC1 *
                               (-8821/2370939035155891200000 +
                                skoXC1 *
                                (skoXC1 *
                                 (269/16596573246091238400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1/
                                   24894859869136857600000)))))))))))))))))))))))))))))) <
     cos(skoXC1) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_96, refute_0_91])).

cnf(refute_0_98, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_97])).

cnf(refute_0_99, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0 | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_99, refute_0_91])).

cnf(refute_0_101, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (14417/19200 +
           skoXC1 *
           (skoXC1 *
            (-379697/3225600 +
             skoXC1 *
             (skoXC1 *
              (1381031/132710400 +
               skoXC1 *
               (skoXC1 *
                (-5119/8601600 +
                 skoXC1 *
                 (skoXC1 *
                  (20351/851558400 +
                   skoXC1 *
                   (skoXC1 *
                    (-413717/581188608000 +
                     skoXC1 *
                     (skoXC1 *
                      (57049/3487131648000 +
                       skoXC1 *
                       (skoXC1 *
                        (-79733/266765571072000 +
                         skoXC1 *
                         (skoXC1 *
                          (281/63243887616000 +
                           skoXC1 *
                           (skoXC1 *
                            (-32051/585417045717504000 +
                             skoXC1 *
                             (skoXC1 *
                              (5813/10752557982566400000 +
                               skoXC1 *
                               (skoXC1 *
                                (-8821/2370939035155891200000 +
                                 skoXC1 *
                                 (skoXC1 *
                                  (269/16596573246091238400000 +
                                   skoXC1 *
                                   (skoXC1 *
                                    -1/
                                    24894859869136857600000))))))))))))))))))))))))))))))) /
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_73, refute_0_100])).

cnf(refute_0_102, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (14417/19200 +
          skoXC1 *
          (skoXC1 *
           (-379697/3225600 +
            skoXC1 *
            (skoXC1 *
             (1381031/132710400 +
              skoXC1 *
              (skoXC1 *
               (-5119/8601600 +
                skoXC1 *
                (skoXC1 *
                 (20351/851558400 +
                  skoXC1 *
                  (skoXC1 *
                   (-413717/581188608000 +
                    skoXC1 *
                    (skoXC1 *
                     (57049/3487131648000 +
                      skoXC1 *
                      (skoXC1 *
                       (-79733/266765571072000 +
                        skoXC1 *
                        (skoXC1 *
                         (281/63243887616000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32051/585417045717504000 +
                            skoXC1 *
                            (skoXC1 *
                             (5813/10752557982566400000 +
                              skoXC1 *
                              (skoXC1 *
                               (-8821/2370939035155891200000 +
                                skoXC1 *
                                (skoXC1 *
                                 (269/16596573246091238400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   -1/
                                   24894859869136857600000)))))))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 *
      (-1 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/1920 +
           skoXC1 *
           (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280))))))))) |
     0 <=
     skoXC1 *
     (-1 +
      skoXC1 *
      (skoXC1 *
       (1/24 +
        skoXC1 *
        (skoXC1 *
         (-1/1920 +
          skoXC1 *
          (skoXC1 * (1/322560 + skoXC1 * (skoXC1 * -1/92897280)))))))) |
     -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_72, refute_0_101])).

cnf(refute_0_103, plain,
    (1 <
     skoXC1 *
     (-27/10 +
      skoXC1 *
      (skoXC1 *
       (209/80 +
        skoXC1 *
        (skoXC1 *
         (-5209/6400 +
          skoXC1 *
          (skoXC1 *
           (43403/358400 +
            skoXC1 *
            (skoXC1 *
             (-3255209/309657600 +
              skoXC1 *
              (skoXC1 *
               (36991/61931520 +
                skoXC1 *
                (skoXC1 *
                 (-39125/1634992128 +
                  skoXC1 *
                  (skoXC1 *
                   (5298061/7439214182400 +
                    skoXC1 *
                    (skoXC1 *
                     (-973799/59513713459200 +
                      skoXC1 *
                      (skoXC1 *
                       (81651389/273167944777728000 +
                        skoXC1 *
                        (skoXC1 *
                         (-58700963/13211395147431936000 +
                          skoXC1 *
                          (skoXC1 *
                           (2625635713/47957364385177927680000 +
                            skoXC1 *
                            (skoXC1 *
                             (-18667120789/34529302357328107929600000 +
                              skoXC1 *
                              (skoXC1 *
                               (2312374349/621527442431905942732800000 +
                                skoXC1 *
                                (skoXC1 *
                                 (-282067069/
                                  17402768388093366396518400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   838861/
                                   20883322065712039675822080000)))))))))))))))))))))))))))))) |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0 | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (skoXC1 *
     (-27/10 +
      skoXC1 *
      (skoXC1 *
       (209/80 +
        skoXC1 *
        (skoXC1 *
         (-5209/6400 +
          skoXC1 *
          (skoXC1 *
           (43403/358400 +
            skoXC1 *
            (skoXC1 *
             (-3255209/309657600 +
              skoXC1 *
              (skoXC1 *
               (36991/61931520 +
                skoXC1 *
                (skoXC1 *
                 (-39125/1634992128 +
                  skoXC1 *
                  (skoXC1 *
                   (5298061/7439214182400 +
                    skoXC1 *
                    (skoXC1 *
                     (-973799/59513713459200 +
                      skoXC1 *
                      (skoXC1 *
                       (81651389/273167944777728000 +
                        skoXC1 *
                        (skoXC1 *
                         (-58700963/13211395147431936000 +
                          skoXC1 *
                          (skoXC1 *
                           (2625635713/47957364385177927680000 +
                            skoXC1 *
                            (skoXC1 *
                             (-18667120789/34529302357328107929600000 +
                              skoXC1 *
                              (skoXC1 *
                               (2312374349/621527442431905942732800000 +
                                skoXC1 *
                                (skoXC1 *
                                 (-282067069/
                                  17402768388093366396518400000 +
                                  skoXC1 *
                                  (skoXC1 *
                                   838861/
                                   20883322065712039675822080000)))))))))))))))))))))))))))))) <=
     1 |
     skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0), inference(decision, [], [refute_0_49, refute_0_50])).

cnf(refute_0_105, plain,
    (skoXC1 *
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/24 +
        skoXC1 *
        (skoXC1 *
         (1/1920 +
          skoXC1 *
          (skoXC1 * (-1/322560 + skoXC1 * (skoXC1 * 1/92897280)))))))) <=
     0 | -7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_104, refute_0_103])).

cnf(refute_0_106, plain,
    (-7/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_105, refute_0_91])).

cnf(refute_0_107, plain,
    (-7/10 + cos(skoXC1) * -1 < cos(skoXC1 * 2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_106, refute_0_71])).

cnf(refute_0_108, plain,
    (sin(skoXC1 * 1/2) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_69, refute_0_107])).

cnf(refute_0_109, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_35, refute_0_108])).

cnf(refute_0_110, plain,
    (skoXC1 < 0 |
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_109])).

cnf(refute_0_111, plain,
    (skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) <
     1 / (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_53, refute_0_110])).

cnf(refute_0_112, plain,
    (1 <
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) *
     (-7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2) |
     0 <= -7/5 + cos(skoXC1) * -2 + cos(skoXC1 * 2) * -2 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_30, refute_0_111])).

cnf(refute_0_113, plain,
    (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     cos(skoXC1 * 2) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2) |
     cos(skoXC1 * 2) <= -7/10 + cos(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_112])).

cnf(refute_0_114, plain,
    (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     cos(skoXC1 * 2) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_113, refute_0_107])).

cnf(refute_0_115, plain,
    (cos(skoXC1 * 2) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_114])).

cnf(refute_0_116, plain,
    (cos(skoXC1 * 2) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_115])).

cnf(refute_0_117, plain, (0 < skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))),
    inference(decision, [], [refute_0_49, refute_0_50])).

cnf(refute_0_118, plain,
    (cos(skoXC1 * 2) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_116, refute_0_117])).

cnf(refute_0_119, plain,
    (1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_29, refute_0_118])).

cnf(refute_0_120, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          -8/2143861251406875))))))))))))))))))))) <
     (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
      cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_119])).

cnf(refute_0_121, plain,
    (1 + skoXC1 * (7/10 + skoXC1 * (skoXC1 * -7/240)) +
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           -8/2143861251406875)))))))))))))))))))))) *
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_28, refute_0_120])).

cnf(refute_0_122, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (3/4 +
          skoXC1 *
          (skoXC1 *
           (-7/60 +
            skoXC1 *
            (skoXC1 *
             (19/1890 +
              skoXC1 *
              (skoXC1 *
               (-31/56700 +
                skoXC1 *
                (skoXC1 *
                 (19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (-139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (-67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (-31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             1/6431583754220625)))))))))))))))))))))))) <
     cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_121])).

cnf(refute_0_123, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (3/4 +
          skoXC1 *
          (skoXC1 *
           (-7/60 +
            skoXC1 *
            (skoXC1 *
             (19/1890 +
              skoXC1 *
              (skoXC1 *
               (-31/56700 +
                skoXC1 *
                (skoXC1 *
                 (19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (-139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (-67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (-31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             1/6431583754220625)))))))))))))))))))))))) <
     cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_122, refute_0_117])).

cnf(refute_0_124, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_123])).

cnf(refute_0_125, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_124])).

cnf(refute_0_126, plain,
    (cos(skoXC1) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_125, refute_0_117])).

cnf(refute_0_127, plain,
    (1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          -1/1124000727777607680000))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (17/10 +
       skoXC1 *
       (skoXC1 *
        (-497/240 +
         skoXC1 *
         (skoXC1 *
          (3/4 +
           skoXC1 *
           (skoXC1 *
            (-7/60 +
             skoXC1 *
             (skoXC1 *
              (19/1890 +
               skoXC1 *
               (skoXC1 *
                (-31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (-139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (-67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (-31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              1/6431583754220625))))))))))))))))))))))))) /
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_27, refute_0_126])).

cnf(refute_0_128, plain,
    (1 +
     skoXC1 *
     (17/10 +
      skoXC1 *
      (skoXC1 *
       (-497/240 +
        skoXC1 *
        (skoXC1 *
         (3/4 +
          skoXC1 *
          (skoXC1 *
           (-7/60 +
            skoXC1 *
            (skoXC1 *
             (19/1890 +
              skoXC1 *
              (skoXC1 *
               (-31/56700 +
                skoXC1 *
                (skoXC1 *
                 (19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (-139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (-67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (-31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             1/6431583754220625)))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           -1/
                           1124000727777607680000)))))))))))))))))))))) *
     (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) |
     0 <= skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)) |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_22, refute_0_127])).

cnf(refute_0_129, plain,
    (1 <
     skoXC1 *
     (-27/10 +
      skoXC1 *
      (skoXC1 *
       (209/80 +
        skoXC1 *
        (skoXC1 *
         (-13/16 +
          skoXC1 *
          (skoXC1 *
           (23/192 +
            skoXC1 *
            (skoXC1 *
             (-613/60480 +
              skoXC1 *
              (skoXC1 *
               (1591/2903040 +
                skoXC1 *
                (skoXC1 *
                 (-19469/958003200 +
                  skoXC1 *
                  (skoXC1 *
                   (569447/1046139494400 +
                    skoXC1 *
                    (skoXC1 *
                     (-229387/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (627281/3658499260416000 +
                        skoXC1 *
                        (skoXC1 *
                         (-31195237/14597412049059840000 +
                          skoXC1 *
                          (skoXC1 *
                           (32505883/1498667637036810240000 +
                            skoXC1 *
                            (skoXC1 *
                             -838861/
                             5395203493332516864000)))))))))))))))))))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_128])).

cnf(refute_0_130, plain,
    (skoXC1 *
     (-27/10 +
      skoXC1 *
      (skoXC1 *
       (209/80 +
        skoXC1 *
        (skoXC1 *
         (-13/16 +
          skoXC1 *
          (skoXC1 *
           (23/192 +
            skoXC1 *
            (skoXC1 *
             (-613/60480 +
              skoXC1 *
              (skoXC1 *
               (1591/2903040 +
                skoXC1 *
                (skoXC1 *
                 (-19469/958003200 +
                  skoXC1 *
                  (skoXC1 *
                   (569447/1046139494400 +
                    skoXC1 *
                    (skoXC1 *
                     (-229387/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (627281/3658499260416000 +
                        skoXC1 *
                        (skoXC1 *
                         (-31195237/14597412049059840000 +
                          skoXC1 *
                          (skoXC1 *
                           (32505883/1498667637036810240000 +
                            skoXC1 *
                            (skoXC1 *
                             -838861/
                             5395203493332516864000)))))))))))))))))))))))) <=
     1 | skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0),
    inference(decision, [], [refute_0_50, refute_0_49])).

cnf(refute_0_131, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_130, refute_0_129])).

cnf(refute_0_132, plain, (-1/2 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_131, refute_0_117])).

cnf(refute_0_133, plain,
    (1 / abs(sin(skoXC1 * 1/2) * 2) <
     3/10 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_132, refute_0_21])).

cnf(refute_0_134, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < 3/10 + cos(skoXC1) + cos(skoXC1 * 2) |
     abs(sin(skoXC1 * 1/2) * 2) != sin(skoXC1 * 1/2) * 2 |
     3/10 + cos(skoXC1) + cos(skoXC1 * 2) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)), introduced(tautology, [equality])).

cnf(refute_0_135, plain,
    (sin(skoXC1 * 1/2) * 2 < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) < 3/10 + cos(skoXC1) + cos(skoXC1 * 2) |
     3/10 + cos(skoXC1) + cos(skoXC1 * 2) <=
     1 / abs(sin(skoXC1 * 1/2) * 2)),
    inference(resolve, [], [refute_0_42, refute_0_134])).

cnf(refute_0_136, plain,
    (sin(skoXC1 * 1/2) * 2 < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) < 3/10 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_135, refute_0_133])).

cnf(refute_0_137, plain,
    (sin(skoXC1 * 1/2) < 0 |
     1 / (sin(skoXC1 * 1/2) * 2) < 3/10 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_136])).

cnf(refute_0_138, plain,
    (1 / (sin(skoXC1 * 1/2) * 2) < 3/10 + cos(skoXC1) + cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_54, refute_0_137])).

cnf(refute_0_139, plain,
    (1 < (3/10 + cos(skoXC1) + cos(skoXC1 * 2)) * (sin(skoXC1 * 1/2) * 2) |
     sin(skoXC1 * 1/2) * 2 <= 0),
    inference(resolve, [], [refute_0_14, refute_0_138])).

cnf(refute_0_140, plain,
    (1 <
     sin(skoXC1 * 1/2) * (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     sin(skoXC1 * 1/2) <= 0), inference(arithmetic, [], [refute_0_139])).

cnf(refute_0_141, plain,
    (1 <
     sin(skoXC1 * 1/2) * (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2)),
    inference(resolve, [], [refute_0_140, refute_0_66])).

cnf(refute_0_142, plain,
    (1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) <
     sin(skoXC1 * 1/2) | 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 <= 0),
    inference(split, [], [refute_0_141])).

cnf(refute_0_143, plain,
    (1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) <
     sin(skoXC1 * 1/2) | cos(skoXC1 * 2) <= -3/10 + cos(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_142])).

cnf(refute_0_144, plain,
    (1 < 0 | 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 < 0 |
     0 < 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2),
    inference(split, [], [refute_0_141])).

cnf(refute_0_145, plain,
    (cos(skoXC1 * 2) < -3/10 + cos(skoXC1) * -1 |
     -3/10 + cos(skoXC1) * -1 < cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_144])).

cnf(refute_0_146, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (493/240 +
        skoXC1 *
        (skoXC1 *
         (-3/4 +
          skoXC1 *
          (skoXC1 *
           (7/60 +
            skoXC1 *
            (skoXC1 *
             (-19/1890 +
              skoXC1 *
              (skoXC1 *
               (31/56700 +
                skoXC1 *
                (skoXC1 *
                 (-19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (-1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             (-1/5478756531373125 +
                              skoXC1 *
                              (skoXC1 *
                               1/
                               887558558082446250)))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) <=
     (1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_147, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(skoXC1) <=
     (1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_148, plain,
    (1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
     cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (-8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             4/147926426347074375)))))))))))))))))))))))) *
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          (-8/2143861251406875 +
                           skoXC1 *
                           (skoXC1 *
                            4/147926426347074375))))))))))))))))))))))) <=
     (1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_149, plain,
    ((1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          (-1/1124000727777607680000 +
                           skoXC1 * 2 *
                           (skoXC1 * 2 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     cos(skoXC1 * 2) <=
     (1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)))),
    inference(subst, [], [refute_0_4])).

cnf(refute_0_150, plain,
    (1 <
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) *
     (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     0 <= 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 |
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) <=
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48))),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_151, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) <=
     sin(skoXC1 * 1/2)), inference(subst, [], [refute_0_34])).

cnf(refute_0_152, plain,
    (sin(skoXC1 * 1/2) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     0 <= 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2),
    inference(split, [], [refute_0_141])).

cnf(refute_0_153, plain,
    (sin(skoXC1 * 1/2) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_152])).

cnf(refute_0_154, plain,
    (skoXC1 * 1/2 < 0 |
     skoXC1 * 1/2 * (1 + skoXC1 * 1/2 * (skoXC1 * 1/2 * -1/6)) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_151, refute_0_153])).

cnf(refute_0_155, plain,
    (skoXC1 < 0 |
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_154])).

cnf(refute_0_156, plain,
    (skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) <
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_53, refute_0_155])).

cnf(refute_0_157, plain,
    (1 <
     skoXC1 * (1/2 + skoXC1 * (skoXC1 * -1/48)) *
     (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     0 <= 3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_150, refute_0_156])).

cnf(refute_0_158, plain,
    (1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
     cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <
     cos(skoXC1 * 2) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_157])).

cnf(refute_0_159, plain,
    ((1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) < cos(skoXC1 * 2) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_158])).

cnf(refute_0_160, plain,
    ((1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) < cos(skoXC1 * 2) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_159, refute_0_117])).

cnf(refute_0_161, plain,
    ((1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          (-1/1124000727777607680000 +
                           skoXC1 * 2 *
                           (skoXC1 * 2 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_149, refute_0_160])).

cnf(refute_0_162, plain,
    ((1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
      cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24)))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          (-8/2143861251406875 +
                           skoXC1 *
                           (skoXC1 *
                            4/147926426347074375))))))))))))))))))))))) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_161])).

cnf(refute_0_163, plain,
    (1 + skoXC1 * (-3/10 + skoXC1 * (skoXC1 * 1/80)) +
     cos(skoXC1) * (skoXC1 * (-1 + skoXC1 * (skoXC1 * 1/24))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (-8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             4/147926426347074375)))))))))))))))))))))))) *
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_148, refute_0_162])).

cnf(refute_0_164, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (493/240 +
        skoXC1 *
        (skoXC1 *
         (-3/4 +
          skoXC1 *
          (skoXC1 *
           (7/60 +
            skoXC1 *
            (skoXC1 *
             (-19/1890 +
              skoXC1 *
              (skoXC1 *
               (31/56700 +
                skoXC1 *
                (skoXC1 *
                 (-19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (-1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             (-1/5478756531373125 +
                              skoXC1 *
                              (skoXC1 *
                               1/
                               887558558082446250)))))))))))))))))))))))))) <
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_163])).

cnf(refute_0_165, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (493/240 +
        skoXC1 *
        (skoXC1 *
         (-3/4 +
          skoXC1 *
          (skoXC1 *
           (7/60 +
            skoXC1 *
            (skoXC1 *
             (-19/1890 +
              skoXC1 *
              (skoXC1 *
               (31/56700 +
                skoXC1 *
                (skoXC1 *
                 (-19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (-1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             (-1/5478756531373125 +
                              skoXC1 *
                              (skoXC1 *
                               1/
                               887558558082446250)))))))))))))))))))))))))) <
     cos(skoXC1) * (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_164, refute_0_117])).

cnf(refute_0_166, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) < cos(skoXC1) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(split, [], [refute_0_165])).

cnf(refute_0_167, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) < cos(skoXC1) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_166, refute_0_117])).

cnf(refute_0_168, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (493/240 +
         skoXC1 *
         (skoXC1 *
          (-3/4 +
           skoXC1 *
           (skoXC1 *
            (7/60 +
             skoXC1 *
             (skoXC1 *
              (-19/1890 +
               skoXC1 *
               (skoXC1 *
                (31/56700 +
                 skoXC1 *
                 (skoXC1 *
                  (-19/935550 +
                   skoXC1 *
                   (skoXC1 *
                    (139/255405150 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/91216125 +
                       skoXC1 *
                       (skoXC1 *
                        (67/390769879500 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/467938721250 +
                           skoXC1 *
                           (skoXC1 *
                            (31/1429240834271250 +
                             skoXC1 *
                             (skoXC1 *
                              (-1/5478756531373125 +
                               skoXC1 *
                               (skoXC1 *
                                1/
                                887558558082446250))))))))))))))))))))))))))) /
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000))))))))))))))))))))))) |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_147, refute_0_167])).

cnf(refute_0_169, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (493/240 +
        skoXC1 *
        (skoXC1 *
         (-3/4 +
          skoXC1 *
          (skoXC1 *
           (7/60 +
            skoXC1 *
            (skoXC1 *
             (-19/1890 +
              skoXC1 *
              (skoXC1 *
               (31/56700 +
                skoXC1 *
                (skoXC1 *
                 (-19/935550 +
                  skoXC1 *
                  (skoXC1 *
                   (139/255405150 +
                    skoXC1 *
                    (skoXC1 *
                     (-1/91216125 +
                      skoXC1 *
                      (skoXC1 *
                       (67/390769879500 +
                        skoXC1 *
                        (skoXC1 *
                         (-1/467938721250 +
                          skoXC1 *
                          (skoXC1 *
                           (31/1429240834271250 +
                            skoXC1 *
                            (skoXC1 *
                             (-1/5478756531373125 +
                              skoXC1 *
                              (skoXC1 *
                               1/
                               887558558082446250)))))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_146, refute_0_168])).

cnf(refute_0_170, plain,
    (1 <
     skoXC1 *
     (23/10 +
      skoXC1 *
      (skoXC1 *
       (-623/240 +
        skoXC1 *
        (skoXC1 *
         (13/16 +
          skoXC1 *
          (skoXC1 *
           (-23/192 +
            skoXC1 *
            (skoXC1 *
             (613/60480 +
              skoXC1 *
              (skoXC1 *
               (-1591/2903040 +
                skoXC1 *
                (skoXC1 *
                 (19469/958003200 +
                  skoXC1 *
                  (skoXC1 *
                   (-569447/1046139494400 +
                    skoXC1 *
                    (skoXC1 *
                     (229387/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-627281/3658499260416000 +
                        skoXC1 *
                        (skoXC1 *
                         (31195237/14597412049059840000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32505883/1498667637036810240000 +
                            skoXC1 *
                            (skoXC1 *
                             (22577/123693859994664960000 +
                              skoXC1 *
                              (skoXC1 *
                               -16777217/
                               14890761641597746544640000)))))))))))))))))))))))))) |
     skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(arithmetic, [], [refute_0_169])).

cnf(refute_0_171, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     skoXC1 *
     (23/10 +
      skoXC1 *
      (skoXC1 *
       (-623/240 +
        skoXC1 *
        (skoXC1 *
         (13/16 +
          skoXC1 *
          (skoXC1 *
           (-23/192 +
            skoXC1 *
            (skoXC1 *
             (613/60480 +
              skoXC1 *
              (skoXC1 *
               (-1591/2903040 +
                skoXC1 *
                (skoXC1 *
                 (19469/958003200 +
                  skoXC1 *
                  (skoXC1 *
                   (-569447/1046139494400 +
                    skoXC1 *
                    (skoXC1 *
                     (229387/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-627281/3658499260416000 +
                        skoXC1 *
                        (skoXC1 *
                         (31195237/14597412049059840000 +
                          skoXC1 *
                          (skoXC1 *
                           (-32505883/1498667637036810240000 +
                            skoXC1 *
                            (skoXC1 *
                             (22577/123693859994664960000 +
                              skoXC1 *
                              (skoXC1 *
                               -16777217/
                               14890761641597746544640000)))))))))))))))))))))))))) <=
     1), inference(decision, [], [refute_0_50, refute_0_49])).

cnf(refute_0_172, plain,
    (skoXC1 * (1 + skoXC1 * (skoXC1 * -1/24)) <= 0 |
     -3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_171, refute_0_170])).

cnf(refute_0_173, plain, (-3/10 + cos(skoXC1) * -1 <= cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_172, refute_0_117])).

cnf(refute_0_174, plain, (-3/10 + cos(skoXC1) * -1 < cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_173, refute_0_145])).

cnf(refute_0_175, plain,
    (1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) <
     sin(skoXC1 * 1/2)),
    inference(resolve, [], [refute_0_143, refute_0_174])).

cnf(refute_0_176, plain,
    (skoXC1 * 1/2 < 0 |
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) < skoXC1 * 1/2),
    inference(resolve, [], [refute_0_13, refute_0_175])).

cnf(refute_0_177, plain,
    (skoXC1 < 0 |
     1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) < skoXC1 * 1/2),
    inference(arithmetic, [], [refute_0_176])).

cnf(refute_0_178, plain,
    (1 / (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) < skoXC1 * 1/2),
    inference(resolve, [], [refute_0_53, refute_0_177])).

cnf(refute_0_179, plain,
    (1 < skoXC1 * 1/2 * (3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2) |
     3/5 + cos(skoXC1) * 2 + cos(skoXC1 * 2) * 2 <= 0),
    inference(resolve, [], [refute_0_8, refute_0_178])).

cnf(refute_0_180, plain,
    (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1) <
     cos(skoXC1 * 2) * skoXC1 |
     cos(skoXC1 * 2) <= -3/10 + cos(skoXC1) * -1),
    inference(arithmetic, [], [refute_0_179])).

cnf(refute_0_181, plain,
    (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1) <
     cos(skoXC1 * 2) * skoXC1),
    inference(resolve, [], [refute_0_180, refute_0_174])).

cnf(refute_0_182, plain,
    ((1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 <
     cos(skoXC1 * 2) | skoXC1 <= 0), inference(split, [], [refute_0_181])).

cnf(refute_0_183, plain,
    ((1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 <
     cos(skoXC1 * 2)),
    inference(resolve, [], [refute_0_182, refute_0_49])).

cnf(refute_0_184, plain,
    ((1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 <
     1 +
     skoXC1 * 2 *
     (skoXC1 * 2 *
      (-1/2 +
       skoXC1 * 2 *
       (skoXC1 * 2 *
        (1/24 +
         skoXC1 * 2 *
         (skoXC1 * 2 *
          (-1/720 +
           skoXC1 * 2 *
           (skoXC1 * 2 *
            (1/40320 +
             skoXC1 * 2 *
             (skoXC1 * 2 *
              (-1/3628800 +
               skoXC1 * 2 *
               (skoXC1 * 2 *
                (1/479001600 +
                 skoXC1 * 2 *
                 (skoXC1 * 2 *
                  (-1/87178291200 +
                   skoXC1 * 2 *
                   (skoXC1 * 2 *
                    (1/20922789888000 +
                     skoXC1 * 2 *
                     (skoXC1 * 2 *
                      (-1/6402373705728000 +
                       skoXC1 * 2 *
                       (skoXC1 * 2 *
                        (1/2432902008176640000 +
                         skoXC1 * 2 *
                         (skoXC1 * 2 *
                          (-1/1124000727777607680000 +
                           skoXC1 * 2 *
                           (skoXC1 * 2 *
                            1/
                            620448401733239439360000)))))))))))))))))))))))),
    inference(resolve, [], [refute_0_7, refute_0_183])).

cnf(refute_0_185, plain,
    ((1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1)) / skoXC1 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-2 +
       skoXC1 *
       (skoXC1 *
        (2/3 +
         skoXC1 *
         (skoXC1 *
          (-4/45 +
           skoXC1 *
           (skoXC1 *
            (2/315 +
             skoXC1 *
             (skoXC1 *
              (-4/14175 +
               skoXC1 *
               (skoXC1 *
                (4/467775 +
                 skoXC1 *
                 (skoXC1 *
                  (-8/42567525 +
                   skoXC1 *
                   (skoXC1 *
                    (2/638512875 +
                     skoXC1 *
                     (skoXC1 *
                      (-4/97692469875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/9280784638125 +
                         skoXC1 *
                         (skoXC1 *
                          (-8/2143861251406875 +
                           skoXC1 *
                           (skoXC1 *
                            4/147926426347074375)))))))))))))))))))))))),
    inference(arithmetic, [], [refute_0_184])).

cnf(refute_0_186, plain,
    (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (-8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             4/147926426347074375)))))))))))))))))))))))) *
     skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_6, refute_0_185])).

cnf(refute_0_187, plain,
    (1 + skoXC1 * -3/10 + cos(skoXC1) * (skoXC1 * -1) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-2 +
        skoXC1 *
        (skoXC1 *
         (2/3 +
          skoXC1 *
          (skoXC1 *
           (-4/45 +
            skoXC1 *
            (skoXC1 *
             (2/315 +
              skoXC1 *
              (skoXC1 *
               (-4/14175 +
                skoXC1 *
                (skoXC1 *
                 (4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (-8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (-4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (-8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             4/147926426347074375)))))))))))))))))))))))) *
     skoXC1), inference(resolve, [], [refute_0_186, refute_0_49])).

cnf(refute_0_188, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (2 +
        skoXC1 *
        (skoXC1 *
         (-2/3 +
          skoXC1 *
          (skoXC1 *
           (4/45 +
            skoXC1 *
            (skoXC1 *
             (-2/315 +
              skoXC1 *
              (skoXC1 *
               (4/14175 +
                skoXC1 *
                (skoXC1 *
                 (-4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (-2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (-4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             -4/
                             147926426347074375)))))))))))))))))))))))) <
     cos(skoXC1) * skoXC1), inference(arithmetic, [], [refute_0_187])).

cnf(refute_0_189, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1 < cos(skoXC1) | skoXC1 <= 0),
    inference(split, [], [refute_0_188])).

cnf(refute_0_190, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1 < cos(skoXC1)),
    inference(resolve, [], [refute_0_189, refute_0_49])).

cnf(refute_0_191, plain,
    ((1 +
      skoXC1 *
      (-13/10 +
       skoXC1 *
       (skoXC1 *
        (2 +
         skoXC1 *
         (skoXC1 *
          (-2/3 +
           skoXC1 *
           (skoXC1 *
            (4/45 +
             skoXC1 *
             (skoXC1 *
              (-2/315 +
               skoXC1 *
               (skoXC1 *
                (4/14175 +
                 skoXC1 *
                 (skoXC1 *
                  (-4/467775 +
                   skoXC1 *
                   (skoXC1 *
                    (8/42567525 +
                     skoXC1 *
                     (skoXC1 *
                      (-2/638512875 +
                       skoXC1 *
                       (skoXC1 *
                        (4/97692469875 +
                         skoXC1 *
                         (skoXC1 *
                          (-4/9280784638125 +
                           skoXC1 *
                           (skoXC1 *
                            (8/2143861251406875 +
                             skoXC1 *
                             (skoXC1 *
                              -4/
                              147926426347074375))))))))))))))))))))))))) /
     skoXC1 <
     1 +
     skoXC1 *
     (skoXC1 *
      (-1/2 +
       skoXC1 *
       (skoXC1 *
        (1/24 +
         skoXC1 *
         (skoXC1 *
          (-1/720 +
           skoXC1 *
           (skoXC1 *
            (1/40320 +
             skoXC1 *
             (skoXC1 *
              (-1/3628800 +
               skoXC1 *
               (skoXC1 *
                (1/479001600 +
                 skoXC1 *
                 (skoXC1 *
                  (-1/87178291200 +
                   skoXC1 *
                   (skoXC1 *
                    (1/20922789888000 +
                     skoXC1 *
                     (skoXC1 *
                      (-1/6402373705728000 +
                       skoXC1 *
                       (skoXC1 *
                        (1/2432902008176640000 +
                         skoXC1 *
                         (skoXC1 *
                          (-1/1124000727777607680000 +
                           skoXC1 *
                           (skoXC1 *
                            1/
                            620448401733239439360000)))))))))))))))))))))))),
    inference(resolve, [], [refute_0_5, refute_0_190])).

cnf(refute_0_192, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (2 +
        skoXC1 *
        (skoXC1 *
         (-2/3 +
          skoXC1 *
          (skoXC1 *
           (4/45 +
            skoXC1 *
            (skoXC1 *
             (-2/315 +
              skoXC1 *
              (skoXC1 *
               (4/14175 +
                skoXC1 *
                (skoXC1 *
                 (-4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (-2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (-4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             -4/
                             147926426347074375)))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     skoXC1 | skoXC1 <= 0),
    inference(resolve, [], [refute_0_0, refute_0_191])).

cnf(refute_0_193, plain,
    (1 +
     skoXC1 *
     (-13/10 +
      skoXC1 *
      (skoXC1 *
       (2 +
        skoXC1 *
        (skoXC1 *
         (-2/3 +
          skoXC1 *
          (skoXC1 *
           (4/45 +
            skoXC1 *
            (skoXC1 *
             (-2/315 +
              skoXC1 *
              (skoXC1 *
               (4/14175 +
                skoXC1 *
                (skoXC1 *
                 (-4/467775 +
                  skoXC1 *
                  (skoXC1 *
                   (8/42567525 +
                    skoXC1 *
                    (skoXC1 *
                     (-2/638512875 +
                      skoXC1 *
                      (skoXC1 *
                       (4/97692469875 +
                        skoXC1 *
                        (skoXC1 *
                         (-4/9280784638125 +
                          skoXC1 *
                          (skoXC1 *
                           (8/2143861251406875 +
                            skoXC1 *
                            (skoXC1 *
                             -4/
                             147926426347074375)))))))))))))))))))))))) <
     (1 +
      skoXC1 *
      (skoXC1 *
       (-1/2 +
        skoXC1 *
        (skoXC1 *
         (1/24 +
          skoXC1 *
          (skoXC1 *
           (-1/720 +
            skoXC1 *
            (skoXC1 *
             (1/40320 +
              skoXC1 *
              (skoXC1 *
               (-1/3628800 +
                skoXC1 *
                (skoXC1 *
                 (1/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-1/87178291200 +
                    skoXC1 *
                    (skoXC1 *
                     (1/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-1/6402373705728000 +
                        skoXC1 *
                        (skoXC1 *
                         (1/2432902008176640000 +
                          skoXC1 *
                          (skoXC1 *
                           (-1/1124000727777607680000 +
                            skoXC1 *
                            (skoXC1 *
                             1/
                             620448401733239439360000)))))))))))))))))))))))) *
     skoXC1), inference(resolve, [], [refute_0_192, refute_0_49])).

cnf(refute_0_194, plain,
    (1 <
     skoXC1 *
     (23/10 +
      skoXC1 *
      (skoXC1 *
       (-5/2 +
        skoXC1 *
        (skoXC1 *
         (17/24 +
          skoXC1 *
          (skoXC1 *
           (-13/144 +
            skoXC1 *
            (skoXC1 *
             (257/40320 +
              skoXC1 *
              (skoXC1 *
               (-41/145152 +
                skoXC1 *
                (skoXC1 *
                 (4097/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-3277/17435658240 +
                    skoXC1 *
                    (skoXC1 *
                     (65537/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-4033/98498057011200 +
                        skoXC1 *
                        (skoXC1 *
                         (61681/143111882833920000 +
                          skoXC1 *
                          (skoXC1 *
                           (-838861/224800145555521536000 +
                            skoXC1 *
                            (skoXC1 *
                             16777217/
                             620448401733239439360000))))))))))))))))))))))))),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (skoXC1 *
     (23/10 +
      skoXC1 *
      (skoXC1 *
       (-5/2 +
        skoXC1 *
        (skoXC1 *
         (17/24 +
          skoXC1 *
          (skoXC1 *
           (-13/144 +
            skoXC1 *
            (skoXC1 *
             (257/40320 +
              skoXC1 *
              (skoXC1 *
               (-41/145152 +
                skoXC1 *
                (skoXC1 *
                 (4097/479001600 +
                  skoXC1 *
                  (skoXC1 *
                   (-3277/17435658240 +
                    skoXC1 *
                    (skoXC1 *
                     (65537/20922789888000 +
                      skoXC1 *
                      (skoXC1 *
                       (-4033/98498057011200 +
                        skoXC1 *
                        (skoXC1 *
                         (61681/143111882833920000 +
                          skoXC1 *
                          (skoXC1 *
                           (-838861/224800145555521536000 +
                            skoXC1 *
                            (skoXC1 *
                             16777217/
                             620448401733239439360000)))))))))))))))))))))))) <=
     1), inference(decision, [], [refute_0_49, refute_0_50])).

cnf(refute_0_196, plain, ($false),
    inference(resolve, [], [refute_0_195, refute_0_194])).
% SZS output end CNFRefutation for trigpoly-352-2-weak.tptp
