% SZS output start CNFRefutation for Tunnel-3-VC-U.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X) < 9/10))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X) < 9/10)),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 817/10000000000) =>
          -65737067/400000000 * exp(-2559889987/10 * X) +
          403501489/1250000000 +
          5645414763/10000000000 * exp(-168440051/40 * X) < 9/10)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 817/10000000000 & 0 <= X &
        9/10 <=
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     9/10 <=
     -65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1) &
     skoXC1 <= 817/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (9/10 <=
     -65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 817/10000000000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) <
     -32402771200/627268307 +
     skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000))))) |
     420717228800/627268307 +
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= 0 |
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) |
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) < exp(skoXC1 * -2559889987/10)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <
     5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 |
     (5771988088/5645414763 +
      exp(skoXC1 * -2559889987/10) * 1643426675/5645414763) *
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (9/10 <=
     -65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <=
     exp(skoXC1 * -168440051/40)),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -168440051/40 |
     exp(skoXC1 * -168440051/40) <
     5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 |
     5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <=
     2304 /
     (2304 +
      skoXC1 * -168440051/40 *
      (-2304 +
       skoXC1 * -168440051/40 *
       (1152 +
        skoXC1 * -168440051/40 *
        (-384 +
         skoXC1 * -168440051/40 *
         (84 +
          skoXC1 * -168440051/40 * (-12 + skoXC1 * -168440051/40))))))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -168440051/40 |
     5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <=
     2304 /
     (2304 +
      skoXC1 * -168440051/40 *
      (-2304 +
       skoXC1 * -168440051/40 *
       (1152 +
        skoXC1 * -168440051/40 *
        (-384 +
         skoXC1 * -168440051/40 *
         (84 +
          skoXC1 * -168440051/40 * (-12 + skoXC1 * -168440051/40))))))),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <=
     2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000))))))), inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (5771988088/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <=
     2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((5771988088/5645414763 +
      exp(skoXC1 * -2559889987/10) * 1643426675/5645414763) *
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_5])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -2559889987/10) *
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) <=
     -32402771200/627268307 +
     skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000))))) |
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain, (skoXC1 <= 817/10000000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_20, plain,
    (-2304 <
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_21, plain,
    (exp(skoXC1 * -2559889987/10) *
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) <=
     -32402771200/627268307 +
     skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000)))))),
    inference(resolve, [], [refute_0_18, refute_0_20])).

cnf(refute_0_22, plain,
    (420717228800/627268307 +
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= 0 |
     exp(skoXC1 * -2559889987/10) <=
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000))))))), inference(split, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= -420717228800/627268307 |
     exp(skoXC1 * -2559889987/10) <=
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000))))))),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (-420717228800/627268307 <
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_25, plain,
    (exp(skoXC1 * -2559889987/10) <=
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000))))))),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (0 <
     (-32402771200/627268307 +
      skoXC1 *
      (-31111486973251599616/3136341535 +
       skoXC1 *
       (-327526278278770942196913776/15681707675 +
        skoXC1 *
        (-6896067877139546215120776059005322/235225615125 +
         skoXC1 *
         (-4065509087236964145071301716884524064329977/150544393680000 +
          skoXC1 *
          (-97827793999308241382997351404059542072352658101261/
           6021775747200000 +
           skoXC1 *
           -16478118610460974143272384403364710873703727420561966004311/
           2890452358656000000)))))) /
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000))))))),
    inference(resolve, [], [refute_0_25, refute_0_4])).

cnf(refute_0_27, plain,
    (0 *
     (420717228800/627268307 +
      skoXC1 *
      (1771640786891266720/627268307 +
       skoXC1 *
       (3730190806220563722142534/627268307 +
        skoXC1 *
        (314156764819761435303219128114617/37636098420 +
         skoXC1 *
         (370416070417649353792152012826611349278269/48174205977600 +
          skoXC1 *
          (8913271684624064064695304062895295689944349078817/
           1926968239104000 +
           skoXC1 *
           1501351937134933266884544315814590813674306345997738499667/
           924944754769920000)))))) <
     -32402771200/627268307 +
     skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000))))) |
     420717228800/627268307 +
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_26])).

cnf(refute_0_28, plain,
    (32402771200/627268307 <
     skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000))))) |
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= -420717228800/627268307),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (-31111486973251599616/3136341535 +
      skoXC1 *
      (-327526278278770942196913776/15681707675 +
       skoXC1 *
       (-6896067877139546215120776059005322/235225615125 +
        skoXC1 *
        (-4065509087236964145071301716884524064329977/150544393680000 +
         skoXC1 *
         (-97827793999308241382997351404059542072352658101261/
          6021775747200000 +
          skoXC1 *
          -16478118610460974143272384403364710873703727420561966004311/
          2890452358656000000))))) <= 32402771200/627268307 |
     skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= -420717228800/627268307),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_30, plain,
    (skoXC1 *
     (1771640786891266720/627268307 +
      skoXC1 *
      (3730190806220563722142534/627268307 +
       skoXC1 *
       (314156764819761435303219128114617/37636098420 +
        skoXC1 *
        (370416070417649353792152012826611349278269/48174205977600 +
         skoXC1 *
         (8913271684624064064695304062895295689944349078817/
          1926968239104000 +
          skoXC1 *
          1501351937134933266884544315814590813674306345997738499667/
          924944754769920000))))) <= -420717228800/627268307),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain, ($false),
    inference(resolve, [], [refute_0_30, refute_0_24])).
% SZS output end CNFRefutation for Tunnel-3-VC-U.tptp
