% SZS output start CNFRefutation for Tunnel-1-IL.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

fof('Tunnel', conjecture,
    (! [X] :
       (0 <= X =>
        3/500 + 3/2500 * exp(-257000000 * X) +
        43/50000 * exp(-54800000 * X) < 3/100))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        3/500 + 3/2500 * exp(-257000000 * X) +
        43/50000 * exp(-54800000 * X) < 3/100)),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          3/500 + 3/2500 * exp(-257000000 * X) +
          43/50000 * exp(-54800000 * X) < 3/100)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        3/100 <=
        3/500 + 3/2500 * exp(-257000000 * X) +
        43/50000 * exp(-54800000 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     3/100 <=
     3/500 + 3/2500 * exp(-257000000 * skoXC1) +
     43/50000 * exp(-54800000 * skoXC1)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (3/100 <=
     3/500 + 3/2500 * exp(-257000000 * skoXC1) +
     43/50000 * exp(-54800000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (2304 <
     (-6141837312/43 +
      skoXC1 *
      (-1927533625344000000/43 +
       skoXC1 *
       (-302109866754048000000000000/43 +
        skoXC1 *
        (-31536983905527398400000000000000000/43 +
         skoXC1 *
         (-2327305901146602577920000000000000000000000/43 +
          skoXC1 *
          (-124769650054290010103808000000000000000000000000000/43 +
           skoXC1 *
           (-4895523074351142388824883200000000000000000000000000000000/
            43 +
            skoXC1 *
            (-133782909008372663558956646400000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (-2598383715912803188474439270400000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (-36899812473883831714408110489600000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (-369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                (-2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                 43 +
                 skoXC1 *
                 -9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                 43)))))))))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) |
     0 <=
     -138240/43 +
     skoXC1 *
     (-7575552000000/43 +
      skoXC1 *
      (-207570124800000000000/43 +
       skoXC1 *
       (-3791614279680000000000000000/43 +
        skoXC1 *
        (-45451976177664000000000000000000000/43 +
         skoXC1 *
         (-355824042076569600000000000000000000000000/43 +
          skoXC1 *
          -1624929792149667840000000000000000000000000000000/43))))) |
     2304 *
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     -6141837312/43 +
     skoXC1 *
     (-1927533625344000000/43 +
      skoXC1 *
      (-302109866754048000000000000/43 +
       skoXC1 *
       (-31536983905527398400000000000000000/43 +
        skoXC1 *
        (-2327305901146602577920000000000000000000000/43 +
         skoXC1 *
         (-124769650054290010103808000000000000000000000000000/43 +
          skoXC1 *
          (-4895523074351142388824883200000000000000000000000000000000/43 +
           skoXC1 *
           (-133782909008372663558956646400000000000000000000000000000000000000/
            43 +
            skoXC1 *
            (-2598383715912803188474439270400000000000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (-36899812473883831714408110489600000000000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (-369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (-2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                -9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                43)))))))))))),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_1, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (592128000000 +
       skoXC1 *
       (76088448000000000000 +
        skoXC1 *
        (6518243712000000000000000000 +
         skoXC1 *
         (366447513684000000000000000000000000 +
          skoXC1 *
          (13453858716684000000000000000000000000000000 +
           skoXC1 *
           288136807515649000000000000000000000000000000000000)))))) <
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) *
     (2304 +
      skoXC1 *
      (592128000000 +
       skoXC1 *
       (76088448000000000000 +
        skoXC1 *
        (6518243712000000000000000000 +
         skoXC1 *
         (366447513684000000000000000000000000 +
          skoXC1 *
          (13453858716684000000000000000000000000000000 +
           skoXC1 *
           288136807515649000000000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (592128000000 +
      skoXC1 *
      (76088448000000000000 +
       skoXC1 *
       (6518243712000000000000000000 +
        skoXC1 *
        (366447513684000000000000000000000000 +
         skoXC1 *
         (13453858716684000000000000000000000000000000 +
          skoXC1 *
          288136807515649000000000000000000000000000000000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_2, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_3, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_4, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (0 < skoXC1 * -257000000 |
     exp(skoXC1 * -257000000) <
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304 /
     (2304 +
      skoXC1 * -257000000 *
      (-2304 +
       skoXC1 * -257000000 *
       (1152 +
        skoXC1 * -257000000 *
        (-384 +
         skoXC1 * -257000000 *
         (84 + skoXC1 * -257000000 * (-12 + skoXC1 * -257000000))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (126259200000 +
       skoXC1 *
       (3459502080000000000 +
        skoXC1 *
        (63193571328000000000000000 +
         skoXC1 *
         (757532936294400000000000000000000 +
          skoXC1 *
          (5930400701276160000000000000000000000000 +
           skoXC1 * 27082163202494464000000000000000000000000000000)))))) <
     1200/43 + exp(skoXC1 * -257000000) * -60/43 |
     (1200/43 + exp(skoXC1 * -257000000) * -60/43) *
     (2304 +
      skoXC1 *
      (126259200000 +
       skoXC1 *
       (3459502080000000000 +
        skoXC1 *
        (63193571328000000000000000 +
         skoXC1 *
         (757532936294400000000000000000000 +
          skoXC1 *
          (5930400701276160000000000000000000000000 +
           skoXC1 *
           27082163202494464000000000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (126259200000 +
      skoXC1 *
      (3459502080000000000 +
       skoXC1 *
       (63193571328000000000000000 +
        skoXC1 *
        (757532936294400000000000000000000 +
         skoXC1 *
         (5930400701276160000000000000000000000000 +
          skoXC1 * 27082163202494464000000000000000000000000000000))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_8, plain,
    (3/100 <=
     3/500 + 3/2500 * exp(-257000000 * skoXC1) +
     43/50000 * exp(-54800000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_9, plain,
    (1200/43 + exp(skoXC1 * -257000000) * -60/43 <=
     exp(skoXC1 * -54800000)), inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (0 < skoXC1 * -54800000 |
     exp(skoXC1 * -54800000) <
     1200/43 + exp(skoXC1 * -257000000) * -60/43 |
     1200/43 + exp(skoXC1 * -257000000) * -60/43 <=
     2304 /
     (2304 +
      skoXC1 * -54800000 *
      (-2304 +
       skoXC1 * -54800000 *
       (1152 +
        skoXC1 * -54800000 *
        (-384 +
         skoXC1 * -54800000 *
         (84 + skoXC1 * -54800000 * (-12 + skoXC1 * -54800000))))))),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_11, plain,
    (0 < skoXC1 * -54800000 |
     1200/43 + exp(skoXC1 * -257000000) * -60/43 <=
     2304 /
     (2304 +
      skoXC1 * -54800000 *
      (-2304 +
       skoXC1 * -54800000 *
       (1152 +
        skoXC1 * -54800000 *
        (-384 +
         skoXC1 * -54800000 *
         (84 + skoXC1 * -54800000 * (-12 + skoXC1 * -54800000))))))),
    inference(resolve, [], [refute_0_9, refute_0_10])).

cnf(refute_0_12, plain,
    (skoXC1 < 0 |
     1200/43 + exp(skoXC1 * -257000000) * -60/43 <=
     2304 /
     (2304 +
      skoXC1 *
      (126259200000 +
       skoXC1 *
       (3459502080000000000 +
        skoXC1 *
        (63193571328000000000000000 +
         skoXC1 *
         (757532936294400000000000000000000 +
          skoXC1 *
          (5930400701276160000000000000000000000000 +
           skoXC1 * 27082163202494464000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain,
    (1200/43 + exp(skoXC1 * -257000000) * -60/43 <=
     2304 /
     (2304 +
      skoXC1 *
      (126259200000 +
       skoXC1 *
       (3459502080000000000 +
        skoXC1 *
        (63193571328000000000000000 +
         skoXC1 *
         (757532936294400000000000000000000 +
          skoXC1 *
          (5930400701276160000000000000000000000000 +
           skoXC1 * 27082163202494464000000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_13, refute_0_12])).

cnf(refute_0_15, plain,
    ((1200/43 + exp(skoXC1 * -257000000) * -60/43) *
     (2304 +
      skoXC1 *
      (126259200000 +
       skoXC1 *
       (3459502080000000000 +
        skoXC1 *
        (63193571328000000000000000 +
         skoXC1 *
         (757532936294400000000000000000000 +
          skoXC1 *
          (5930400701276160000000000000000000000000 +
           skoXC1 *
           27082163202494464000000000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (126259200000 +
      skoXC1 *
      (3459502080000000000 +
       skoXC1 *
       (63193571328000000000000000 +
        skoXC1 *
        (757532936294400000000000000000000 +
         skoXC1 *
         (5930400701276160000000000000000000000000 +
          skoXC1 * 27082163202494464000000000000000000000000000000))))) <=
     0), inference(resolve, [], [refute_0_14, refute_0_7])).

cnf(refute_0_16, plain,
    (exp(skoXC1 * -257000000) *
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     -2665728/43 +
     skoXC1 *
     (-151511040000000/43 +
      skoXC1 *
      (-4151402496000000000000/43 +
       skoXC1 *
       (-75832285593600000000000000000/43 +
        skoXC1 *
        (-909039523553280000000000000000000000/43 +
         skoXC1 *
         (-7116480841531392000000000000000000000000000/43 +
          skoXC1 *
          -32498595842993356800000000000000000000000000000000/43))))) |
     skoXC1 *
     (126259200000 +
      skoXC1 *
      (3459502080000000000 +
       skoXC1 *
       (63193571328000000000000000 +
        skoXC1 *
        (757532936294400000000000000000000 +
         skoXC1 *
         (5930400701276160000000000000000000000000 +
          skoXC1 * 27082163202494464000000000000000000000000000000))))) <=
     -2304), inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (-2304 <
     skoXC1 *
     (126259200000 +
      skoXC1 *
      (3459502080000000000 +
       skoXC1 *
       (63193571328000000000000000 +
        skoXC1 *
        (757532936294400000000000000000000 +
         skoXC1 *
         (5930400701276160000000000000000000000000 +
          skoXC1 * 27082163202494464000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -257000000) *
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     -2665728/43 +
     skoXC1 *
     (-151511040000000/43 +
      skoXC1 *
      (-4151402496000000000000/43 +
       skoXC1 *
       (-75832285593600000000000000000/43 +
        skoXC1 *
        (-909039523553280000000000000000000000/43 +
         skoXC1 *
         (-7116480841531392000000000000000000000000000/43 +
          skoXC1 *
          -32498595842993356800000000000000000000000000000000/43)))))),
    inference(resolve, [], [refute_0_16, refute_0_17])).

cnf(refute_0_19, plain,
    (0 <=
     -138240/43 +
     skoXC1 *
     (-7575552000000/43 +
      skoXC1 *
      (-207570124800000000000/43 +
       skoXC1 *
       (-3791614279680000000000000000/43 +
        skoXC1 *
        (-45451976177664000000000000000000000/43 +
         skoXC1 *
         (-355824042076569600000000000000000000000000/43 +
          skoXC1 *
          -1624929792149667840000000000000000000000000000000/43))))) |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     exp(skoXC1 * -257000000)), inference(split, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (skoXC1 *
     (7575552000000/43 +
      skoXC1 *
      (207570124800000000000/43 +
       skoXC1 *
       (3791614279680000000000000000/43 +
        skoXC1 *
        (45451976177664000000000000000000000/43 +
         skoXC1 *
         (355824042076569600000000000000000000000000/43 +
          skoXC1 *
          1624929792149667840000000000000000000000000000000/43))))) <=
     -138240/43 |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     exp(skoXC1 * -257000000)), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (-138240/43 <
     skoXC1 *
     (7575552000000/43 +
      skoXC1 *
      (207570124800000000000/43 +
       skoXC1 *
       (3791614279680000000000000000/43 +
        skoXC1 *
        (45451976177664000000000000000000000/43 +
         skoXC1 *
         (355824042076569600000000000000000000000000/43 +
          skoXC1 *
          1624929792149667840000000000000000000000000000000/43)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_22, plain,
    ((-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     exp(skoXC1 * -257000000)),
    inference(resolve, [], [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (0 < skoXC1 * -257000000 |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304 /
     (2304 +
      skoXC1 * -257000000 *
      (-2304 +
       skoXC1 * -257000000 *
       (1152 +
        skoXC1 * -257000000 *
        (-384 +
         skoXC1 * -257000000 *
         (84 + skoXC1 * -257000000 * (-12 + skoXC1 * -257000000))))))),
    inference(resolve, [], [refute_0_22, refute_0_6])).

cnf(refute_0_24, plain,
    (skoXC1 < 0 |
     (-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (592128000000 +
       skoXC1 *
       (76088448000000000000 +
        skoXC1 *
        (6518243712000000000000000000 +
         skoXC1 *
         (366447513684000000000000000000000000 +
          skoXC1 *
          (13453858716684000000000000000000000000000000 +
           skoXC1 *
           288136807515649000000000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    ((-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304 /
     (2304 +
      skoXC1 *
      (592128000000 +
       skoXC1 *
       (76088448000000000000 +
        skoXC1 *
        (6518243712000000000000000000 +
         skoXC1 *
         (366447513684000000000000000000000000 +
          skoXC1 *
          (13453858716684000000000000000000000000000000 +
           skoXC1 *
           288136807515649000000000000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_13, refute_0_24])).

cnf(refute_0_26, plain,
    ((-2665728/43 +
      skoXC1 *
      (-151511040000000/43 +
       skoXC1 *
       (-4151402496000000000000/43 +
        skoXC1 *
        (-75832285593600000000000000000/43 +
         skoXC1 *
         (-909039523553280000000000000000000000/43 +
          skoXC1 *
          (-7116480841531392000000000000000000000000000/43 +
           skoXC1 *
           -32498595842993356800000000000000000000000000000000/43)))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) *
     (2304 +
      skoXC1 *
      (592128000000 +
       skoXC1 *
       (76088448000000000000 +
        skoXC1 *
        (6518243712000000000000000000 +
         skoXC1 *
         (366447513684000000000000000000000000 +
          skoXC1 *
          (13453858716684000000000000000000000000000000 +
           skoXC1 *
           288136807515649000000000000000000000000000000000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (592128000000 +
      skoXC1 *
      (76088448000000000000 +
       skoXC1 *
       (6518243712000000000000000000 +
        skoXC1 *
        (366447513684000000000000000000000000 +
         skoXC1 *
         (13453858716684000000000000000000000000000000 +
          skoXC1 *
          288136807515649000000000000000000000000000000000000))))) <= 0),
    inference(resolve, [], [refute_0_25, refute_0_1])).

cnf(refute_0_27, plain,
    (skoXC1 *
     (592128000000 +
      skoXC1 *
      (76088448000000000000 +
       skoXC1 *
       (6518243712000000000000000000 +
        skoXC1 *
        (366447513684000000000000000000000000 +
         skoXC1 *
         (13453858716684000000000000000000000000000000 +
          skoXC1 *
          288136807515649000000000000000000000000000000000000))))) <=
     -2304 |
     (-6141837312/43 +
      skoXC1 *
      (-1927533625344000000/43 +
       skoXC1 *
       (-302109866754048000000000000/43 +
        skoXC1 *
        (-31536983905527398400000000000000000/43 +
         skoXC1 *
         (-2327305901146602577920000000000000000000000/43 +
          skoXC1 *
          (-124769650054290010103808000000000000000000000000000/43 +
           skoXC1 *
           (-4895523074351142388824883200000000000000000000000000000000/
            43 +
            skoXC1 *
            (-133782909008372663558956646400000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (-2598383715912803188474439270400000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (-36899812473883831714408110489600000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (-369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                (-2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                 43 +
                 skoXC1 *
                 -9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                 43)))))))))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (-2304 <
     skoXC1 *
     (592128000000 +
      skoXC1 *
      (76088448000000000000 +
       skoXC1 *
       (6518243712000000000000000000 +
        skoXC1 *
        (366447513684000000000000000000000000 +
         skoXC1 *
         (13453858716684000000000000000000000000000000 +
          skoXC1 *
          288136807515649000000000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_13])).

cnf(refute_0_29, plain,
    ((-6141837312/43 +
      skoXC1 *
      (-1927533625344000000/43 +
       skoXC1 *
       (-302109866754048000000000000/43 +
        skoXC1 *
        (-31536983905527398400000000000000000/43 +
         skoXC1 *
         (-2327305901146602577920000000000000000000000/43 +
          skoXC1 *
          (-124769650054290010103808000000000000000000000000000/43 +
           skoXC1 *
           (-4895523074351142388824883200000000000000000000000000000000/
            43 +
            skoXC1 *
            (-133782909008372663558956646400000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (-2598383715912803188474439270400000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (-36899812473883831714408110489600000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (-369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                (-2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                 43 +
                 skoXC1 *
                 -9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                 43)))))))))))) /
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     2304), inference(resolve, [], [refute_0_27, refute_0_28])).

cnf(refute_0_30, plain,
    (0 <=
     -138240/43 +
     skoXC1 *
     (-7575552000000/43 +
      skoXC1 *
      (-207570124800000000000/43 +
       skoXC1 *
       (-3791614279680000000000000000/43 +
        skoXC1 *
        (-45451976177664000000000000000000000/43 +
         skoXC1 *
         (-355824042076569600000000000000000000000000/43 +
          skoXC1 *
          -1624929792149667840000000000000000000000000000000/43))))) |
     2304 *
     (-138240/43 +
      skoXC1 *
      (-7575552000000/43 +
       skoXC1 *
       (-207570124800000000000/43 +
        skoXC1 *
        (-3791614279680000000000000000/43 +
         skoXC1 *
         (-45451976177664000000000000000000000/43 +
          skoXC1 *
          (-355824042076569600000000000000000000000000/43 +
           skoXC1 *
           -1624929792149667840000000000000000000000000000000/43)))))) <=
     -6141837312/43 +
     skoXC1 *
     (-1927533625344000000/43 +
      skoXC1 *
      (-302109866754048000000000000/43 +
       skoXC1 *
       (-31536983905527398400000000000000000/43 +
        skoXC1 *
        (-2327305901146602577920000000000000000000000/43 +
         skoXC1 *
         (-124769650054290010103808000000000000000000000000000/43 +
          skoXC1 *
          (-4895523074351142388824883200000000000000000000000000000000/43 +
           skoXC1 *
           (-133782909008372663558956646400000000000000000000000000000000000000/
            43 +
            skoXC1 *
            (-2598383715912803188474439270400000000000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (-36899812473883831714408110489600000000000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (-369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (-2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                -9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                43)))))))))))),
    inference(resolve, [], [refute_0_29, refute_0_0])).

cnf(refute_0_31, plain,
    (skoXC1 *
     (7575552000000/43 +
      skoXC1 *
      (207570124800000000000/43 +
       skoXC1 *
       (3791614279680000000000000000/43 +
        skoXC1 *
        (45451976177664000000000000000000000/43 +
         skoXC1 *
         (355824042076569600000000000000000000000000/43 +
          skoXC1 *
          1624929792149667840000000000000000000000000000000/43))))) <=
     -138240/43 |
     skoXC1 *
     (1910079553536000000/43 +
      skoXC1 *
      (301631625186508800000000000/43 +
       skoXC1 *
       (31528248026227015680000000000000000/43 +
        skoXC1 *
        (2327201179793489240064000000000000000000000/43 +
         skoXC1 *
         (124768830235697065687449600000000000000000000000000/43 +
          skoXC1 *
          (4895519330512901275990179840000000000000000000000000000000/43 +
           skoXC1 *
           (133782909008372663558956646400000000000000000000000000000000000000/
            43 +
            skoXC1 *
            (2598383715912803188474439270400000000000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (36899812473883831714408110489600000000000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                43))))))))))) <= -5823332352/43),
    inference(arithmetic, [], [refute_0_30])).

cnf(refute_0_32, plain,
    (-5823332352/43 <
     skoXC1 *
     (1910079553536000000/43 +
      skoXC1 *
      (301631625186508800000000000/43 +
       skoXC1 *
       (31528248026227015680000000000000000/43 +
        skoXC1 *
        (2327201179793489240064000000000000000000000/43 +
         skoXC1 *
         (124768830235697065687449600000000000000000000000000/43 +
          skoXC1 *
          (4895519330512901275990179840000000000000000000000000000000/43 +
           skoXC1 *
           (133782909008372663558956646400000000000000000000000000000000000000/
            43 +
            skoXC1 *
            (2598383715912803188474439270400000000000000000000000000000000000000000000/
             43 +
             skoXC1 *
             (36899812473883831714408110489600000000000000000000000000000000000000000000000000/
              43 +
              skoXC1 *
              (369580903669026614431594517299200000000000000000000000000000000000000000000000000000000/
               43 +
               skoXC1 *
               (2487751587387381091070550078259200000000000000000000000000000000000000000000000000000000000000/
                43 +
                skoXC1 *
                9364041654941447598407419040563200000000000000000000000000000000000000000000000000000000000000000000/
                43))))))))))) |
     skoXC1 *
     (7575552000000/43 +
      skoXC1 *
      (207570124800000000000/43 +
       skoXC1 *
       (3791614279680000000000000000/43 +
        skoXC1 *
        (45451976177664000000000000000000000/43 +
         skoXC1 *
         (355824042076569600000000000000000000000000/43 +
          skoXC1 *
          1624929792149667840000000000000000000000000000000/43))))) <=
     -138240/43), inference(decision, [], [refute_0_13])).

cnf(refute_0_33, plain,
    (skoXC1 *
     (7575552000000/43 +
      skoXC1 *
      (207570124800000000000/43 +
       skoXC1 *
       (3791614279680000000000000000/43 +
        skoXC1 *
        (45451976177664000000000000000000000/43 +
         skoXC1 *
         (355824042076569600000000000000000000000000/43 +
          skoXC1 *
          1624929792149667840000000000000000000000000000000/43))))) <=
     -138240/43), inference(resolve, [], [refute_0_31, refute_0_32])).

cnf(refute_0_34, plain, ($false),
    inference(resolve, [], [refute_0_33, refute_0_21])).
% SZS output end CNFRefutation for Tunnel-1-IL.tptp
