% SZS output start CNFRefutation for Tunnel-2-VC-L.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        0 <
        695936057/2500000000 +
        2472636127/1000000000000 * exp(8785654241/100 * X) -
        4454058919/1000000000000 * exp(1104345759/100 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        0 <
        695936057/2500000000 +
        2472636127/1000000000000 * exp(8785654241/100 * X) -
        4454058919/1000000000000 * exp(1104345759/100 * X))),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 593/10000000000) =>
          0 <
          695936057/2500000000 +
          2472636127/1000000000000 * exp(8785654241/100 * X) -
          4454058919/1000000000000 * exp(1104345759/100 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 593/10000000000 & 0 <= X &
        695936057/2500000000 +
        2472636127/1000000000000 * exp(8785654241/100 * X) -
        4454058919/1000000000000 * exp(1104345759/100 * X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     695936057/2500000000 +
     2472636127/1000000000000 * exp(8785654241/100 * skoXC1) -
     4454058919/1000000000000 * exp(1104345759/100 * skoXC1) <= 0 &
     skoXC1 <= 593/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (695936057/2500000000 +
     2472636127/1000000000000 * exp(8785654241/100 * skoXC1) -
     4454058919/1000000000000 * exp(1104345759/100 * skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 593/10000000000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (278374422800/4454058919 *
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) < 21743271936 |
     21743271936 +
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) <=
     278374422800/4454058919),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (695936057/2500000000 +
     2472636127/1000000000000 * exp(8785654241/100 * skoXC1) -
     4454058919/1000000000000 * exp(1104345759/100 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain,
    (exp(skoXC1 * 8785654241/100) <=
     -278374422800/2472636127 +
     exp(skoXC1 * 1104345759/100) * 4454058919/2472636127),
    inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_4, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_5, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (0 <
     -278374422800/2472636127 +
     exp(skoXC1 * 1104345759/100) * 4454058919/2472636127 |
     -278374422800/2472636127 +
     exp(skoXC1 * 1104345759/100) * 4454058919/2472636127 <
     exp(skoXC1 * 8785654241/100)), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (0 <
     -278374422800/2472636127 +
     exp(skoXC1 * 1104345759/100) * 4454058919/2472636127),
    inference(resolve, [], [refute_0_2, refute_0_6])).

cnf(refute_0_8, plain,
    (278374422800/4454058919 < exp(skoXC1 * 1104345759/100)),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, exp(X_000061), X_000062) | exp(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(0, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_11, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | exp(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (3 < X_000061 |
     X_000062 <
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061)))))))))))) |
     exp(X_000061) <= X_000062), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (3 < skoXC1 * 1104345759/100 |
     278374422800/4454058919 <
     21743271936 /
     (21743271936 +
      skoXC1 * 1104345759/100 *
      (-21743271936 +
       skoXC1 * 1104345759/100 *
       (10871635968 +
        skoXC1 * 1104345759/100 *
        (-3623878656 +
         skoXC1 * 1104345759/100 *
         (891813888 +
          skoXC1 * 1104345759/100 *
          (-169869312 +
           skoXC1 * 1104345759/100 *
           (25657344 +
            skoXC1 * 1104345759/100 *
            (-3096576 +
             skoXC1 * 1104345759/100 *
             (297216 +
              skoXC1 * 1104345759/100 *
              (-22272 +
               skoXC1 * 1104345759/100 *
               (1248 +
                skoXC1 * 1104345759/100 *
                (-48 + skoXC1 * 1104345759/100)))))))))))) |
     exp(skoXC1 * 1104345759/100) <= 278374422800/4454058919),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (3 < skoXC1 * 1104345759/100 |
     278374422800/4454058919 <
     21743271936 /
     (21743271936 +
      skoXC1 * 1104345759/100 *
      (-21743271936 +
       skoXC1 * 1104345759/100 *
       (10871635968 +
        skoXC1 * 1104345759/100 *
        (-3623878656 +
         skoXC1 * 1104345759/100 *
         (891813888 +
          skoXC1 * 1104345759/100 *
          (-169869312 +
           skoXC1 * 1104345759/100 *
           (25657344 +
            skoXC1 * 1104345759/100 *
            (-3096576 +
             skoXC1 * 1104345759/100 *
             (297216 +
              skoXC1 * 1104345759/100 *
              (-22272 +
               skoXC1 * 1104345759/100 *
               (1248 +
                skoXC1 * 1104345759/100 *
                (-48 + skoXC1 * 1104345759/100))))))))))))),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    (278374422800/4454058919 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) |
     100/368115253 < skoXC1), inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (278374422800/4454058919 *
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) <= 278374422800/4454058919),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_17, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_18, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_19, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (0 < skoXC1 * 1104345759/100 |
     278374422800/4454058919 <
     2304 /
     (2304 +
      skoXC1 * 1104345759/100 *
      (-2304 +
       skoXC1 * 1104345759/100 *
       (1152 +
        skoXC1 * 1104345759/100 *
        (-384 +
         skoXC1 * 1104345759/100 *
         (84 +
          skoXC1 * 1104345759/100 * (-12 + skoXC1 * 1104345759/100)))))) |
     exp(skoXC1 * 1104345759/100) <= 278374422800/4454058919),
    inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (0 < skoXC1 * 1104345759/100 |
     278374422800/4454058919 <
     2304 /
     (2304 +
      skoXC1 * 1104345759/100 *
      (-2304 +
       skoXC1 * 1104345759/100 *
       (1152 +
        skoXC1 * 1104345759/100 *
        (-384 +
         skoXC1 * 1104345759/100 *
         (84 +
          skoXC1 * 1104345759/100 * (-12 + skoXC1 * 1104345759/100))))))),
    inference(resolve, [], [refute_0_21, refute_0_8])).

cnf(refute_0_23, plain,
    (278374422800/4454058919 <
     2304 /
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) | 0 < skoXC1),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (278374422800/4454058919 *
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) < 2304 | 0 < skoXC1 |
     2304 +
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_23])).

cnf(refute_0_25, plain,
    (631112518381824/4454058919 <
     skoXC1 *
     (7082993968894146935808/4454058919 +
      skoXC1 *
      (-977759293821353636060301254784/111351472975 +
       skoXC1 *
       (89982027454537232635201929915257405088/2783786824375 +
        skoXC1 *
        (-21737465401233697099311609933553381653623273517/278378682437500 +
         skoXC1 *
         (3429382532465952394216482607083135547041038441742280629/
          27837868243750000 +
          skoXC1 *
          -1262408018572484779482956231676508033932305267364685394541334137/
          11135147297500000000))))) | 0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 <= 593/10000000000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_28, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304 |
     skoXC1 *
     (7082993968894146935808/4454058919 +
      skoXC1 *
      (-977759293821353636060301254784/111351472975 +
       skoXC1 *
       (89982027454537232635201929915257405088/2783786824375 +
        skoXC1 *
        (-21737465401233697099311609933553381653623273517/278378682437500 +
         skoXC1 *
         (3429382532465952394216482607083135547041038441742280629/
          27837868243750000 +
          skoXC1 *
          -1262408018572484779482956231676508033932305267364685394541334137/
          11135147297500000000))))) <= 631112518381824/4454058919),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_28, refute_0_25])).

cnf(refute_0_30, plain,
    (-2304 <
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) | 0 < skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_31, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain, (skoXC1 <= 100/368115253),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_33, plain,
    (278374422800/4454058919 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_32, refute_0_15])).

cnf(refute_0_34, plain,
    (278374422800/4454058919 *
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) < 21743271936 |
     21743271936 +
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_33])).

cnf(refute_0_35, plain,
    (5955924960672655343616/4454058919 <
     skoXC1 *
     (66843517355344341156256284672/4454058919 +
      skoXC1 *
      (-9227294363502177392570098036827488256/111351472975 +
       skoXC1 *
       (849176949781519499229205489765388539177992192/2783786824375 +
        skoXC1 *
        (-57695754567677688710569669578516854012100121555009536/
         69594670609375 +
         skoXC1 *
         (3034098184243796857449513955860305011346900186793298394578944/
          1739866765234375 +
          skoXC1 *
          (-126523581253408714538899805363546341192831363116358228982510672546688/
           43496669130859375 +
           skoXC1 *
           (4215864062908865204149696051957349076180569911285235499616835451699173948592/
            1087416728271484375 +
            skoXC1 *
            (-1787483738669092170922603567438215270982062079360953667166769904568283750348454879617/
             434966691308593750000 +
             skoXC1 *
             (147922487070644989810893877625894683054532828553211182851169190711436040523441851774493011201/
              43496669130859375000000 +
              skoXC1 *
              (-73229256080813402158598801890271833226034221869212668378269272806802606688744722005982198511535517423/
               34797335304687500000000000 +
               skoXC1 *
               (6220801414428557072631540942307914029255860100779235614816544552631268080527712829303108735054643172612358389/
                6959467060937500000000000000 +
                skoXC1 *
                -2289971886535126137216699069757536176781770342730953815494802846577631065174283381670259685691149940310920979293407417/
                11135147297500000000000000000000))))))))))) |
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936 |
     skoXC1 *
     (66843517355344341156256284672/4454058919 +
      skoXC1 *
      (-9227294363502177392570098036827488256/111351472975 +
       skoXC1 *
       (849176949781519499229205489765388539177992192/2783786824375 +
        skoXC1 *
        (-57695754567677688710569669578516854012100121555009536/
         69594670609375 +
         skoXC1 *
         (3034098184243796857449513955860305011346900186793298394578944/
          1739866765234375 +
          skoXC1 *
          (-126523581253408714538899805363546341192831363116358228982510672546688/
           43496669130859375 +
           skoXC1 *
           (4215864062908865204149696051957349076180569911285235499616835451699173948592/
            1087416728271484375 +
            skoXC1 *
            (-1787483738669092170922603567438215270982062079360953667166769904568283750348454879617/
             434966691308593750000 +
             skoXC1 *
             (147922487070644989810893877625894683054532828553211182851169190711436040523441851774493011201/
              43496669130859375000000 +
              skoXC1 *
              (-73229256080813402158598801890271833226034221869212668378269272806802606688744722005982198511535517423/
               34797335304687500000000000 +
               skoXC1 *
               (6220801414428557072631540942307914029255860100779235614816544552631268080527712829303108735054643172612358389/
                6959467060937500000000000000 +
                skoXC1 *
                -2289971886535126137216699069757536176781770342730953815494802846577631065174283381670259685691149940310920979293407417/
                11135147297500000000000000000000))))))))))) <=
     5955924960672655343616/4454058919),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_37, plain,
    (skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (-21743271936 <
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000)))))))))))),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_39, plain, ($false),
    inference(resolve, [], [refute_0_37, refute_0_38])).
% SZS output end CNFRefutation for Tunnel-2-VC-L.tptp
