% SZS output start CNFRefutation for Tunnel-3-IL-U.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        -1606860999/2500000000000 * exp(-2559889987/10 * X) -
        7600397043/100000000000 +
        1443470449/10000000000 * exp(-168440051/40 * X) < 2/25))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        -1606860999/2500000000000 * exp(-2559889987/10 * X) -
        7600397043/100000000000 +
        1443470449/10000000000 * exp(-168440051/40 * X) < 2/25)),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 817/10000000000) =>
          -1606860999/2500000000000 * exp(-2559889987/10 * X) -
          7600397043/100000000000 +
          1443470449/10000000000 * exp(-168440051/40 * X) < 2/25)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 817/10000000000 & 0 <= X &
        2/25 <=
        -1606860999/2500000000000 * exp(-2559889987/10 * X) -
        7600397043/100000000000 +
        1443470449/10000000000 * exp(-168440051/40 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     2/25 <=
     -1606860999/2500000000000 * exp(-2559889987/10 * skoXC1) -
     7600397043/100000000000 +
     1443470449/10000000000 * exp(-168440051/40 * skoXC1) &
     skoXC1 <= 817/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (2/25 <=
     -1606860999/2500000000000 * exp(-2559889987/10 * skoXC1) -
     7600397043/100000000000 +
     1443470449/10000000000 * exp(-168440051/40 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 817/10000000000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) <
     -1342877821056/7217352245 +
     skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000))))) |
     1851103870848/180433806125 +
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= 0 |
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) |
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) < exp(skoXC1 * -2559889987/10)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <
     15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 |
     (15600397043/14434704490 +
      exp(skoXC1 * -2559889987/10) * 1606860999/360867612250) *
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (2/25 <=
     -1606860999/2500000000000 * exp(-2559889987/10 * skoXC1) -
     7600397043/100000000000 +
     1443470449/10000000000 * exp(-168440051/40 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 <=
     exp(skoXC1 * -168440051/40)),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -168440051/40 |
     exp(skoXC1 * -168440051/40) <
     15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 |
     15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 <=
     2304 /
     (2304 +
      skoXC1 * -168440051/40 *
      (-2304 +
       skoXC1 * -168440051/40 *
       (1152 +
        skoXC1 * -168440051/40 *
        (-384 +
         skoXC1 * -168440051/40 *
         (84 +
          skoXC1 * -168440051/40 * (-12 + skoXC1 * -168440051/40))))))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -168440051/40 |
     15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 <=
     2304 /
     (2304 +
      skoXC1 * -168440051/40 *
      (-2304 +
       skoXC1 * -168440051/40 *
       (1152 +
        skoXC1 * -168440051/40 *
        (-384 +
         skoXC1 * -168440051/40 *
         (84 +
          skoXC1 * -168440051/40 * (-12 + skoXC1 * -168440051/40))))))),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 <=
     2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000))))))), inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (15600397043/14434704490 +
     exp(skoXC1 * -2559889987/10) * 1606860999/360867612250 <=
     2304 /
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((15600397043/14434704490 +
      exp(skoXC1 * -2559889987/10) * 1606860999/360867612250) *
     (2304 +
      skoXC1 *
      (48510734688/5 +
       skoXC1 *
       (510696914055886818/25 +
        skoXC1 *
        (14336969041519365412357953/500 +
         skoXC1 *
         (16904438575772601192816467435029221/640000 +
          skoXC1 *
          (406769213689929187045809515485165881675753/25600000 +
           skoXC1 *
           22838742366387190150794898041535543617641266927801/
           4096000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_5])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -2559889987/10) *
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) <=
     -1342877821056/7217352245 +
     skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000))))) |
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain, (skoXC1 <= 817/10000000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_20, plain,
    (-2304 <
     skoXC1 *
     (48510734688/5 +
      skoXC1 *
      (510696914055886818/25 +
       skoXC1 *
       (14336969041519365412357953/500 +
        skoXC1 *
        (16904438575772601192816467435029221/640000 +
         skoXC1 *
         (406769213689929187045809515485165881675753/25600000 +
          skoXC1 *
          22838742366387190150794898041535543617641266927801/
          4096000000)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_21, plain,
    (exp(skoXC1 * -2559889987/10) *
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) <=
     -1342877821056/7217352245 +
     skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000)))))),
    inference(resolve, [], [refute_0_18, refute_0_20])).

cnf(refute_0_22, plain,
    (1851103870848/180433806125 +
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= 0 |
     exp(skoXC1 * -2559889987/10) <=
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000))))))),
    inference(split, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= -1851103870848/180433806125 |
     exp(skoXC1 * -2559889987/10) <=
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000))))))),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (-1851103870848/180433806125 <
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_25, plain,
    (exp(skoXC1 * -2559889987/10) <=
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000))))))),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (0 <
     (-1342877821056/7217352245 +
      skoXC1 *
      (-378393360990216363792/36086761225 +
       skoXC1 *
       (-3983537313953340926134939587/180433806125 +
        skoXC1 *
        (-223662409440901252406185485638732979/7217352245000 +
         skoXC1 *
         (-263715953571058019088832291415135653906993503/9238210873600000 +
          skoXC1 *
          (-6345761238431806408468840670916044530858904985998379/
           369528434944000000 +
           skoXC1 *
           -356293448878425543821539451506657585832048343215240414892443/
           59124549591040000000)))))) /
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000))))))),
    inference(resolve, [], [refute_0_25, refute_0_4])).

cnf(refute_0_27, plain,
    (0 *
     (1851103870848/180433806125 +
      skoXC1 *
      (38975003801491816656/902169030625 +
       skoXC1 *
       (410309476753029717101205591/4510845153125 +
        skoXC1 *
        (23037516396687879984347547303175047/180433806125000 +
         skoXC1 *
         (27163083057400099149517660486302021650251779/230955271840000000 +
          skoXC1 *
          (653621585072244089735687336816199618310216259657247/
           9238210873600000000 +
           skoXC1 *
           36698684374756544386509250511124947111451120199231975733199/
           1478113739776000000000)))))) <
     -1342877821056/7217352245 +
     skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000))))) |
     1851103870848/180433806125 +
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_26])).

cnf(refute_0_28, plain,
    (1342877821056/7217352245 <
     skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000))))) |
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= -1851103870848/180433806125),
    inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (-378393360990216363792/36086761225 +
      skoXC1 *
      (-3983537313953340926134939587/180433806125 +
       skoXC1 *
       (-223662409440901252406185485638732979/7217352245000 +
        skoXC1 *
        (-263715953571058019088832291415135653906993503/9238210873600000 +
         skoXC1 *
         (-6345761238431806408468840670916044530858904985998379/
          369528434944000000 +
          skoXC1 *
          -356293448878425543821539451506657585832048343215240414892443/
          59124549591040000000))))) <= 1342877821056/7217352245 |
     skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= -1851103870848/180433806125),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_30, plain,
    (skoXC1 *
     (38975003801491816656/902169030625 +
      skoXC1 *
      (410309476753029717101205591/4510845153125 +
       skoXC1 *
       (23037516396687879984347547303175047/180433806125000 +
        skoXC1 *
        (27163083057400099149517660486302021650251779/230955271840000000 +
         skoXC1 *
         (653621585072244089735687336816199618310216259657247/
          9238210873600000000 +
          skoXC1 *
          36698684374756544386509250511124947111451120199231975733199/
          1478113739776000000000))))) <= -1851103870848/180433806125),
    inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain, ($false),
    inference(resolve, [], [refute_0_30, refute_0_24])).
% SZS output end CNFRefutation for Tunnel-3-IL-U.tptp
