% SZS output start CNFRefutation for Tunnel-1-IL-L.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       (0 <= X =>
        0 <
        1822450701/25000000000 +
        449447377/1000000000000 * exp(-258222872 * X) -
        3392373771/50000000000 * exp(-1044282011/250 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        0 <
        1822450701/25000000000 +
        449447377/1000000000000 * exp(-258222872 * X) -
        3392373771/50000000000 * exp(-1044282011/250 * X))),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          0 <
          1822450701/25000000000 +
          449447377/1000000000000 * exp(-258222872 * X) -
          3392373771/50000000000 * exp(-1044282011/250 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        1822450701/25000000000 +
        449447377/1000000000000 * exp(-258222872 * X) -
        3392373771/50000000000 * exp(-1044282011/250 * X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     1822450701/25000000000 +
     449447377/1000000000000 * exp(-258222872 * skoXC1) -
     3392373771/50000000000 * exp(-1044282011/250 * skoXC1) <= 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (1822450701/25000000000 +
     449447377/1000000000000 * exp(-258222872 * skoXC1) -
     3392373771/50000000000 * exp(-1044282011/250 * skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) <
     -21549024512/125643473 +
     skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000))))) |
     28764632128/1884652095 +
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= 0 |
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) |
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) < exp(skoXC1 * -258222872)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (1203012876672/125 +
       skoXC1 *
       (314071176527482786848/15625 +
        skoXC1 *
        (54663146636875953569585631888/1953125 +
         skoXC1 *
         (24974136555175809553645411545809035461/976562500 +
          skoXC1 *
          (3725720220689658122747691821425725453311913153/244140625000 +
           skoXC1 *
           1296900868161719997108481486962236487839483759224063561/
           244140625000000)))))) <
     1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 |
     (1214967134/1130791257 +
      exp(skoXC1 * -258222872) * 449447377/67847475420) *
     (2304 +
      skoXC1 *
      (1203012876672/125 +
       skoXC1 *
       (314071176527482786848/15625 +
        skoXC1 *
        (54663146636875953569585631888/1953125 +
         skoXC1 *
         (24974136555175809553645411545809035461/976562500 +
          skoXC1 *
          (3725720220689658122747691821425725453311913153/244140625000 +
           skoXC1 *
           1296900868161719997108481486962236487839483759224063561/
           244140625000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (1203012876672/125 +
      skoXC1 *
      (314071176527482786848/15625 +
       skoXC1 *
       (54663146636875953569585631888/1953125 +
        skoXC1 *
        (24974136555175809553645411545809035461/976562500 +
         skoXC1 *
         (3725720220689658122747691821425725453311913153/244140625000 +
          skoXC1 *
          1296900868161719997108481486962236487839483759224063561/
          244140625000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (1822450701/25000000000 +
     449447377/1000000000000 * exp(-258222872 * skoXC1) -
     3392373771/50000000000 * exp(-1044282011/250 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 <=
     exp(skoXC1 * -1044282011/250)),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -1044282011/250 |
     exp(skoXC1 * -1044282011/250) <
     1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 |
     1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 <=
     2304 /
     (2304 +
      skoXC1 * -1044282011/250 *
      (-2304 +
       skoXC1 * -1044282011/250 *
       (1152 +
        skoXC1 * -1044282011/250 *
        (-384 +
         skoXC1 * -1044282011/250 *
         (84 +
          skoXC1 * -1044282011/250 *
          (-12 + skoXC1 * -1044282011/250))))))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -1044282011/250 |
     1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 <=
     2304 /
     (2304 +
      skoXC1 * -1044282011/250 *
      (-2304 +
       skoXC1 * -1044282011/250 *
       (1152 +
        skoXC1 * -1044282011/250 *
        (-384 +
         skoXC1 * -1044282011/250 *
         (84 +
          skoXC1 * -1044282011/250 *
          (-12 + skoXC1 * -1044282011/250))))))),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 <=
     2304 /
     (2304 +
      skoXC1 *
      (1203012876672/125 +
       skoXC1 *
       (314071176527482786848/15625 +
        skoXC1 *
        (54663146636875953569585631888/1953125 +
         skoXC1 *
         (24974136555175809553645411545809035461/976562500 +
          skoXC1 *
          (3725720220689658122747691821425725453311913153/244140625000 +
           skoXC1 *
           1296900868161719997108481486962236487839483759224063561/
           244140625000000))))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (1214967134/1130791257 +
     exp(skoXC1 * -258222872) * 449447377/67847475420 <=
     2304 /
     (2304 +
      skoXC1 *
      (1203012876672/125 +
       skoXC1 *
       (314071176527482786848/15625 +
        skoXC1 *
        (54663146636875953569585631888/1953125 +
         skoXC1 *
         (24974136555175809553645411545809035461/976562500 +
          skoXC1 *
          (3725720220689658122747691821425725453311913153/244140625000 +
           skoXC1 *
           1296900868161719997108481486962236487839483759224063561/
           244140625000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((1214967134/1130791257 +
      exp(skoXC1 * -258222872) * 449447377/67847475420) *
     (2304 +
      skoXC1 *
      (1203012876672/125 +
       skoXC1 *
       (314071176527482786848/15625 +
        skoXC1 *
        (54663146636875953569585631888/1953125 +
         skoXC1 *
         (24974136555175809553645411545809035461/976562500 +
          skoXC1 *
          (3725720220689658122747691821425725453311913153/244140625000 +
           skoXC1 *
           1296900868161719997108481486962236487839483759224063561/
           244140625000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (1203012876672/125 +
      skoXC1 *
      (314071176527482786848/15625 +
       skoXC1 *
       (54663146636875953569585631888/1953125 +
        skoXC1 *
        (24974136555175809553645411545809035461/976562500 +
         skoXC1 *
         (3725720220689658122747691821425725453311913153/244140625000 +
          skoXC1 *
          1296900868161719997108481486962236487839483759224063561/
          244140625000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_5])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -258222872) *
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) <=
     -21549024512/125643473 +
     skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000))))) |
     skoXC1 *
     (1203012876672/125 +
      skoXC1 *
      (314071176527482786848/15625 +
       skoXC1 *
       (54663146636875953569585631888/1953125 +
        skoXC1 *
        (24974136555175809553645411545809035461/976562500 +
         skoXC1 *
         (3725720220689658122747691821425725453311913153/244140625000 +
          skoXC1 *
          1296900868161719997108481486962236487839483759224063561/
          244140625000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (-2304 <
     skoXC1 *
     (1203012876672/125 +
      skoXC1 *
      (314071176527482786848/15625 +
       skoXC1 *
       (54663146636875953569585631888/1953125 +
        skoXC1 *
        (24974136555175809553645411545809035461/976562500 +
         skoXC1 *
         (3725720220689658122747691821425725453311913153/244140625000 +
          skoXC1 *
          1296900868161719997108481486962236487839483759224063561/
          244140625000000)))))), inference(decision, [], [refute_0_15])).

cnf(refute_0_20, plain,
    (exp(skoXC1 * -258222872) *
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) <=
     -21549024512/125643473 +
     skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000)))))),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (28764632128/1884652095 +
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= 0 |
     exp(skoXC1 * -258222872) <=
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000))))))),
    inference(split, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= -28764632128/1884652095 |
     exp(skoXC1 * -258222872) <=
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000))))))),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-28764632128/1884652095 <
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000)))))),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_24, plain,
    (exp(skoXC1 * -258222872) <=
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000))))))),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (0 <
     (-21549024512/125643473 +
      skoXC1 *
      (-162402345215030588672/15705434125 +
       skoXC1 *
       (-42398461913067092641227494848/1963179265625 +
        skoXC1 *
        (-22137975534942305340652174914180789664/736192224609375 +
         skoXC1 *
         (-5057125852427764366587064153010352254225923129/
          184048056152343750 +
          skoXC1 *
          (-754437936436193572139097222898808907980204321926218917/
           46012014038085937500 +
           skoXC1 *
           -787845965436278396698690019653283415930281717492003284271002087/
           138036042114257812500000)))))) /
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000))))))),
    inference(resolve, [], [refute_0_24, refute_0_4])).

cnf(refute_0_26, plain,
    (0 *
     (28764632128/1884652095 +
      skoXC1 *
      (15019193942151524704/235581511875 +
       skoXC1 *
       (3921068513377252971152324936/29447688984375 +
        skoXC1 *
        (2047350656209189067185337435745763148/11042883369140625 +
         skoXC1 *
         (3741520055854527792579157002449795443615478599/
          22085766738281250000 +
          skoXC1 *
          (558171726874942658098564707314714901771058443155883227/
           5521441684570312500000 +
           skoXC1 *
           582888693424307864508854588768236887513148372617354922772729497/
           16564325053710937500000000)))))) <
     -21549024512/125643473 +
     skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000))))) |
     28764632128/1884652095 +
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_25])).

cnf(refute_0_27, plain,
    (21549024512/125643473 <
     skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000))))) |
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= -28764632128/1884652095),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 *
     (-162402345215030588672/15705434125 +
      skoXC1 *
      (-42398461913067092641227494848/1963179265625 +
       skoXC1 *
       (-22137975534942305340652174914180789664/736192224609375 +
        skoXC1 *
        (-5057125852427764366587064153010352254225923129/
         184048056152343750 +
         skoXC1 *
         (-754437936436193572139097222898808907980204321926218917/
          46012014038085937500 +
          skoXC1 *
          -787845965436278396698690019653283415930281717492003284271002087/
          138036042114257812500000))))) <= 21549024512/125643473 |
     skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= -28764632128/1884652095),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (15019193942151524704/235581511875 +
      skoXC1 *
      (3921068513377252971152324936/29447688984375 +
       skoXC1 *
       (2047350656209189067185337435745763148/11042883369140625 +
        skoXC1 *
        (3741520055854527792579157002449795443615478599/
         22085766738281250000 +
         skoXC1 *
         (558171726874942658098564707314714901771058443155883227/
          5521441684570312500000 +
          skoXC1 *
          582888693424307864508854588768236887513148372617354922772729497/
          16564325053710937500000000))))) <= -28764632128/1884652095),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain, ($false),
    inference(resolve, [], [refute_0_29, refute_0_23])).
% SZS output end CNFRefutation for Tunnel-1-IL-L.tptp
