% SZS output start CNFRefutation for Tunnel-3-IL.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 239/100000000000) =>
        -59/10000 - 1/62500 * exp(-255000000 * X) +
        31/1000 * exp(-54900000 * X) <= 3/100))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 239/100000000000) =>
        -59/10000 - 1/62500 * exp(-255000000 * X) +
        31/1000 * exp(-54900000 * X) <= 3/100)),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 239/100000000000) =>
          -59/10000 - 1/62500 * exp(-255000000 * X) +
          31/1000 * exp(-54900000 * X) <= 3/100)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (3/100 <
        -59/10000 - 1/62500 * exp(-255000000 * X) +
        31/1000 * exp(-54900000 * X) & X <= 239/100000000000 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (3/100 <
     -59/10000 - 1/62500 * exp(-255000000 * skoXC1) +
     31/1000 * exp(-54900000 * skoXC1) & 0 <= skoXC1 &
     skoXC1 <= 239/100000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (3/100 <
     -59/10000 - 1/62500 * exp(-255000000 * skoXC1) +
     31/1000 * exp(-54900000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 239/100000000000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (56448/155 +
     skoXC1 *
     (4540976640000/31 +
      skoXC1 *
      (124649808768000000000/31 +
       skoXC1 *
       (2281091500454400000000000000/31 +
        skoXC1 *
        (27394483238269560000000000000000000/31 +
         skoXC1 *
         (214851018540142692000000000000000000000000/31 +
          skoXC1 *
          982943409821152815900000000000000000000000000000/31))))) <
     0 *
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     0 <=
     -4608/3875 +
     skoXC1 *
     (-2023833600/31 +
      skoXC1 *
      (-55554232320000000/31 +
       skoXC1 *
       (-1016642451456000000000000/31 +
        skoXC1 *
        (-12209240440454400000000000000000/31 +
         skoXC1 *
         (-95755328597278080000000000000000000000/31 +
          skoXC1 * -438080628332547216000000000000000000000000000/31))))) |
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) <= 0),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000010, exp(X_000009)) | X_000010 <= exp(X_000009)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(0, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 <= exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) <=
     exp(skoXC1 * -255000000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    ((359/310 + exp(skoXC1 * -255000000) * 2/3875) *
     (2304 +
      skoXC1 *
      (126489600000 +
       skoXC1 *
       (3472139520000000000 +
        skoXC1 *
        (63540153216000000000000000 +
         skoXC1 *
         (763077527528400000000000000000000 +
          skoXC1 *
          (5984708037329880000000000000000000000000 +
           skoXC1 * 27380039270784201000000000000000000000000000000)))))) <
     2304 |
     2304 +
     skoXC1 *
     (126489600000 +
      skoXC1 *
      (3472139520000000000 +
       skoXC1 *
       (63540153216000000000000000 +
        skoXC1 *
        (763077527528400000000000000000000 +
         skoXC1 *
         (5984708037329880000000000000000000000000 +
          skoXC1 * 27380039270784201000000000000000000000000000000))))) <=
     0 |
     2304 /
     (2304 +
      skoXC1 *
      (126489600000 +
       skoXC1 *
       (3472139520000000000 +
        skoXC1 *
        (63540153216000000000000000 +
         skoXC1 *
         (763077527528400000000000000000000 +
          skoXC1 *
          (5984708037329880000000000000000000000000 +
           skoXC1 *
           27380039270784201000000000000000000000000000000)))))) <=
     359/310 + exp(skoXC1 * -255000000) * 2/3875),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_6, plain,
    (3/100 <
     -59/10000 - 1/62500 * exp(-255000000 * skoXC1) +
     31/1000 * exp(-54900000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (359/310 + exp(skoXC1 * -255000000) * 2/3875 <
     exp(skoXC1 * -54900000)), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -54900000 |
     359/310 + exp(skoXC1 * -255000000) * 2/3875 <
     2304 /
     (2304 +
      skoXC1 * -54900000 *
      (-2304 +
       skoXC1 * -54900000 *
       (1152 +
        skoXC1 * -54900000 *
        (-384 +
         skoXC1 * -54900000 *
         (84 + skoXC1 * -54900000 * (-12 + skoXC1 * -54900000)))))) |
     exp(skoXC1 * -54900000) <=
     359/310 + exp(skoXC1 * -255000000) * 2/3875),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -54900000 |
     359/310 + exp(skoXC1 * -255000000) * 2/3875 <
     2304 /
     (2304 +
      skoXC1 * -54900000 *
      (-2304 +
       skoXC1 * -54900000 *
       (1152 +
        skoXC1 * -54900000 *
        (-384 +
         skoXC1 * -54900000 *
         (84 + skoXC1 * -54900000 * (-12 + skoXC1 * -54900000))))))),
    inference(resolve, [], [refute_0_12, refute_0_7])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     359/310 + exp(skoXC1 * -255000000) * 2/3875 <
     2304 /
     (2304 +
      skoXC1 *
      (126489600000 +
       skoXC1 *
       (3472139520000000000 +
        skoXC1 *
        (63540153216000000000000000 +
         skoXC1 *
         (763077527528400000000000000000000 +
          skoXC1 *
          (5984708037329880000000000000000000000000 +
           skoXC1 * 27380039270784201000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (359/310 + exp(skoXC1 * -255000000) * 2/3875 <
     2304 /
     (2304 +
      skoXC1 *
      (126489600000 +
       skoXC1 *
       (3472139520000000000 +
        skoXC1 *
        (63540153216000000000000000 +
         skoXC1 *
         (763077527528400000000000000000000 +
          skoXC1 *
          (5984708037329880000000000000000000000000 +
           skoXC1 * 27380039270784201000000000000000000000000000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((359/310 + exp(skoXC1 * -255000000) * 2/3875) *
     (2304 +
      skoXC1 *
      (126489600000 +
       skoXC1 *
       (3472139520000000000 +
        skoXC1 *
        (63540153216000000000000000 +
         skoXC1 *
         (763077527528400000000000000000000 +
          skoXC1 *
          (5984708037329880000000000000000000000000 +
           skoXC1 * 27380039270784201000000000000000000000000000000)))))) <
     2304 |
     2304 +
     skoXC1 *
     (126489600000 +
      skoXC1 *
      (3472139520000000000 +
       skoXC1 *
       (63540153216000000000000000 +
        skoXC1 *
        (763077527528400000000000000000000 +
         skoXC1 *
         (5984708037329880000000000000000000000000 +
          skoXC1 * 27380039270784201000000000000000000000000000000))))) <=
     0), inference(resolve, [], [refute_0_5, refute_0_16])).

cnf(refute_0_18, plain,
    (56448/155 +
     skoXC1 *
     (4540976640000/31 +
      skoXC1 *
      (124649808768000000000/31 +
       skoXC1 *
       (2281091500454400000000000000/31 +
        skoXC1 *
        (27394483238269560000000000000000000/31 +
         skoXC1 *
         (214851018540142692000000000000000000000000/31 +
          skoXC1 *
          982943409821152815900000000000000000000000000000/31))))) <
     exp(skoXC1 * -255000000) *
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     skoXC1 *
     (126489600000 +
      skoXC1 *
      (3472139520000000000 +
       skoXC1 *
       (63540153216000000000000000 +
        skoXC1 *
        (763077527528400000000000000000000 +
         skoXC1 *
         (5984708037329880000000000000000000000000 +
          skoXC1 * 27380039270784201000000000000000000000000000000))))) <=
     -2304), inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain, (skoXC1 <= 239/100000000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_20, plain,
    (-2304 <
     skoXC1 *
     (126489600000 +
      skoXC1 *
      (3472139520000000000 +
       skoXC1 *
       (63540153216000000000000000 +
        skoXC1 *
        (763077527528400000000000000000000 +
         skoXC1 *
         (5984708037329880000000000000000000000000 +
          skoXC1 * 27380039270784201000000000000000000000000000000)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_21, plain,
    (56448/155 +
     skoXC1 *
     (4540976640000/31 +
      skoXC1 *
      (124649808768000000000/31 +
       skoXC1 *
       (2281091500454400000000000000/31 +
        skoXC1 *
        (27394483238269560000000000000000000/31 +
         skoXC1 *
         (214851018540142692000000000000000000000000/31 +
          skoXC1 *
          982943409821152815900000000000000000000000000000/31))))) <
     exp(skoXC1 * -255000000) *
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31))))))),
    inference(resolve, [], [refute_0_18, refute_0_20])).

cnf(refute_0_22, plain,
    (exp(skoXC1 * -255000000) <
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     0 <=
     -4608/3875 +
     skoXC1 *
     (-2023833600/31 +
      skoXC1 *
      (-55554232320000000/31 +
       skoXC1 *
       (-1016642451456000000000000/31 +
        skoXC1 *
        (-12209240440454400000000000000000/31 +
         skoXC1 *
         (-95755328597278080000000000000000000000/31 +
          skoXC1 * -438080628332547216000000000000000000000000000/31)))))),
    inference(split, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (exp(skoXC1 * -255000000) <
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     skoXC1 *
     (2023833600/31 +
      skoXC1 *
      (55554232320000000/31 +
       skoXC1 *
       (1016642451456000000000000/31 +
        skoXC1 *
        (12209240440454400000000000000000/31 +
         skoXC1 *
         (95755328597278080000000000000000000000/31 +
          skoXC1 * 438080628332547216000000000000000000000000000/31))))) <=
     -4608/3875), inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (-4608/3875 <
     skoXC1 *
     (2023833600/31 +
      skoXC1 *
      (55554232320000000/31 +
       skoXC1 *
       (1016642451456000000000000/31 +
        skoXC1 *
        (12209240440454400000000000000000/31 +
         skoXC1 *
         (95755328597278080000000000000000000000/31 +
          skoXC1 * 438080628332547216000000000000000000000000000/31)))))),
    inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_25, plain,
    (exp(skoXC1 * -255000000) <
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31))))))),
    inference(resolve, [], [refute_0_23, refute_0_24])).

cnf(refute_0_26, plain,
    (0 <
     (56448/155 +
      skoXC1 *
      (4540976640000/31 +
       skoXC1 *
       (124649808768000000000/31 +
        skoXC1 *
        (2281091500454400000000000000/31 +
         skoXC1 *
         (27394483238269560000000000000000000/31 +
          skoXC1 *
          (214851018540142692000000000000000000000000/31 +
           skoXC1 *
           982943409821152815900000000000000000000000000000/31)))))) /
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31))))))),
    inference(resolve, [], [refute_0_4, refute_0_25])).

cnf(refute_0_27, plain,
    (56448/155 +
     skoXC1 *
     (4540976640000/31 +
      skoXC1 *
      (124649808768000000000/31 +
       skoXC1 *
       (2281091500454400000000000000/31 +
        skoXC1 *
        (27394483238269560000000000000000000/31 +
         skoXC1 *
         (214851018540142692000000000000000000000000/31 +
          skoXC1 *
          982943409821152815900000000000000000000000000000/31))))) <
     0 *
     (-4608/3875 +
      skoXC1 *
      (-2023833600/31 +
       skoXC1 *
       (-55554232320000000/31 +
        skoXC1 *
        (-1016642451456000000000000/31 +
         skoXC1 *
         (-12209240440454400000000000000000/31 +
          skoXC1 *
          (-95755328597278080000000000000000000000/31 +
           skoXC1 *
           -438080628332547216000000000000000000000000000/31)))))) |
     0 <=
     -4608/3875 +
     skoXC1 *
     (-2023833600/31 +
      skoXC1 *
      (-55554232320000000/31 +
       skoXC1 *
       (-1016642451456000000000000/31 +
        skoXC1 *
        (-12209240440454400000000000000000/31 +
         skoXC1 *
         (-95755328597278080000000000000000000000/31 +
          skoXC1 * -438080628332547216000000000000000000000000000/31)))))),
    inference(resolve, [], [refute_0_0, refute_0_26])).

cnf(refute_0_28, plain,
    (56448/155 <
     skoXC1 *
     (-4540976640000/31 +
      skoXC1 *
      (-124649808768000000000/31 +
       skoXC1 *
       (-2281091500454400000000000000/31 +
        skoXC1 *
        (-27394483238269560000000000000000000/31 +
         skoXC1 *
         (-214851018540142692000000000000000000000000/31 +
          skoXC1 *
          -982943409821152815900000000000000000000000000000/31))))) |
     skoXC1 *
     (2023833600/31 +
      skoXC1 *
      (55554232320000000/31 +
       skoXC1 *
       (1016642451456000000000000/31 +
        skoXC1 *
        (12209240440454400000000000000000/31 +
         skoXC1 *
         (95755328597278080000000000000000000000/31 +
          skoXC1 * 438080628332547216000000000000000000000000000/31))))) <=
     -4608/3875), inference(arithmetic, [], [refute_0_27])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (-4540976640000/31 +
      skoXC1 *
      (-124649808768000000000/31 +
       skoXC1 *
       (-2281091500454400000000000000/31 +
        skoXC1 *
        (-27394483238269560000000000000000000/31 +
         skoXC1 *
         (-214851018540142692000000000000000000000000/31 +
          skoXC1 *
          -982943409821152815900000000000000000000000000000/31))))) <=
     56448/155 |
     skoXC1 *
     (2023833600/31 +
      skoXC1 *
      (55554232320000000/31 +
       skoXC1 *
       (1016642451456000000000000/31 +
        skoXC1 *
        (12209240440454400000000000000000/31 +
         skoXC1 *
         (95755328597278080000000000000000000000/31 +
          skoXC1 * 438080628332547216000000000000000000000000000/31))))) <=
     -4608/3875), inference(decision, [], [refute_0_19, refute_0_15])).

cnf(refute_0_30, plain,
    (skoXC1 *
     (2023833600/31 +
      skoXC1 *
      (55554232320000000/31 +
       skoXC1 *
       (1016642451456000000000000/31 +
        skoXC1 *
        (12209240440454400000000000000000/31 +
         skoXC1 *
         (95755328597278080000000000000000000000/31 +
          skoXC1 * 438080628332547216000000000000000000000000000/31))))) <=
     -4608/3875), inference(resolve, [], [refute_0_29, refute_0_28])).

cnf(refute_0_31, plain, ($false),
    inference(resolve, [], [refute_0_30, refute_0_24])).
% SZS output end CNFRefutation for Tunnel-3-IL.tptp
