% SZS output start CNFRefutation for Tunnel-1-VC-L.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       (0 <= X =>
        0 <
        27813/100000 + 57961/500000 * exp(-258220000 * X) -
        263053/1000000 * exp(-4177120 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       (0 <= X =>
        0 <
        27813/100000 + 57961/500000 * exp(-258220000 * X) -
        263053/1000000 * exp(-4177120 * X))),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         (0 <= X =>
          0 <
          27813/100000 + 57961/500000 * exp(-258220000 * X) -
          263053/1000000 * exp(-4177120 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (0 <= X &
        27813/100000 + 57961/500000 * exp(-258220000 * X) -
        263053/1000000 * exp(-4177120 * X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     27813/100000 + 57961/500000 * exp(-258220000 * skoXC1) -
     263053/1000000 * exp(-4177120 * skoXC1) <= 0),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (27813/100000 + 57961/500000 * exp(-258220000 * skoXC1) -
     263053/1000000 * exp(-4177120 * skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <
     -34737408/263053 +
     skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053))))) |
     267084288/263053 +
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <= 0 |
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 <
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) |
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <
     exp(skoXC1 * -258220000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (9624084480 +
       skoXC1 *
       (20100477881548800 +
        skoXC1 *
        (27987369389525041152000 +
         skoXC1 *
         (25573318842831558727434240000 +
          skoXC1 *
          (15260403086395508655934301798400000 +
           skoXC1 * 5312044578353700593073024227344384000000)))))) <
     278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 |
     (278130/263053 + exp(skoXC1 * -258220000) * 115922/263053) *
     (2304 +
      skoXC1 *
      (9624084480 +
       skoXC1 *
       (20100477881548800 +
        skoXC1 *
        (27987369389525041152000 +
         skoXC1 *
         (25573318842831558727434240000 +
          skoXC1 *
          (15260403086395508655934301798400000 +
           skoXC1 * 5312044578353700593073024227344384000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (9624084480 +
      skoXC1 *
      (20100477881548800 +
       skoXC1 *
       (27987369389525041152000 +
        skoXC1 *
        (25573318842831558727434240000 +
         skoXC1 *
         (15260403086395508655934301798400000 +
          skoXC1 * 5312044578353700593073024227344384000000))))) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_6, plain,
    (27813/100000 + 57961/500000 * exp(-258220000 * skoXC1) -
     263053/1000000 * exp(-4177120 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_7, plain,
    (278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 <=
     exp(skoXC1 * -4177120)), inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_9, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_10, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)), inference(resolve, [], [refute_0_9, refute_0_8])).

cnf(refute_0_11, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < skoXC1 * -4177120 |
     exp(skoXC1 * -4177120) <
     278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 |
     278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 <=
     2304 /
     (2304 +
      skoXC1 * -4177120 *
      (-2304 +
       skoXC1 * -4177120 *
       (1152 +
        skoXC1 * -4177120 *
        (-384 +
         skoXC1 * -4177120 *
         (84 + skoXC1 * -4177120 * (-12 + skoXC1 * -4177120))))))),
    inference(subst, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -4177120 |
     278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 <=
     2304 /
     (2304 +
      skoXC1 * -4177120 *
      (-2304 +
       skoXC1 * -4177120 *
       (1152 +
        skoXC1 * -4177120 *
        (-384 +
         skoXC1 * -4177120 *
         (84 + skoXC1 * -4177120 * (-12 + skoXC1 * -4177120))))))),
    inference(resolve, [], [refute_0_7, refute_0_12])).

cnf(refute_0_14, plain,
    (skoXC1 < 0 |
     278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 <=
     2304 /
     (2304 +
      skoXC1 *
      (9624084480 +
       skoXC1 *
       (20100477881548800 +
        skoXC1 *
        (27987369389525041152000 +
         skoXC1 *
         (25573318842831558727434240000 +
          skoXC1 *
          (15260403086395508655934301798400000 +
           skoXC1 * 5312044578353700593073024227344384000000))))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_16, plain,
    (278130/263053 + exp(skoXC1 * -258220000) * 115922/263053 <=
     2304 /
     (2304 +
      skoXC1 *
      (9624084480 +
       skoXC1 *
       (20100477881548800 +
        skoXC1 *
        (27987369389525041152000 +
         skoXC1 *
         (25573318842831558727434240000 +
          skoXC1 *
          (15260403086395508655934301798400000 +
           skoXC1 * 5312044578353700593073024227344384000000))))))),
    inference(resolve, [], [refute_0_15, refute_0_14])).

cnf(refute_0_17, plain,
    ((278130/263053 + exp(skoXC1 * -258220000) * 115922/263053) *
     (2304 +
      skoXC1 *
      (9624084480 +
       skoXC1 *
       (20100477881548800 +
        skoXC1 *
        (27987369389525041152000 +
         skoXC1 *
         (25573318842831558727434240000 +
          skoXC1 *
          (15260403086395508655934301798400000 +
           skoXC1 * 5312044578353700593073024227344384000000)))))) <=
     2304 |
     2304 +
     skoXC1 *
     (9624084480 +
      skoXC1 *
      (20100477881548800 +
       skoXC1 *
       (27987369389525041152000 +
        skoXC1 *
        (25573318842831558727434240000 +
         skoXC1 *
         (15260403086395508655934301798400000 +
          skoXC1 * 5312044578353700593073024227344384000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_5])).

cnf(refute_0_18, plain,
    (exp(skoXC1 * -258220000) *
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <=
     -34737408/263053 +
     skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053))))) |
     skoXC1 *
     (9624084480 +
      skoXC1 *
      (20100477881548800 +
       skoXC1 *
       (27987369389525041152000 +
        skoXC1 *
        (25573318842831558727434240000 +
         skoXC1 *
         (15260403086395508655934301798400000 +
          skoXC1 * 5312044578353700593073024227344384000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_17])).

cnf(refute_0_19, plain,
    (-2304 <
     skoXC1 *
     (9624084480 +
      skoXC1 *
      (20100477881548800 +
       skoXC1 *
       (27987369389525041152000 +
        skoXC1 *
        (25573318842831558727434240000 +
         skoXC1 *
         (15260403086395508655934301798400000 +
          skoXC1 * 5312044578353700593073024227344384000000)))))),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_20, plain,
    (exp(skoXC1 * -258220000) *
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <=
     -34737408/263053 +
     skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053)))))),
    inference(resolve, [], [refute_0_18, refute_0_19])).

cnf(refute_0_21, plain,
    (267084288/263053 +
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <= 0 |
     exp(skoXC1 * -258220000) <=
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053))))))),
    inference(split, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <=
     -267084288/263053 |
     exp(skoXC1 * -258220000) <=
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053))))))),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (-267084288/263053 <
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053)))))),
    inference(decision, [], [refute_0_15])).

cnf(refute_0_24, plain,
    (exp(skoXC1 * -258220000) <=
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053))))))),
    inference(resolve, [], [refute_0_22, refute_0_23])).

cnf(refute_0_25, plain,
    (0 <
     (-34737408/263053 +
      skoXC1 *
      (-2676746616422400/263053 +
       skoXC1 *
       (-5590545913195167744000/263053 +
        skoXC1 *
        (-7784127048308599695605760000/263053 +
         skoXC1 *
         (-1016101024250963061265897881600000/37579 +
          skoXC1 *
          (-4244375910419182822475007359188992000000/263053 +
           skoXC1 *
           -1477438958577514745951400228351293521920000000/263053)))))) /
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053))))))),
    inference(resolve, [], [refute_0_24, refute_0_4])).

cnf(refute_0_26, plain,
    (0 *
     (267084288/263053 +
      skoXC1 *
      (1115643121090560/263053 +
       skoXC1 *
       (2330087596984899993600/263053 +
        skoXC1 *
        (3244351834372521820422144000/263053 +
         skoXC1 *
         (423501466699817135828804567040000/37579 +
          skoXC1 *
          (1769016446581140154413216133074124800000/263053 +
           skoXC1 *
           615782831611917680150211114482215682048000000/263053)))))) <
     -34737408/263053 +
     skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053))))) |
     267084288/263053 +
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_25])).

cnf(refute_0_27, plain,
    (34737408/263053 <
     skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053))))) |
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <=
     -267084288/263053), inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (skoXC1 *
     (-2676746616422400/263053 +
      skoXC1 *
      (-5590545913195167744000/263053 +
       skoXC1 *
       (-7784127048308599695605760000/263053 +
        skoXC1 *
        (-1016101024250963061265897881600000/37579 +
         skoXC1 *
         (-4244375910419182822475007359188992000000/263053 +
          skoXC1 *
          -1477438958577514745951400228351293521920000000/263053))))) <=
     34737408/263053 |
     skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <=
     -267084288/263053), inference(decision, [], [refute_0_15])).

cnf(refute_0_29, plain,
    (skoXC1 *
     (1115643121090560/263053 +
      skoXC1 *
      (2330087596984899993600/263053 +
       skoXC1 *
       (3244351834372521820422144000/263053 +
        skoXC1 *
        (423501466699817135828804567040000/37579 +
         skoXC1 *
         (1769016446581140154413216133074124800000/263053 +
          skoXC1 *
          615782831611917680150211114482215682048000000/263053))))) <=
     -267084288/263053),
    inference(resolve, [], [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain, ($false),
    inference(resolve, [], [refute_0_29, refute_0_23])).
% SZS output end CNFRefutation for Tunnel-1-VC-L.tptp
