% SZS output start CNFRefutation for Tunnel-2-VC-U.tptp
cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_mul_divide_neg, axiom, (Y / Z < X | Y <= X * Z | 0 <= Z)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_upper_bound_cf3, axiom,
    (X < 0 | 4 < X |
     ~ lgen(R,
            -(X ^ 3 + 12 * X ^ 2 + 60 * X + 120) /
            (X ^ 3 - 12 * X ^ 2 + 60 * X - 120), Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

cnf(exp_upper_bound_cf5, axiom,
    (X < 0 | 7 < X |
     ~ lgen(R,
            -(X ^ 5 + 30 * X ^ 4 + 420 * X ^ 3 + 3360 * X ^ 2 + 15120 * X +
             30240) /
            (X ^ 5 - 30 * X ^ 4 + 420 * X ^ 3 - 3360 * X ^ 2 + 15120 * X -
             30240), Y) | lgen(R, exp(X), Y))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        27/100 + 3/1250 * exp(87800000 * X) - 11/2500 * exp(11000000 * X) <
        9/10))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        27/100 + 3/1250 * exp(87800000 * X) - 11/2500 * exp(11000000 * X) <
        9/10)), inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 593/10000000000) =>
          27/100 + 3/1250 * exp(87800000 * X) -
          11/2500 * exp(11000000 * X) < 9/10)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 593/10000000000 & 0 <= X &
        9/10 <=
        27/100 + 3/1250 * exp(87800000 * X) -
        11/2500 * exp(11000000 * X))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     9/10 <=
     27/100 + 3/1250 * exp(87800000 * skoXC1) -
     11/2500 * exp(11000000 * skoXC1) & skoXC1 <= 593/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (9/10 <=
     27/100 + 3/1250 * exp(87800000 * skoXC1) -
     11/2500 * exp(11000000 * skoXC1)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 593/10000000000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (0 *
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) <
     -7907760 +
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) |
     55440 +
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <= 0 |
     (-7907760 +
      skoXC1 *
      (349805736000000 +
       skoXC1 *
       (-6773295177600000000000 +
        skoXC1 *
        (74905456941840000000000000000 +
         skoXC1 *
         (-466198649972232000000000000000000000 +
          skoXC1 * 1374843292122699680000000000000000000000000))))) /
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_2, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_3, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (0 < X_000029 / X_000030 | X_000029 / X_000030 < exp(X_000009)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (X_000029 < exp(X_000009) * X_000030 | X_000030 <= 0 |
     exp(X_000009) <= X_000029 / X_000030),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (X_000029 < exp(X_000009) * X_000030 | 0 < X_000029 / X_000030 |
     X_000030 <= 0), inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (-7907760 +
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) <
     exp(skoXC1 * 11000000) *
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) |
     0 <
     (-7907760 +
      skoXC1 *
      (349805736000000 +
       skoXC1 *
       (-6773295177600000000000 +
        skoXC1 *
        (74905456941840000000000000000 +
         skoXC1 *
         (-466198649972232000000000000000000000 +
          skoXC1 * 1374843292122699680000000000000000000000000))))) /
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) |
     55440 +
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <= 0),
    inference(subst, [], [refute_0_6])).

cnf(refute_0_8, plain,
    ((-30240 +
      skoXC1 *
      (-1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (-284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * -5217621601983680000000000000000000000000))))) /
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000))))) <
     525/2 + exp(skoXC1 * 11000000) * 11/6 |
     0 <=
     -30240 +
     skoXC1 *
     (1327536000000 +
      skoXC1 *
      (-25901702400000000000 +
       skoXC1 *
       (284271183840000000000000000 +
        skoXC1 *
        (-1782786424368000000000000000000000 +
         skoXC1 * 5217621601983680000000000000000000000000)))) |
     -30240 +
     skoXC1 *
     (-1327536000000 +
      skoXC1 *
      (-25901702400000000000 +
       skoXC1 *
       (-284271183840000000000000000 +
        skoXC1 *
        (-1782786424368000000000000000000000 +
         skoXC1 * -5217621601983680000000000000000000000000)))) <=
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000)))))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_9, plain,
    (9/10 <=
     27/100 + 3/1250 * exp(87800000 * skoXC1) -
     11/2500 * exp(11000000 * skoXC1)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_10, plain,
    (525/2 + exp(skoXC1 * 11000000) * 11/6 <= exp(skoXC1 * 87800000)),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (exp(X_000158) < X_000159 | ~ lgen(1, exp(X_000158), X_000159)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_12, plain,
    (7 < X_000158 | X_000158 < 0 |
     ~ lgen(1,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159) |
     lgen(1, exp(X_000158), X_000159)),
    inference(subst, [], [exp_upper_bound_cf5])).

cnf(refute_0_13, plain,
    (7 < X_000158 | exp(X_000158) < X_000159 | X_000158 < 0 |
     ~ lgen(1,
            -(X_000158 ^ 5 + 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 +
             3360 * X_000158 ^ 2 + 15120 * X_000158 + 30240) /
            (X_000158 ^ 5 - 30 * X_000158 ^ 4 + 420 * X_000158 ^ 3 -
             3360 * X_000158 ^ 2 + 15120 * X_000158 - 30240), X_000159)),
    inference(resolve, [], [refute_0_12, refute_0_11])).

cnf(refute_0_14, plain,
    (7 < X_000158 | exp(X_000158) < X_000159 | X_000158 < 0 |
     X_000159 <=
     (-30240 +
      X_000158 *
      (-15120 +
       X_000158 *
       (-3360 + X_000158 * (-420 + X_000158 * (-30 + X_000158 * -1))))) /
     (-30240 +
      X_000158 *
      (15120 +
       X_000158 *
       (-3360 + X_000158 * (420 + X_000158 * (-30 + X_000158)))))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (skoXC1 * 87800000 < 0 | 7 < skoXC1 * 87800000 |
     exp(skoXC1 * 87800000) < 525/2 + exp(skoXC1 * 11000000) * 11/6 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-30240 +
      skoXC1 * 87800000 *
      (-15120 +
       skoXC1 * 87800000 *
       (-3360 +
        skoXC1 * 87800000 *
        (-420 + skoXC1 * 87800000 * (-30 + skoXC1 * 87800000 * -1))))) /
     (-30240 +
      skoXC1 * 87800000 *
      (15120 +
       skoXC1 * 87800000 *
       (-3360 +
        skoXC1 * 87800000 *
        (420 + skoXC1 * 87800000 * (-30 + skoXC1 * 87800000)))))),
    inference(subst, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (skoXC1 * 87800000 < 0 | 7 < skoXC1 * 87800000 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-30240 +
      skoXC1 * 87800000 *
      (-15120 +
       skoXC1 * 87800000 *
       (-3360 +
        skoXC1 * 87800000 *
        (-420 + skoXC1 * 87800000 * (-30 + skoXC1 * 87800000 * -1))))) /
     (-30240 +
      skoXC1 * 87800000 *
      (15120 +
       skoXC1 * 87800000 *
       (-3360 +
        skoXC1 * 87800000 *
        (420 + skoXC1 * 87800000 * (-30 + skoXC1 * 87800000)))))),
    inference(resolve, [], [refute_0_10, refute_0_15])).

cnf(refute_0_17, plain,
    (skoXC1 < 0 | 7/87800000 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-30240 +
      skoXC1 *
      (-1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (-284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * -5217621601983680000000000000000000000000))))) /
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000)))))),
    inference(arithmetic, [], [refute_0_16])).

cnf(refute_0_18, plain,
    (0 *
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <
     -31380 +
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000)) |
     220 +
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <= 0 |
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_19, plain,
    (0 <
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) |
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <
     exp(skoXC1 * 11000000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_20, plain,
    ((-120 +
      skoXC1 *
      (-5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * -676836152000000000000000))) /
     (-120 +
      skoXC1 *
      (5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * 676836152000000000000000))) <
     525/2 + exp(skoXC1 * 11000000) * 11/6 |
     0 <=
     -120 +
     skoXC1 *
     (5268000000 +
      skoXC1 * (-92506080000000000 + skoXC1 * 676836152000000000000000)) |
     -120 +
     skoXC1 *
     (-5268000000 +
      skoXC1 *
      (-92506080000000000 + skoXC1 * -676836152000000000000000)) <=
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (-120 +
      skoXC1 *
      (5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * 676836152000000000000000)))),
    inference(subst, [], [leq_left_mul_divide_neg])).

cnf(refute_0_21, plain,
    (exp(X_000138) < X_000139 | ~ lgen(1, exp(X_000138), X_000139)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_22, plain,
    (4 < X_000138 | X_000138 < 0 |
     ~ lgen(1,
            -(X_000138 ^ 3 + 12 * X_000138 ^ 2 + 60 * X_000138 + 120) /
            (X_000138 ^ 3 - 12 * X_000138 ^ 2 + 60 * X_000138 - 120),
            X_000139) | lgen(1, exp(X_000138), X_000139)),
    inference(subst, [], [exp_upper_bound_cf3])).

cnf(refute_0_23, plain,
    (4 < X_000138 | exp(X_000138) < X_000139 | X_000138 < 0 |
     ~ lgen(1,
            -(X_000138 ^ 3 + 12 * X_000138 ^ 2 + 60 * X_000138 + 120) /
            (X_000138 ^ 3 - 12 * X_000138 ^ 2 + 60 * X_000138 - 120),
            X_000139)),
    inference(resolve, [], [refute_0_22, refute_0_21])).

cnf(refute_0_24, plain,
    (4 < X_000138 | exp(X_000138) < X_000139 | X_000138 < 0 |
     X_000139 <=
     (-120 + X_000138 * (-60 + X_000138 * (-12 + X_000138 * -1))) /
     (-120 + X_000138 * (60 + X_000138 * (-12 + X_000138)))),
    inference(arithmetic, [], [refute_0_23])).

cnf(refute_0_25, plain,
    (skoXC1 * 87800000 < 0 | 4 < skoXC1 * 87800000 |
     exp(skoXC1 * 87800000) < 525/2 + exp(skoXC1 * 11000000) * 11/6 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-120 +
      skoXC1 * 87800000 *
      (-60 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000 * -1))) /
     (-120 +
      skoXC1 * 87800000 *
      (60 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000)))),
    inference(subst, [], [refute_0_24])).

cnf(refute_0_26, plain,
    (skoXC1 * 87800000 < 0 | 4 < skoXC1 * 87800000 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-120 +
      skoXC1 * 87800000 *
      (-60 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000 * -1))) /
     (-120 +
      skoXC1 * 87800000 *
      (60 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000)))),
    inference(resolve, [], [refute_0_10, refute_0_25])).

cnf(refute_0_27, plain,
    (skoXC1 < 0 | 1/21950000 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-120 +
      skoXC1 *
      (-5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * -676836152000000000000000))) /
     (-120 +
      skoXC1 *
      (5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * 676836152000000000000000)))),
    inference(arithmetic, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (0 *
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) <
     -5685865611264 +
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000))))))))))) |
     39862665216 +
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= 0 |
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_29, plain,
    (0 <
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) |
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) < exp(skoXC1 * 11000000)),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_30, plain,
    (21743271936 /
     (21743271936 +
      skoXC1 *
      (-1909059275980800000 +
       skoXC1 *
       (83807702215557120000000000 +
        skoXC1 *
        (-2452772084841971712000000000000000 +
         skoXC1 *
         (52997123086308134092800000000000000000000 +
          skoXC1 *
          (-886313791805305556828160000000000000000000000000 +
           skoXC1 *
           (11753813420284734420811776000000000000000000000000000000 +
            skoXC1 *
            (-124549891863913754741912371200000000000000000000000000000000000 +
             skoXC1 *
             (1049610852104955334046017781760000000000000000000000000000000000000000 +
              skoXC1 *
              (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
               skoXC1 *
               (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000)))))))))))) <
     525/2 + exp(skoXC1 * 11000000) * 11/6 |
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (21743271936 +
      skoXC1 *
      (-1909059275980800000 +
       skoXC1 *
       (83807702215557120000000000 +
        skoXC1 *
        (-2452772084841971712000000000000000 +
         skoXC1 *
         (52997123086308134092800000000000000000000 +
          skoXC1 *
          (-886313791805305556828160000000000000000000000000 +
           skoXC1 *
           (11753813420284734420811776000000000000000000000000000000 +
            skoXC1 *
            (-124549891863913754741912371200000000000000000000000000000000000 +
             skoXC1 *
             (1049610852104955334046017781760000000000000000000000000000000000000000 +
              skoXC1 *
              (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
               skoXC1 *
               (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000)))))))))))) <=
     21743271936 |
     21743271936 +
     skoXC1 *
     (-1909059275980800000 +
      skoXC1 *
      (83807702215557120000000000 +
       skoXC1 *
       (-2452772084841971712000000000000000 +
        skoXC1 *
        (52997123086308134092800000000000000000000 +
         skoXC1 *
         (-886313791805305556828160000000000000000000000000 +
          skoXC1 *
          (11753813420284734420811776000000000000000000000000000000 +
           skoXC1 *
           (-124549891863913754741912371200000000000000000000000000000000000 +
            skoXC1 *
            (1049610852104955334046017781760000000000000000000000000000000000000000 +
             skoXC1 *
             (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
              skoXC1 *
              (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000))))))))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_31, plain,
    (exp(X_000061) < X_000062 | ~ lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_32, plain,
    (3 < X_000061 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(1, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_33, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     ~ lgen(1,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062)),
    inference(resolve, [], [refute_0_32, refute_0_31])).

cnf(refute_0_34, plain,
    (3 < X_000061 | exp(X_000061) < X_000062 |
     X_000062 <=
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061))))))))))))),
    inference(arithmetic, [], [refute_0_33])).

cnf(refute_0_35, plain,
    (3 < skoXC1 * 87800000 |
     exp(skoXC1 * 87800000) < 525/2 + exp(skoXC1 * 11000000) * 11/6 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     21743271936 /
     (21743271936 +
      skoXC1 * 87800000 *
      (-21743271936 +
       skoXC1 * 87800000 *
       (10871635968 +
        skoXC1 * 87800000 *
        (-3623878656 +
         skoXC1 * 87800000 *
         (891813888 +
          skoXC1 * 87800000 *
          (-169869312 +
           skoXC1 * 87800000 *
           (25657344 +
            skoXC1 * 87800000 *
            (-3096576 +
             skoXC1 * 87800000 *
             (297216 +
              skoXC1 * 87800000 *
              (-22272 +
               skoXC1 * 87800000 *
               (1248 +
                skoXC1 * 87800000 * (-48 + skoXC1 * 87800000))))))))))))),
    inference(subst, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (3 < skoXC1 * 87800000 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     21743271936 /
     (21743271936 +
      skoXC1 * 87800000 *
      (-21743271936 +
       skoXC1 * 87800000 *
       (10871635968 +
        skoXC1 * 87800000 *
        (-3623878656 +
         skoXC1 * 87800000 *
         (891813888 +
          skoXC1 * 87800000 *
          (-169869312 +
           skoXC1 * 87800000 *
           (25657344 +
            skoXC1 * 87800000 *
            (-3096576 +
             skoXC1 * 87800000 *
             (297216 +
              skoXC1 * 87800000 *
              (-22272 +
               skoXC1 * 87800000 *
               (1248 +
                skoXC1 * 87800000 * (-48 + skoXC1 * 87800000))))))))))))),
    inference(resolve, [], [refute_0_10, refute_0_35])).

cnf(refute_0_37, plain,
    (3/87800000 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-1909059275980800000 +
       skoXC1 *
       (83807702215557120000000000 +
        skoXC1 *
        (-2452772084841971712000000000000000 +
         skoXC1 *
         (52997123086308134092800000000000000000000 +
          skoXC1 *
          (-886313791805305556828160000000000000000000000000 +
           skoXC1 *
           (11753813420284734420811776000000000000000000000000000000 +
            skoXC1 *
            (-124549891863913754741912371200000000000000000000000000000000000 +
             skoXC1 *
             (1049610852104955334046017781760000000000000000000000000000000000000000 +
              skoXC1 *
              (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
               skoXC1 *
               (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000))))))))))))),
    inference(arithmetic, [], [refute_0_36])).

cnf(refute_0_38, plain,
    (3/87800000 < skoXC1 |
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (21743271936 +
      skoXC1 *
      (-1909059275980800000 +
       skoXC1 *
       (83807702215557120000000000 +
        skoXC1 *
        (-2452772084841971712000000000000000 +
         skoXC1 *
         (52997123086308134092800000000000000000000 +
          skoXC1 *
          (-886313791805305556828160000000000000000000000000 +
           skoXC1 *
           (11753813420284734420811776000000000000000000000000000000 +
            skoXC1 *
            (-124549891863913754741912371200000000000000000000000000000000000 +
             skoXC1 *
             (1049610852104955334046017781760000000000000000000000000000000000000000 +
              skoXC1 *
              (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
               skoXC1 *
               (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000)))))))))))) <=
     21743271936 |
     21743271936 +
     skoXC1 *
     (-1909059275980800000 +
      skoXC1 *
      (83807702215557120000000000 +
       skoXC1 *
       (-2452772084841971712000000000000000 +
        skoXC1 *
        (52997123086308134092800000000000000000000 +
         skoXC1 *
         (-886313791805305556828160000000000000000000000000 +
          skoXC1 *
          (11753813420284734420811776000000000000000000000000000000 +
           skoXC1 *
           (-124549891863913754741912371200000000000000000000000000000000000 +
            skoXC1 *
            (1049610852104955334046017781760000000000000000000000000000000000000000 +
             skoXC1 *
             (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
              skoXC1 *
              (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000))))))))))) <=
     0), inference(resolve, [], [refute_0_37, refute_0_30])).

cnf(refute_0_39, plain,
    (3/87800000 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) <=
     -5685865611264 +
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000))))))))))) |
     skoXC1 *
     (-1909059275980800000 +
      skoXC1 *
      (83807702215557120000000000 +
       skoXC1 *
       (-2452772084841971712000000000000000 +
        skoXC1 *
        (52997123086308134092800000000000000000000 +
         skoXC1 *
         (-886313791805305556828160000000000000000000000000 +
          skoXC1 *
          (11753813420284734420811776000000000000000000000000000000 +
           skoXC1 *
           (-124549891863913754741912371200000000000000000000000000000000000 +
            skoXC1 *
            (1049610852104955334046017781760000000000000000000000000000000000000000 +
             skoXC1 *
             (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
              skoXC1 *
              (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000))))))))))) <=
     -21743271936), inference(arithmetic, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (0 *
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <
     -602496 +
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000))))) |
     4224 +
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <= 0 |
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_41, plain,
    (0 <
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) |
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <
     exp(skoXC1 * 11000000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_42, plain,
    (2304 /
     (2304 +
      skoXC1 *
      (-202291200000 +
       skoXC1 *
       (8880583680000000000 +
        skoXC1 *
        (-259905082368000000000000000 +
         skoXC1 *
         (4991801988230400000000000000000000 +
          skoXC1 *
          (-62611459223804160000000000000000000000000 +
           skoXC1 *
           458107176654167104000000000000000000000000000000)))))) <
     525/2 + exp(skoXC1 * 11000000) * 11/6 |
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (2304 +
      skoXC1 *
      (-202291200000 +
       skoXC1 *
       (8880583680000000000 +
        skoXC1 *
        (-259905082368000000000000000 +
         skoXC1 *
         (4991801988230400000000000000000000 +
          skoXC1 *
          (-62611459223804160000000000000000000000000 +
           skoXC1 *
           458107176654167104000000000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (-202291200000 +
      skoXC1 *
      (8880583680000000000 +
       skoXC1 *
       (-259905082368000000000000000 +
        skoXC1 *
        (4991801988230400000000000000000000 +
         skoXC1 *
         (-62611459223804160000000000000000000000000 +
          skoXC1 * 458107176654167104000000000000000000000000000000))))) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_43, plain,
    (exp(X_000054) < X_000055 | ~ lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_44, plain,
    (0 < X_000054 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(1, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_45, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     ~ lgen(1,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055)),
    inference(resolve, [], [refute_0_44, refute_0_43])).

cnf(refute_0_46, plain,
    (0 < X_000054 | exp(X_000054) < X_000055 |
     X_000055 <=
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054))))))),
    inference(arithmetic, [], [refute_0_45])).

cnf(refute_0_47, plain,
    (0 < skoXC1 * 87800000 |
     exp(skoXC1 * 87800000) < 525/2 + exp(skoXC1 * 11000000) * 11/6 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     2304 /
     (2304 +
      skoXC1 * 87800000 *
      (-2304 +
       skoXC1 * 87800000 *
       (1152 +
        skoXC1 * 87800000 *
        (-384 +
         skoXC1 * 87800000 *
         (84 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000))))))),
    inference(subst, [], [refute_0_46])).

cnf(refute_0_48, plain,
    (0 < skoXC1 * 87800000 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     2304 /
     (2304 +
      skoXC1 * 87800000 *
      (-2304 +
       skoXC1 * 87800000 *
       (1152 +
        skoXC1 * 87800000 *
        (-384 +
         skoXC1 * 87800000 *
         (84 + skoXC1 * 87800000 * (-12 + skoXC1 * 87800000))))))),
    inference(resolve, [], [refute_0_10, refute_0_47])).

cnf(refute_0_49, plain,
    (0 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     2304 /
     (2304 +
      skoXC1 *
      (-202291200000 +
       skoXC1 *
       (8880583680000000000 +
        skoXC1 *
        (-259905082368000000000000000 +
         skoXC1 *
         (4991801988230400000000000000000000 +
          skoXC1 *
          (-62611459223804160000000000000000000000000 +
           skoXC1 *
           458107176654167104000000000000000000000000000000))))))),
    inference(arithmetic, [], [refute_0_48])).

cnf(refute_0_50, plain,
    (0 < skoXC1 |
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (2304 +
      skoXC1 *
      (-202291200000 +
       skoXC1 *
       (8880583680000000000 +
        skoXC1 *
        (-259905082368000000000000000 +
         skoXC1 *
         (4991801988230400000000000000000000 +
          skoXC1 *
          (-62611459223804160000000000000000000000000 +
           skoXC1 *
           458107176654167104000000000000000000000000000000)))))) <= 2304 |
     2304 +
     skoXC1 *
     (-202291200000 +
      skoXC1 *
      (8880583680000000000 +
       skoXC1 *
       (-259905082368000000000000000 +
        skoXC1 *
        (4991801988230400000000000000000000 +
         skoXC1 *
         (-62611459223804160000000000000000000000000 +
          skoXC1 * 458107176654167104000000000000000000000000000000))))) <=
     0), inference(resolve, [], [refute_0_49, refute_0_42])).

cnf(refute_0_51, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <=
     -602496 +
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000))))) |
     skoXC1 *
     (-202291200000 +
      skoXC1 *
      (8880583680000000000 +
       skoXC1 *
       (-259905082368000000000000000 +
        skoXC1 *
        (4991801988230400000000000000000000 +
         skoXC1 *
         (-62611459223804160000000000000000000000000 +
          skoXC1 * 458107176654167104000000000000000000000000000000))))) <=
     -2304), inference(arithmetic, [], [refute_0_50])).

cnf(refute_0_52, plain, (skoXC1 <= 593/10000000000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_53, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_54, plain,
    (-2304 <
     skoXC1 *
     (-202291200000 +
      skoXC1 *
      (8880583680000000000 +
       skoXC1 *
       (-259905082368000000000000000 +
        skoXC1 *
        (4991801988230400000000000000000000 +
         skoXC1 *
         (-62611459223804160000000000000000000000000 +
          skoXC1 * 458107176654167104000000000000000000000000000000))))) |
     0 < skoXC1), inference(decision, [], [refute_0_52, refute_0_53])).

cnf(refute_0_55, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <=
     -602496 +
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000)))))),
    inference(resolve, [], [refute_0_51, refute_0_54])).

cnf(refute_0_56, plain,
    (0 < skoXC1 |
     4224 +
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <= 0 |
     exp(skoXC1 * 11000000) <=
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3))))))),
    inference(split, [], [refute_0_55])).

cnf(refute_0_57, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <=
     -4224 |
     exp(skoXC1 * 11000000) <=
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3))))))),
    inference(arithmetic, [], [refute_0_56])).

cnf(refute_0_58, plain,
    (-4224 <
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) |
     0 < skoXC1), inference(decision, [], [refute_0_52, refute_0_53])).

cnf(refute_0_59, plain,
    (0 < skoXC1 |
     exp(skoXC1 * 11000000) <=
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3))))))),
    inference(resolve, [], [refute_0_57, refute_0_58])).

cnf(refute_0_60, plain,
    (0 <
     (-602496 +
      skoXC1 *
      (53101440000000 +
       skoXC1 *
       (-2331153216000000000000 +
        skoXC1 *
        (68225084121600000000000000000 +
         skoXC1 *
         (-1310348021910480000000000000000000000 +
          skoXC1 *
          (16435508046248592000000000000000000000000000 +
           skoXC1 *
           -120253133871718864800000000000000000000000000000000)))))) /
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) |
     0 < skoXC1), inference(resolve, [], [refute_0_59, refute_0_41])).

cnf(refute_0_61, plain,
    (0 *
     (4224 +
      skoXC1 *
      (-370867200000 +
       skoXC1 *
       (16281070080000000000 +
        skoXC1 *
        (-476492651008000000000000000 +
         skoXC1 *
         (9151636978422400000000000000000000 +
          skoXC1 *
          (-114787675243640960000000000000000000000000 +
           skoXC1 *
           2519589471597919072000000000000000000000000000000/3)))))) <
     -602496 +
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000))))) |
     0 < skoXC1 |
     4224 +
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <= 0),
    inference(resolve, [], [refute_0_40, refute_0_60])).

cnf(refute_0_62, plain,
    (602496 <
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000))))) |
     0 < skoXC1 |
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <=
     -4224), inference(arithmetic, [], [refute_0_61])).

cnf(refute_0_63, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <=
     -4224 |
     skoXC1 *
     (53101440000000 +
      skoXC1 *
      (-2331153216000000000000 +
       skoXC1 *
       (68225084121600000000000000000 +
        skoXC1 *
        (-1310348021910480000000000000000000000 +
         skoXC1 *
         (16435508046248592000000000000000000000000000 +
          skoXC1 *
          -120253133871718864800000000000000000000000000000000))))) <=
     602496), inference(decision, [], [refute_0_52, refute_0_53])).

cnf(refute_0_64, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-370867200000 +
      skoXC1 *
      (16281070080000000000 +
       skoXC1 *
       (-476492651008000000000000000 +
        skoXC1 *
        (9151636978422400000000000000000000 +
         skoXC1 *
         (-114787675243640960000000000000000000000000 +
          skoXC1 *
          2519589471597919072000000000000000000000000000000/3))))) <=
     -4224), inference(resolve, [], [refute_0_63, refute_0_62])).

cnf(refute_0_65, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_64, refute_0_58])).

cnf(refute_0_66, plain,
    (-21743271936 <
     skoXC1 *
     (-1909059275980800000 +
      skoXC1 *
      (83807702215557120000000000 +
       skoXC1 *
       (-2452772084841971712000000000000000 +
        skoXC1 *
        (52997123086308134092800000000000000000000 +
         skoXC1 *
         (-886313791805305556828160000000000000000000000000 +
          skoXC1 *
          (11753813420284734420811776000000000000000000000000000000 +
           skoXC1 *
           (-124549891863913754741912371200000000000000000000000000000000000 +
            skoXC1 *
            (1049610852104955334046017781760000000000000000000000000000000000000000 +
             skoXC1 *
             (-6905734241937047213302249291776000000000000000000000000000000000000000000000 +
              skoXC1 *
              (33975021826495455557168910955315200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-114731035244857730689208860841410560000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                209862185302052265719011207955746816000000000000000000000000000000000000000000000000000000000000))))))))))) |
     3/87800000 < skoXC1),
    inference(decision, [], [refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_67, plain,
    (3/87800000 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) <=
     -5685865611264 +
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))),
    inference(resolve, [], [refute_0_39, refute_0_66])).

cnf(refute_0_68, plain,
    (3/87800000 < skoXC1 |
     39862665216 +
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= 0 |
     exp(skoXC1 * 11000000) <=
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3))))))))))))), inference(split, [], [refute_0_67])).

cnf(refute_0_69, plain,
    (3/87800000 < skoXC1 |
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= -39862665216 |
     exp(skoXC1 * 11000000) <=
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3))))))))))))), inference(arithmetic, [], [refute_0_68])).

cnf(refute_0_70, plain,
    (-39862665216 <
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) | 3/87800000 < skoXC1),
    inference(decision, [], [refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_71, plain,
    (3/87800000 < skoXC1 |
     exp(skoXC1 * 11000000) <=
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3))))))))))))),
    inference(resolve, [], [refute_0_69, refute_0_70])).

cnf(refute_0_72, plain,
    (0 <
     (-5685865611264 +
      skoXC1 *
      (501128059944960000000 +
       skoXC1 *
       (-21999521831583744000000000000 +
        skoXC1 *
        (643852672271017574400000000000000000 +
         skoXC1 *
         (-13911744810155885199360000000000000000000000 +
          skoXC1 *
          (232657370348892708667392000000000000000000000000000 +
           skoXC1 *
           (-3085376022824742785463091200000000000000000000000000000000 +
            skoXC1 *
            (32694346614277360619751997440000000000000000000000000000000000000 +
             skoXC1 *
             (-275522848677550775187079667712000000000000000000000000000000000000000000 +
              skoXC1 *
              (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000)))))))))))) /
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) | 3/87800000 < skoXC1),
    inference(resolve, [], [refute_0_71, refute_0_29])).

cnf(refute_0_73, plain,
    (0 *
     (39862665216 +
      skoXC1 *
      (-3499942005964800000 +
       skoXC1 *
       (153647454061854720000000000 +
        skoXC1 *
        (-4496748822210281472000000000000000 +
         skoXC1 *
         (97161392324898245836800000000000000000000 +
          skoXC1 *
          (-1624908618309726854184960000000000000000000000000 +
           skoXC1 *
           (21548657937188679771488256000000000000000000000000000000 +
            skoXC1 *
            (-228341468417175217026839347200000000000000000000000000000000000 +
             skoXC1 *
             (1924286562192418112417699266560000000000000000000000000000000000000000 +
              skoXC1 *
              (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
               skoXC1 *
               (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
                skoXC1 *
                (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                 skoXC1 *
                 1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                 3)))))))))))) <
     -5685865611264 +
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000))))))))))) |
     3/87800000 < skoXC1 |
     39862665216 +
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= 0),
    inference(resolve, [], [refute_0_28, refute_0_72])).

cnf(refute_0_74, plain,
    (5685865611264 <
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000))))))))))) |
     3/87800000 < skoXC1 |
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= -39862665216),
    inference(arithmetic, [], [refute_0_73])).

cnf(refute_0_75, plain,
    (3/87800000 < skoXC1 |
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= -39862665216 |
     skoXC1 *
     (501128059944960000000 +
      skoXC1 *
      (-21999521831583744000000000000 +
       skoXC1 *
       (643852672271017574400000000000000000 +
        skoXC1 *
        (-13911744810155885199360000000000000000000000 +
         skoXC1 *
         (232657370348892708667392000000000000000000000000000 +
          skoXC1 *
          (-3085376022824742785463091200000000000000000000000000000000 +
           skoXC1 *
           (32694346614277360619751997440000000000000000000000000000000000000 +
            skoXC1 *
            (-275522848677550775187079667712000000000000000000000000000000000000000000 +
             skoXC1 *
             (1812755238508474893491840439091200000000000000000000000000000000000000000000000 +
              skoXC1 *
              (-8918443229455057083756839125770240000000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (30116896751775154305917325970870272000000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                -55088823641788719751240442088383539200000000000000000000000000000000000000000000000000000000000000))))))))))) <=
     5685865611264),
    inference(decision, [], [refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_76, plain,
    (3/87800000 < skoXC1 |
     skoXC1 *
     (-3499942005964800000 +
      skoXC1 *
      (153647454061854720000000000 +
       skoXC1 *
       (-4496748822210281472000000000000000 +
        skoXC1 *
        (97161392324898245836800000000000000000000 +
         skoXC1 *
         (-1624908618309726854184960000000000000000000000000 +
          skoXC1 *
          (21548657937188679771488256000000000000000000000000000000 +
           skoXC1 *
           (-228341468417175217026839347200000000000000000000000000000000000 +
            skoXC1 *
            (1924286562192418112417699266560000000000000000000000000000000000000000 +
             skoXC1 *
             (-12660512776884586557720790368256000000000000000000000000000000000000000000000 +
              skoXC1 *
              (62287540015241668521476336751411200000000000000000000000000000000000000000000000000 +
               skoXC1 *
               (-210340231282239172930216244875919360000000000000000000000000000000000000000000000000000000 +
                skoXC1 *
                1154242019161287461454561643756607488000000000000000000000000000000000000000000000000000000000000/
                3))))))))))) <= -39862665216),
    inference(resolve, [], [refute_0_75, refute_0_74])).

cnf(refute_0_77, plain, (3/87800000 < skoXC1),
    inference(resolve, [], [refute_0_76, refute_0_70])).

cnf(refute_0_78, plain, (1/21950000 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_77, refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_79, plain,
    (1/21950000 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-120 +
      skoXC1 *
      (-5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * -676836152000000000000000))) /
     (-120 +
      skoXC1 *
      (5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * 676836152000000000000000)))),
    inference(resolve, [], [refute_0_78, refute_0_27])).

cnf(refute_0_80, plain,
    (1/21950000 < skoXC1 |
     0 <=
     -120 +
     skoXC1 *
     (5268000000 +
      skoXC1 * (-92506080000000000 + skoXC1 * 676836152000000000000000)) |
     -120 +
     skoXC1 *
     (-5268000000 +
      skoXC1 *
      (-92506080000000000 + skoXC1 * -676836152000000000000000)) <=
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (-120 +
      skoXC1 *
      (5268000000 +
       skoXC1 *
       (-92506080000000000 + skoXC1 * 676836152000000000000000)))),
    inference(resolve, [], [refute_0_79, refute_0_20])).

cnf(refute_0_81, plain,
    (1/21950000 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <=
     -31380 +
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000)) |
     skoXC1 *
     (-5268000000 +
      skoXC1 * (92506080000000000 + skoXC1 * -676836152000000000000000)) <=
     -120), inference(arithmetic, [], [refute_0_80])).

cnf(refute_0_82, plain,
    (-120 <
     skoXC1 *
     (-5268000000 +
      skoXC1 * (92506080000000000 + skoXC1 * -676836152000000000000000)) |
     1/21950000 < skoXC1),
    inference(decision, [],
              [refute_0_77, refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_83, plain,
    (1/21950000 < skoXC1 |
     exp(skoXC1 * 11000000) *
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <=
     -31380 +
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000))),
    inference(resolve, [], [refute_0_81, refute_0_82])).

cnf(refute_0_84, plain,
    (1/21950000 < skoXC1 |
     220 +
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <= 0 |
     exp(skoXC1 * 11000000) <=
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3)))),
    inference(split, [], [refute_0_83])).

cnf(refute_0_85, plain,
    (1/21950000 < skoXC1 |
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <=
     -220 |
     exp(skoXC1 * 11000000) <=
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3)))),
    inference(arithmetic, [], [refute_0_84])).

cnf(refute_0_86, plain,
    (-220 <
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) |
     1/21950000 < skoXC1),
    inference(decision, [],
              [refute_0_77, refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_87, plain,
    (1/21950000 < skoXC1 |
     exp(skoXC1 * 11000000) <=
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3)))),
    inference(resolve, [], [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (0 <
     (-31380 +
      skoXC1 *
      (1388118000000 +
       skoXC1 *
       (-24190339920000000000 + skoXC1 * 178346326052000000000000000))) /
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) |
     1/21950000 < skoXC1),
    inference(resolve, [], [refute_0_87, refute_0_19])).

cnf(refute_0_89, plain,
    (0 *
     (220 +
      skoXC1 *
      (-9658000000 +
       skoXC1 *
       (169594480000000000 + skoXC1 * -3722598836000000000000000/3))) <
     -31380 +
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000)) |
     1/21950000 < skoXC1 |
     220 +
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <= 0),
    inference(resolve, [], [refute_0_18, refute_0_88])).

cnf(refute_0_90, plain,
    (31380 <
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000)) |
     1/21950000 < skoXC1 |
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <=
     -220), inference(arithmetic, [], [refute_0_89])).

cnf(refute_0_91, plain,
    (1/21950000 < skoXC1 |
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <=
     -220 |
     skoXC1 *
     (1388118000000 +
      skoXC1 *
      (-24190339920000000000 + skoXC1 * 178346326052000000000000000)) <=
     31380),
    inference(decision, [],
              [refute_0_77, refute_0_65, refute_0_52, refute_0_53])).

cnf(refute_0_92, plain,
    (1/21950000 < skoXC1 |
     skoXC1 *
     (-9658000000 +
      skoXC1 *
      (169594480000000000 + skoXC1 * -3722598836000000000000000/3)) <=
     -220), inference(resolve, [], [refute_0_91, refute_0_90])).

cnf(refute_0_93, plain, (1/21950000 < skoXC1),
    inference(resolve, [], [refute_0_92, refute_0_86])).

cnf(refute_0_94, plain, (7/87800000 < skoXC1 | 0 <= skoXC1),
    inference(decision, [],
              [refute_0_93, refute_0_77, refute_0_65, refute_0_52,
               refute_0_53])).

cnf(refute_0_95, plain,
    (7/87800000 < skoXC1 |
     525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-30240 +
      skoXC1 *
      (-1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (-284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * -5217621601983680000000000000000000000000))))) /
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000)))))),
    inference(resolve, [], [refute_0_94, refute_0_17])).

cnf(refute_0_96, plain, (skoXC1 <= 7/87800000),
    inference(decision, [],
              [refute_0_93, refute_0_77, refute_0_65, refute_0_52,
               refute_0_53])).

cnf(refute_0_97, plain,
    (525/2 + exp(skoXC1 * 11000000) * 11/6 <=
     (-30240 +
      skoXC1 *
      (-1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (-284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * -5217621601983680000000000000000000000000))))) /
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000)))))),
    inference(resolve, [], [refute_0_96, refute_0_95])).

cnf(refute_0_98, plain,
    (0 <=
     -30240 +
     skoXC1 *
     (1327536000000 +
      skoXC1 *
      (-25901702400000000000 +
       skoXC1 *
       (284271183840000000000000000 +
        skoXC1 *
        (-1782786424368000000000000000000000 +
         skoXC1 * 5217621601983680000000000000000000000000)))) |
     -30240 +
     skoXC1 *
     (-1327536000000 +
      skoXC1 *
      (-25901702400000000000 +
       skoXC1 *
       (-284271183840000000000000000 +
        skoXC1 *
        (-1782786424368000000000000000000000 +
         skoXC1 * -5217621601983680000000000000000000000000)))) <=
     (525/2 + exp(skoXC1 * 11000000) * 11/6) *
     (-30240 +
      skoXC1 *
      (1327536000000 +
       skoXC1 *
       (-25901702400000000000 +
        skoXC1 *
        (284271183840000000000000000 +
         skoXC1 *
         (-1782786424368000000000000000000000 +
          skoXC1 * 5217621601983680000000000000000000000000)))))),
    inference(resolve, [], [refute_0_97, refute_0_8])).

cnf(refute_0_99, plain,
    (exp(skoXC1 * 11000000) *
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) <=
     -7907760 +
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) |
     skoXC1 *
     (-1327536000000 +
      skoXC1 *
      (25901702400000000000 +
       skoXC1 *
       (-284271183840000000000000000 +
        skoXC1 *
        (1782786424368000000000000000000000 +
         skoXC1 * -5217621601983680000000000000000000000000)))) <= -30240),
    inference(arithmetic, [], [refute_0_98])).

cnf(refute_0_100, plain,
    (-30240 <
     skoXC1 *
     (-1327536000000 +
      skoXC1 *
      (25901702400000000000 +
       skoXC1 *
       (-284271183840000000000000000 +
        skoXC1 *
        (1782786424368000000000000000000000 +
         skoXC1 * -5217621601983680000000000000000000000000))))),
    inference(decision, [],
              [refute_0_93, refute_0_77, refute_0_65, refute_0_52,
               refute_0_53])).

cnf(refute_0_101, plain,
    (exp(skoXC1 * 11000000) *
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) <=
     -7907760 +
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000))))),
    inference(resolve, [], [refute_0_99, refute_0_100])).

cnf(refute_0_102, plain,
    (0 <
     (-7907760 +
      skoXC1 *
      (349805736000000 +
       skoXC1 *
       (-6773295177600000000000 +
        skoXC1 *
        (74905456941840000000000000000 +
         skoXC1 *
         (-466198649972232000000000000000000000 +
          skoXC1 * 1374843292122699680000000000000000000000000))))) /
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) |
     55440 +
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <= 0),
    inference(resolve, [], [refute_0_101, refute_0_7])).

cnf(refute_0_103, plain,
    (0 <
     (-7907760 +
      skoXC1 *
      (349805736000000 +
       skoXC1 *
       (-6773295177600000000000 +
        skoXC1 *
        (74905456941840000000000000000 +
         skoXC1 *
         (-466198649972232000000000000000000000 +
          skoXC1 * 1374843292122699680000000000000000000000000))))) /
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) |
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <=
     -55440), inference(arithmetic, [], [refute_0_102])).

cnf(refute_0_104, plain,
    (-55440 <
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3))))),
    inference(decision, [],
              [refute_0_93, refute_0_77, refute_0_65, refute_0_52,
               refute_0_53])).

cnf(refute_0_105, plain,
    (0 <
     (-7907760 +
      skoXC1 *
      (349805736000000 +
       skoXC1 *
       (-6773295177600000000000 +
        skoXC1 *
        (74905456941840000000000000000 +
         skoXC1 *
         (-466198649972232000000000000000000000 +
          skoXC1 * 1374843292122699680000000000000000000000000))))) /
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3)))))),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain,
    (0 *
     (55440 +
      skoXC1 *
      (-2433816000000 +
       skoXC1 *
       (47486454400000000000 +
        skoXC1 *
        (-521163837040000000000000000 +
         skoXC1 *
         (3268441778008000000000000000000000 +
          skoXC1 * -28696918810910240000000000000000000000000/3))))) <
     -7907760 +
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) |
     55440 +
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_105])).

cnf(refute_0_107, plain,
    (7907760 <
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) |
     skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <=
     -55440), inference(arithmetic, [], [refute_0_106])).

cnf(refute_0_108, plain,
    (skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <=
     -55440 |
     skoXC1 *
     (349805736000000 +
      skoXC1 *
      (-6773295177600000000000 +
       skoXC1 *
       (74905456941840000000000000000 +
        skoXC1 *
        (-466198649972232000000000000000000000 +
         skoXC1 * 1374843292122699680000000000000000000000000)))) <=
     7907760),
    inference(decision, [],
              [refute_0_93, refute_0_77, refute_0_65, refute_0_52,
               refute_0_53])).

cnf(refute_0_109, plain,
    (skoXC1 *
     (-2433816000000 +
      skoXC1 *
      (47486454400000000000 +
       skoXC1 *
       (-521163837040000000000000000 +
        skoXC1 *
        (3268441778008000000000000000000000 +
         skoXC1 * -28696918810910240000000000000000000000000/3)))) <=
     -55440), inference(resolve, [], [refute_0_108, refute_0_107])).

cnf(refute_0_110, plain, ($false),
    inference(resolve, [], [refute_0_109, refute_0_104])).
% SZS output end CNFRefutation for Tunnel-2-VC-U.tptp
