% SZS output start CNFRefutation for Tunnel-2-IL-U.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_upper_bound_case_5, axiom,
    (3 < X |
     ~ lgen(R, 21743271936 / (-(X ^ 3) + 12 * X ^ 2 - 96 * X + 384) ^ 4,
            Y) | lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        -1402412057/50000000000000 * exp(8785654241/100 * X) +
        3926544433/10000000000000 * exp(1104345759/100 * X) +
        288341029/4000000000 < 2/25))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 593/10000000000) =>
        -1402412057/50000000000000 * exp(8785654241/100 * X) +
        3926544433/10000000000000 * exp(1104345759/100 * X) +
        288341029/4000000000 < 2/25)), inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 593/10000000000) =>
          -1402412057/50000000000000 * exp(8785654241/100 * X) +
          3926544433/10000000000000 * exp(1104345759/100 * X) +
          288341029/4000000000 < 2/25)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 593/10000000000 & 0 <= X &
        2/25 <=
        -1402412057/50000000000000 * exp(8785654241/100 * X) +
        3926544433/10000000000000 * exp(1104345759/100 * X) +
        288341029/4000000000)), inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     2/25 <=
     -1402412057/50000000000000 * exp(8785654241/100 * skoXC1) +
     3926544433/10000000000000 * exp(1104345759/100 * skoXC1) +
     288341029/4000000000 & skoXC1 <= 593/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (2/25 <=
     -1402412057/50000000000000 * exp(8785654241/100 * skoXC1) +
     3926544433/10000000000000 * exp(1104345759/100 * skoXC1) +
     288341029/4000000000), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoXC1 <= 593/10000000000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (79147427500/3926544433 *
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) < 21743271936 |
     21743271936 +
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= 0 |
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) <=
     79147427500/3926544433),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (2/25 <=
     -1402412057/50000000000000 * exp(8785654241/100 * skoXC1) +
     3926544433/10000000000000 * exp(1104345759/100 * skoXC1) +
     288341029/4000000000), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain,
    (exp(skoXC1 * 8785654241/100) <=
     -395737137500/1402412057 +
     exp(skoXC1 * 1104345759/100) * 19632722165/1402412057),
    inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_4, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_5, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (0 <
     -395737137500/1402412057 +
     exp(skoXC1 * 1104345759/100) * 19632722165/1402412057 |
     -395737137500/1402412057 +
     exp(skoXC1 * 1104345759/100) * 19632722165/1402412057 <
     exp(skoXC1 * 8785654241/100)), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (0 <
     -395737137500/1402412057 +
     exp(skoXC1 * 1104345759/100) * 19632722165/1402412057),
    inference(resolve, [], [refute_0_2, refute_0_6])).

cnf(refute_0_8, plain,
    (79147427500/3926544433 < exp(skoXC1 * 1104345759/100)),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, exp(X_000061), X_000062) | exp(X_000061) <= X_000062),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | lgen(0, exp(X_000061), X_000062)),
    inference(subst, [], [exp_upper_bound_case_5])).

cnf(refute_0_11, plain,
    (3 < X_000061 |
     ~ lgen(0,
            21743271936 /
            (-(X_000061 ^ 3) + 12 * X_000061 ^ 2 - 96 * X_000061 + 384) ^
            4, X_000062) | exp(X_000061) <= X_000062),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (3 < X_000061 |
     X_000062 <
     21743271936 /
     (21743271936 +
      X_000061 *
      (-21743271936 +
       X_000061 *
       (10871635968 +
        X_000061 *
        (-3623878656 +
         X_000061 *
         (891813888 +
          X_000061 *
          (-169869312 +
           X_000061 *
           (25657344 +
            X_000061 *
            (-3096576 +
             X_000061 *
             (297216 +
              X_000061 *
              (-22272 +
               X_000061 * (1248 + X_000061 * (-48 + X_000061)))))))))))) |
     exp(X_000061) <= X_000062), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (3 < skoXC1 * 1104345759/100 |
     79147427500/3926544433 <
     21743271936 /
     (21743271936 +
      skoXC1 * 1104345759/100 *
      (-21743271936 +
       skoXC1 * 1104345759/100 *
       (10871635968 +
        skoXC1 * 1104345759/100 *
        (-3623878656 +
         skoXC1 * 1104345759/100 *
         (891813888 +
          skoXC1 * 1104345759/100 *
          (-169869312 +
           skoXC1 * 1104345759/100 *
           (25657344 +
            skoXC1 * 1104345759/100 *
            (-3096576 +
             skoXC1 * 1104345759/100 *
             (297216 +
              skoXC1 * 1104345759/100 *
              (-22272 +
               skoXC1 * 1104345759/100 *
               (1248 +
                skoXC1 * 1104345759/100 *
                (-48 + skoXC1 * 1104345759/100)))))))))))) |
     exp(skoXC1 * 1104345759/100) <= 79147427500/3926544433),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (3 < skoXC1 * 1104345759/100 |
     79147427500/3926544433 <
     21743271936 /
     (21743271936 +
      skoXC1 * 1104345759/100 *
      (-21743271936 +
       skoXC1 * 1104345759/100 *
       (10871635968 +
        skoXC1 * 1104345759/100 *
        (-3623878656 +
         skoXC1 * 1104345759/100 *
         (891813888 +
          skoXC1 * 1104345759/100 *
          (-169869312 +
           skoXC1 * 1104345759/100 *
           (25657344 +
            skoXC1 * 1104345759/100 *
            (-3096576 +
             skoXC1 * 1104345759/100 *
             (297216 +
              skoXC1 * 1104345759/100 *
              (-22272 +
               skoXC1 * 1104345759/100 *
               (1248 +
                skoXC1 * 1104345759/100 *
                (-48 + skoXC1 * 1104345759/100))))))))))))),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    (79147427500/3926544433 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) |
     100/368115253 < skoXC1), inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (79147427500/3926544433 *
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) <= 79147427500/3926544433),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_17, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_18, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_19, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (0 < skoXC1 * 1104345759/100 |
     79147427500/3926544433 <
     2304 /
     (2304 +
      skoXC1 * 1104345759/100 *
      (-2304 +
       skoXC1 * 1104345759/100 *
       (1152 +
        skoXC1 * 1104345759/100 *
        (-384 +
         skoXC1 * 1104345759/100 *
         (84 +
          skoXC1 * 1104345759/100 * (-12 + skoXC1 * 1104345759/100)))))) |
     exp(skoXC1 * 1104345759/100) <= 79147427500/3926544433),
    inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (0 < skoXC1 * 1104345759/100 |
     79147427500/3926544433 <
     2304 /
     (2304 +
      skoXC1 * 1104345759/100 *
      (-2304 +
       skoXC1 * 1104345759/100 *
       (1152 +
        skoXC1 * 1104345759/100 *
        (-384 +
         skoXC1 * 1104345759/100 *
         (84 +
          skoXC1 * 1104345759/100 * (-12 + skoXC1 * 1104345759/100))))))),
    inference(resolve, [], [refute_0_21, refute_0_8])).

cnf(refute_0_23, plain,
    (79147427500/3926544433 <
     2304 /
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) | 0 < skoXC1),
    inference(arithmetic, [], [refute_0_22])).

cnf(refute_0_24, plain,
    (79147427500/3926544433 *
     (2304 +
      skoXC1 *
      (-636103157184/25 +
       skoXC1 *
       (87809727990332597832/625 +
        skoXC1 *
        (-8081025058755616451268482874/15625 +
         skoXC1 *
         (31234860132033217958861427312931109781/25000000 +
          skoXC1 *
          (-4927726474252723471427321960246033849401538397/2500000000 +
           skoXC1 *
           1813971277784339286623506894508424692719011205600869441/
           1000000000000)))))) < 2304 | 0 < skoXC1 |
     2304 +
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_23])).

cnf(refute_0_25, plain,
    (173308914586368/3926544433 <
     skoXC1 *
     (2013837140629669766400/3926544433 +
      skoXC1 *
      (-11119862527855311980471803488/3926544433 +
       skoXC1 *
       (1023347751941669429271327250087850616/98163610825 +
        skoXC1 *
        (-141266218729870831199467588616956275650642193/5609349190000 +
         skoXC1 *
         (156006749544299219052936914547092338503221671466009487/
          3926544433000000 +
          skoXC1 *
          -57428464078207341729374292691542276602475086906792963207405211/
          1570617773200000000))))) | 0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_24])).

cnf(refute_0_26, plain, (skoXC1 <= 593/10000000000),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_27, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_28, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304 |
     skoXC1 *
     (2013837140629669766400/3926544433 +
      skoXC1 *
      (-11119862527855311980471803488/3926544433 +
       skoXC1 *
       (1023347751941669429271327250087850616/98163610825 +
        skoXC1 *
        (-141266218729870831199467588616956275650642193/5609349190000 +
         skoXC1 *
         (156006749544299219052936914547092338503221671466009487/
          3926544433000000 +
          skoXC1 *
          -57428464078207341729374292691542276602475086906792963207405211/
          1570617773200000000))))) <= 173308914586368/3926544433),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_29, plain,
    (0 < skoXC1 |
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) <= -2304),
    inference(resolve, [], [refute_0_28, refute_0_25])).

cnf(refute_0_30, plain,
    (-2304 <
     skoXC1 *
     (-636103157184/25 +
      skoXC1 *
      (87809727990332597832/625 +
       skoXC1 *
       (-8081025058755616451268482874/15625 +
        skoXC1 *
        (31234860132033217958861427312931109781/25000000 +
         skoXC1 *
         (-4927726474252723471427321960246033849401538397/2500000000 +
          skoXC1 *
          1813971277784339286623506894508424692719011205600869441/
          1000000000000))))) | 0 < skoXC1),
    inference(decision, [], [refute_0_26, refute_0_27])).

cnf(refute_0_31, plain, (0 < skoXC1),
    inference(resolve, [], [refute_0_29, refute_0_30])).

cnf(refute_0_32, plain, (skoXC1 <= 100/368115253),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_33, plain,
    (79147427500/3926544433 <
     21743271936 /
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000))))))))))))),
    inference(resolve, [], [refute_0_32, refute_0_15])).

cnf(refute_0_34, plain,
    (79147427500/3926544433 *
     (21743271936 +
      skoXC1 *
      (-6003022537326329856/25 +
       skoXC1 *
       (828676560034718946938585088/625 +
        skoXC1 *
        (-76262120388087563484047706282786816/15625 +
         skoXC1 *
         (5181488477582726460759604454087515939350528/390625 +
          skoXC1 *
          (-272483563120278597076045195113764687736507467200512/9765625 +
           skoXC1 *
           (11362716084041101291409647187200772275191886191883846178424/
            244140625 +
            skoXC1 *
            (-378614531150530795474153745776293808644097340476392673423518585891/
             6103515625 +
             skoXC1 *
             (2568459732312867029567636797180944557643550978363740200871313533002190387081/
              39062500000000 +
              skoXC1 *
              (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
               3906250000000000 +
               skoXC1 *
               (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
                3125000000000000000 +
                skoXC1 *
                (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                 625000000000000000000 +
                 skoXC1 *
                 3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                 1000000000000000000000000)))))))))))) < 21743271936 |
     21743271936 +
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_33])).

cnf(refute_0_35, plain,
    (1635548115791838707712/3926544433 <
     skoXC1 *
     (19004951642156069444753817600/3926544433 +
      skoXC1 *
      (-104940188730075704537116816328097792/3926544433 +
       skoXC1 *
       (9657521031059891671208501183293062427705344/98163610825 +
        skoXC1 *
        (-93737481970643249784069117361958570620534928449536/350584324375 +
         skoXC1 *
         (34506196891206278467004798507184071799496582181537842372608/
          61352256765625 +
          skoXC1 *
          (-1438927595943563154371202277679283282551615777536604486124744966816/
           1533806419140625 +
           skoXC1 *
           (6849455121070429193441808392613176100196015441896709618587226073533816092/
            5477880068359375 +
            skoXC1 *
            (-81314792179960820215937955900484405503045008761039278470919089873223888401076153651/
             61352256765625000000 +
             skoXC1 *
             (6729172430876684240716650719299845991278543910412856787446009173838255337424585266911918003/
              6135225676562500000000 +
              skoXC1 *
              (-3331287222863414814200978912174486096816584221278492631716370370477373222067979702897406060386013469/
               4908180541250000000000000 +
               skoXC1 *
               (282991762813853845101586464254486706847989091202709299672978654577688147942379891891756184222168605954386767/
                981636108250000000000000000 +
                skoXC1 *
                -104173584365137800094593311990415830476484336849325408134571354693465280734910603325846476369258204589556391194057051/
                1570617773200000000000000000000))))))))))) |
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936),
    inference(arithmetic, [], [refute_0_34])).

cnf(refute_0_36, plain,
    (skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936 |
     skoXC1 *
     (19004951642156069444753817600/3926544433 +
      skoXC1 *
      (-104940188730075704537116816328097792/3926544433 +
       skoXC1 *
       (9657521031059891671208501183293062427705344/98163610825 +
        skoXC1 *
        (-93737481970643249784069117361958570620534928449536/350584324375 +
         skoXC1 *
         (34506196891206278467004798507184071799496582181537842372608/
          61352256765625 +
          skoXC1 *
          (-1438927595943563154371202277679283282551615777536604486124744966816/
           1533806419140625 +
           skoXC1 *
           (6849455121070429193441808392613176100196015441896709618587226073533816092/
            5477880068359375 +
            skoXC1 *
            (-81314792179960820215937955900484405503045008761039278470919089873223888401076153651/
             61352256765625000000 +
             skoXC1 *
             (6729172430876684240716650719299845991278543910412856787446009173838255337424585266911918003/
              6135225676562500000000 +
              skoXC1 *
              (-3331287222863414814200978912174486096816584221278492631716370370477373222067979702897406060386013469/
               4908180541250000000000000 +
               skoXC1 *
               (282991762813853845101586464254486706847989091202709299672978654577688147942379891891756184222168605954386767/
                981636108250000000000000000 +
                skoXC1 *
                -104173584365137800094593311990415830476484336849325408134571354693465280734910603325846476369258204589556391194057051/
                1570617773200000000000000000000))))))))))) <=
     1635548115791838707712/3926544433),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_37, plain,
    (skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000))))))))))) <= -21743271936),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (-21743271936 <
     skoXC1 *
     (-6003022537326329856/25 +
      skoXC1 *
      (828676560034718946938585088/625 +
       skoXC1 *
       (-76262120388087563484047706282786816/15625 +
        skoXC1 *
        (5181488477582726460759604454087515939350528/390625 +
         skoXC1 *
         (-272483563120278597076045195113764687736507467200512/9765625 +
          skoXC1 *
          (11362716084041101291409647187200772275191886191883846178424/
           244140625 +
           skoXC1 *
           (-378614531150530795474153745776293808644097340476392673423518585891/
            6103515625 +
            skoXC1 *
            (2568459732312867029567636797180944557643550978363740200871313533002190387081/
             39062500000000 +
             skoXC1 *
             (-212551836598753770004611038030890075084201059801117881798685408121390153123567574793/
              3906250000000000 +
              skoXC1 *
              (105224115555221766816141273580069487944399210614852031410508268587673718835270410491149761639/
               3125000000000000000 +
               skoXC1 *
               (-8938754289071929883684042171000652764361453541958432561594584188402337774729945957237718946145426077/
                625000000000000000000 +
                skoXC1 *
                3290491796626548604330611755840574555518065854045774517894828441943526235736171062724890091003216895127652481/
                1000000000000000000000000)))))))))))),
    inference(decision, [], [refute_0_31, refute_0_26, refute_0_27])).

cnf(refute_0_39, plain, ($false),
    inference(resolve, [], [refute_0_37, refute_0_38])).
% SZS output end CNFRefutation for Tunnel-2-IL-U.tptp
