% SZS output start CNFRefutation for Tunnel-3-VC-L.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(exp_upper_bound_case_3, axiom,
    (0 < X |
     ~ lgen(R, 2304 / (-(X ^ 3) + 6 * X ^ 2 - 24 * X + 48) ^ 2, Y) |
     lgen(R, exp(X), Y))).

cnf(exp_positive, axiom, (0 < Y | lgen(R, Y, exp(X)))).

fof('Tunnel', conjecture,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        0 <
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X)))).

fof(subgoal_0, plain,
    (! [X] :
       ((0 <= X & X <= 817/10000000000) =>
        0 <
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X))),
    inference(strip, [], ['Tunnel'])).

fof(negate_0_0, plain,
    (~ ! [X] :
         ((0 <= X & X <= 817/10000000000) =>
          0 <
          -65737067/400000000 * exp(-2559889987/10 * X) +
          403501489/1250000000 +
          5645414763/10000000000 * exp(-168440051/40 * X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [X] :
       (X <= 817/10000000000 & 0 <= X &
        -65737067/400000000 * exp(-2559889987/10 * X) +
        403501489/1250000000 +
        5645414763/10000000000 * exp(-168440051/40 * X) <= 0)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (0 <= skoXC1 &
     -65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1) <= 0 &
     skoXC1 <= 817/10000000000),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (-65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1) <= 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoXC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoXC1 <= 817/10000000000),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (3228011912/1643426675 *
     (2304 +
      skoXC1 *
      (2948993265024/5 +
       skoXC1 *
       (1887274582716343728672/25 +
        skoXC1 *
        (805202551169195262046282934544/125 +
         skoXC1 *
         (901788102357134166195052856434820879781/2500 +
          skoXC1 *
          (329782619088536978579758528017606641184273236121/25000 +
           skoXC1 *
           281402408163793625981852445583376733490774293006078206809/
           1000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000))))) <= 0 |
     2304 /
     (2304 +
      skoXC1 *
      (2948993265024/5 +
       skoXC1 *
       (1887274582716343728672/25 +
        skoXC1 *
        (805202551169195262046282934544/125 +
         skoXC1 *
         (901788102357134166195052856434820879781/2500 +
          skoXC1 *
          (329782619088536978579758528017606641184273236121/25000 +
           skoXC1 *
           281402408163793625981852445583376733490774293006078206809/
           1000000)))))) <= 3228011912/1643426675),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_1, plain,
    (-65737067/400000000 * exp(-2559889987/10 * skoXC1) +
     403501489/1250000000 +
     5645414763/10000000000 * exp(-168440051/40 * skoXC1) <= 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_2, plain,
    (exp(skoXC1 * -168440051/40) <=
     -3228011912/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763),
    inference(arithmetic, [], [refute_0_1])).

cnf(refute_0_3, plain,
    (X_000010 < exp(X_000009) | ~ lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_4, plain, (0 < X_000010 | lgen(1, X_000010, exp(X_000009))),
    inference(subst, [], [exp_positive])).

cnf(refute_0_5, plain, (0 < X_000010 | X_000010 < exp(X_000009)),
    inference(resolve, [], [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (0 <
     -3228011912/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 |
     -3228011912/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763 <
     exp(skoXC1 * -168440051/40)), inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (0 <
     -3228011912/5645414763 +
     exp(skoXC1 * -2559889987/10) * 1643426675/5645414763),
    inference(resolve, [], [refute_0_2, refute_0_6])).

cnf(refute_0_8, plain,
    (3228011912/1643426675 < exp(skoXC1 * -2559889987/10)),
    inference(arithmetic, [], [refute_0_7])).

cnf(refute_0_9, plain,
    (~ lgen(0, exp(X_000054), X_000055) | exp(X_000054) <= X_000055),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_10, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | lgen(0, exp(X_000054), X_000055)),
    inference(subst, [], [exp_upper_bound_case_3])).

cnf(refute_0_11, plain,
    (0 < X_000054 |
     ~ lgen(0,
            2304 /
            (-(X_000054 ^ 3) + 6 * X_000054 ^ 2 - 24 * X_000054 + 48) ^ 2,
            X_000055) | exp(X_000054) <= X_000055),
    inference(resolve, [], [refute_0_10, refute_0_9])).

cnf(refute_0_12, plain,
    (0 < X_000054 |
     X_000055 <
     2304 /
     (2304 +
      X_000054 *
      (-2304 +
       X_000054 *
       (1152 +
        X_000054 *
        (-384 + X_000054 * (84 + X_000054 * (-12 + X_000054)))))) |
     exp(X_000054) <= X_000055), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (0 < skoXC1 * -2559889987/10 |
     3228011912/1643426675 <
     2304 /
     (2304 +
      skoXC1 * -2559889987/10 *
      (-2304 +
       skoXC1 * -2559889987/10 *
       (1152 +
        skoXC1 * -2559889987/10 *
        (-384 +
         skoXC1 * -2559889987/10 *
         (84 +
          skoXC1 * -2559889987/10 * (-12 + skoXC1 * -2559889987/10)))))) |
     exp(skoXC1 * -2559889987/10) <= 3228011912/1643426675),
    inference(subst, [], [refute_0_12])).

cnf(refute_0_14, plain,
    (0 < skoXC1 * -2559889987/10 |
     3228011912/1643426675 <
     2304 /
     (2304 +
      skoXC1 * -2559889987/10 *
      (-2304 +
       skoXC1 * -2559889987/10 *
       (1152 +
        skoXC1 * -2559889987/10 *
        (-384 +
         skoXC1 * -2559889987/10 *
         (84 +
          skoXC1 * -2559889987/10 * (-12 + skoXC1 * -2559889987/10))))))),
    inference(resolve, [], [refute_0_13, refute_0_8])).

cnf(refute_0_15, plain,
    (skoXC1 < 0 |
     3228011912/1643426675 <
     2304 /
     (2304 +
      skoXC1 *
      (2948993265024/5 +
       skoXC1 *
       (1887274582716343728672/25 +
        skoXC1 *
        (805202551169195262046282934544/125 +
         skoXC1 *
         (901788102357134166195052856434820879781/2500 +
          skoXC1 *
          (329782619088536978579758528017606641184273236121/25000 +
           skoXC1 *
           281402408163793625981852445583376733490774293006078206809/
           1000000))))))), inference(arithmetic, [], [refute_0_14])).

cnf(refute_0_16, plain, (0 <= skoXC1),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_17, plain,
    (3228011912/1643426675 <
     2304 /
     (2304 +
      skoXC1 *
      (2948993265024/5 +
       skoXC1 *
       (1887274582716343728672/25 +
        skoXC1 *
        (805202551169195262046282934544/125 +
         skoXC1 *
         (901788102357134166195052856434820879781/2500 +
          skoXC1 *
          (329782619088536978579758528017606641184273236121/25000 +
           skoXC1 *
           281402408163793625981852445583376733490774293006078206809/
           1000000))))))),
    inference(resolve, [], [refute_0_16, refute_0_15])).

cnf(refute_0_18, plain,
    (3228011912/1643426675 *
     (2304 +
      skoXC1 *
      (2948993265024/5 +
       skoXC1 *
       (1887274582716343728672/25 +
        skoXC1 *
        (805202551169195262046282934544/125 +
         skoXC1 *
         (901788102357134166195052856434820879781/2500 +
          skoXC1 *
          (329782619088536978579758528017606641184273236121/25000 +
           skoXC1 *
           281402408163793625981852445583376733490774293006078206809/
           1000000)))))) < 2304 |
     2304 +
     skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000))))) <= 0),
    inference(resolve, [], [refute_0_0, refute_0_17])).

cnf(refute_0_19, plain,
    (3650884386048/1643426675 <
     skoXC1 *
     (-9519385387905244965888/8217133375 +
      skoXC1 *
      (-6092144834223186873239711940864/41085666875 +
       skoXC1 *
       (-2599203426746951833339362808030348288128/205428334375 +
        skoXC1 *
        (-727745684127176091664954584010306912879846987818/1027141671875 +
         skoXC1 *
         (-133067777848544493688493671315552497934142974157672084169/
          5135708359375 +
          skoXC1 *
          -113546290702276483972386548771183985611483594990874842497881438601/
          205428334375000))))) |
     skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000))))) <= -2304),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain, (skoXC1 <= 817/10000000000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_21, plain,
    (skoXC1 *
     (-9519385387905244965888/8217133375 +
      skoXC1 *
      (-6092144834223186873239711940864/41085666875 +
       skoXC1 *
       (-2599203426746951833339362808030348288128/205428334375 +
        skoXC1 *
        (-727745684127176091664954584010306912879846987818/1027141671875 +
         skoXC1 *
         (-133067777848544493688493671315552497934142974157672084169/
          5135708359375 +
          skoXC1 *
          -113546290702276483972386548771183985611483594990874842497881438601/
          205428334375000))))) <= 3650884386048/1643426675 |
     skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000))))) <= -2304),
    inference(decision, [], [refute_0_20, refute_0_16])).

cnf(refute_0_22, plain,
    (skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000))))) <= -2304),
    inference(resolve, [], [refute_0_21, refute_0_19])).

cnf(refute_0_23, plain,
    (-2304 <
     skoXC1 *
     (2948993265024/5 +
      skoXC1 *
      (1887274582716343728672/25 +
       skoXC1 *
       (805202551169195262046282934544/125 +
        skoXC1 *
        (901788102357134166195052856434820879781/2500 +
         skoXC1 *
         (329782619088536978579758528017606641184273236121/25000 +
          skoXC1 *
          281402408163793625981852445583376733490774293006078206809/
          1000000)))))),
    inference(decision, [], [refute_0_20, refute_0_16])).

cnf(refute_0_24, plain, ($false),
    inference(resolve, [], [refute_0_22, refute_0_23])).
% SZS output end CNFRefutation for Tunnel-3-VC-L.tptp
